# Repetir
### Semana 4 · Unidad 4 · Estructuras repetitivas e integración del control de flujo

Sabes hacer que un programa calcule y sabes hacer que decida. Te falta una sola
cosa, y con ella podrás escribir —literalmente— cualquier algoritmo que exista:
**repetir**.

Secuencia, decisión y repetición. Tres estructuras. No hay una cuarta.

**Empieza ejecutando la celda de abajo.**


> ### Antes de empezar, dos cosas
>
> **1. Ejecuta las celdas en orden, de arriba abajo.** Una por una, con
> `Shift+Enter`. Cada celda usa lo que dejaron las de arriba, así que saltarte
> una hace que las siguientes fallen aunque estén bien escritas.
>
> **2. Al terminar, entrega.** Tu trabajo **no le llega a tu profesor** hasta que
> pulses **Guardar y entregar** — el botón está arriba y también al final del
> cuadernillo. Puedes entregar las veces que quieras: siempre cuenta la última.


### Qué mide este cuadernillo

Cada ejercicio calificable está asociado a una microcompetencia del programa. Estas son las de esta sesión:

| Microcompetencia | Qué significa | Ejercicios |
|---|---|---:|
| **mCP17** | Aplica conocimientos matemáticos para la solución de problemas usando programación. | 1 |
| **mCC87** | Identifica variables, conceptos y aspectos relevantes de un problema para desarrollar algoritmos que permitan solucionarlo. | 6 |

No hace falta que hagas nada con esto: está aquí para que sepas qué se está midiendo y por qué estos ejercicios y no otros.


```python
# Motor del cuadernillo: barra de progreso, pistas y verificadores.
# No necesitas leer este codigo. Ejecutalo una vez (Shift+Enter) y sigue.
# Codigo fuente legible en el repositorio del curso:
#   notebook/cuadernillos/**/ava_motor.py
#   notebook/cuadernillos/**/pseudo_uis.py
#   notebook/cuadernillos/**/contenido.py
import base64, zlib
exec(zlib.decompress(base64.b64decode("eNrlfd1y20iW5j2fIhsOr0kXBcvyT7loSxuyre5yjWzXWq7y9GoUVBJIUpBBgA2AtGQuI2beYV6g9q6io67qYiK6LybCepN5kvnOyUwgAZKS5aqdiIlVd1kkfk5mnt/vnPyR53kv0yLNRHzxtzAKUhEqEae5CKYyVFkSxfQlVLHY/XHXb7X+1/TiZ6FycS7+Qp9evcaX1sbST2svL5QYM+FJlBRSxFIEciKJ/CS7+CkookAyXaehHohGgcq7eCQdxGosqelWmoUqkRmuRnmBSyoPZJwmMpR4ciCzzFBNR5nKU6FElOTRKIlk7ovnUmSqyFIZR2OFfgTRxa9JK0rGKowkunUu8ihBz/IiGqVoQYLYX6ZKoGMqL6Z4KMFAciVkkamZFFIMZRzLDKy4fXtfiiQt6B+Ry1hRLyTY8ot/+7aw985p5IXCYNCNi18kP5kwlzGUQMWhHmQyGGXECNE+5g/9KOzhvbzoq1OVBVEQpf1Xx52u7l0OGfGnWIHYcTDNi3Tsn+bH4rx1LCdRX51N0qzwJ+fHvjiIhIxHF39Dq9QBVYkmVDlkkqlAdUVdFGDLCOTzSCVh2gK/oyEEBpn4Yo/1IRoodBMis4LV+oMvaZahhwkRfA9q+EVilvF0DD6di4uf49ZkqsCATEGSWTSIslg+FuivIPElpIJBhPskYR5kMIXoFFj+PZ5hvZtFM25jmojxxa/hNCbKJIapiotU35Gauyu086qf1m7Z3zwFaYg2jtVIiuMPafb+jsMmP5qcJ4PjLnVEZsFJNEsFGA076lF3TuQ5eP5rAvlO0gTCxWPRmCUj9mIxUgkzkQQBJmK0rgRu346SIINkJcsGw4SK0pMkwJYZNjQNg4UMJhkUPDNj9sVL0ltQBx1WSTFR2TiC4KH3hcRlbWP0CSoD1Y5y/hyk47RsK0mzsYypgZn6yH2ku8xZSSYpRtFIaskcwIrAiw9ROFJFTh3WGkYspLGpHt4nxSYTIgozleVki0TWMDkoptr6WuMpxpjxuFUyTYIITIKAodQxq9jFr+gHjB1MBbt8sZtjhA0FpveYGlQAIxjEMoF3Iy2DEKx/YXYm6GiS+i3P81qtGwIKIL6HkRLfBfQwKaJQhlC5fApmXFOTboiDdJBFNGq0Obn4NY8K6AbxRxYKN2JLOE8TdrxwOeDTmHwbmEdX1Dg9jXK/tfu/f9jvCvq3//rgmdgW3o0t+fWj8KHXxce7m1vf3A+91o97b57vdQX/Kp/bHD6S9yU/tykfBPfueq03r797DWovn+6+4UeCR/e+2RrwI4MHX29tbnqtP715cdAV9G//2f7uGzz+9DWo8uMPth7cva/48eHD4dfDb/hjOFT31CN04sXr/b23u/rJwf1BcB+s7T87OKArYPOTvDiP1U7LlzO5EchTOR+Ql896dydnpA5RKG62+VInf6xvbcRqWPTuO/flx2lc3YZuRdO893By9rgl4MjDMEpGvbtbeOEu3no8ltkoSnqP8H3z8UAG70dZOk3C3s12AH+egtAwTYqNoRxH8XkvP4eRjTemUXdDTiax2tAXuiB960CNUiV+eHGr+yYdQMu7uUzyjRzSHGoaefRR9e7e9x+g2ThK1MaJikYnRe+u/+DB4wDKnvVu3A3of48XFQv89P3cGeqGfvBmG5ai2VD1+cZwayiHD2pvw1JXvp6lp2nz7XB4f3iv9jaH1pXvyzG8xDKBb9SwRgBRNl3d/SglW2pS+Hr4YBhaCkVUzJlzHzSjHj7YfFy2DymneTBlEWkhboDriHekDA0GipJcjXvl+w2GLT9uuLXEmuUnLV/Mo4xDNjCIMJ1bcUP3aoNW3yDSDlcobIpeDuP0Q+8kCuFw6jQzheCTVFQ3N2/edMmShslsg5wronXR/mYTnrZrzaPrSKBDlrHcOjwtAFMRpUkP/rs4Ef69B/A6Mle2I8Aucr7aPFbr/j2QtYwaZVFeia5IJ67cCmCKUu/wRiwnuerZD4+v06Zj5A/ZyN0mRHHSdb+Flzkc6zyIzN1N4pA6KzaAgUZJj3S7SXm+0p2sVOCalm+Wfcxno/lYnm0w/42IjbzltEgfh1E+ieV5bxCnwXs7Suqa4NuGyiRS15LSFnuohpygIs5wA0JemSu8jYrP1g5Z0TVdcus9IAzEafTqzm3xzCJcIZMQ3ZKngAI5gZiAgjcQm4UbgEQM6JADfHQRD7R/RvkDOsYxn7CBgZkBfwAWQHemhO8BayRhAUKNoTplSMO5x4yNGLg5A46DgNAA0Sux9WOLGYfAvUAeFI23Hoh/eEo9P/jxT0wR8VgSeJMJY1Nf3L7T8gG44t6JzNs7wk+nxWRa9D9kiBlAeyyUNJjGMD6B21GCu3MrywSI0EpOP9S89eSOCZMImeKmmFOPhcda6vV0LEbQJf3CV40OPNY+fHWjtsdOEBc1NtBkyNPhkoYBHnszokJ4gN7QPoPe0aHcNKQV2TRH8MJQc9ysbYfvCo/UyvSn21oQwHoJpGxUIohUVqR3VJYBl7EkwWsGPvyFkQ+rywzigWQ0Gi3SMKUMKAexPCIdl7nG/wHGgZTNoOSI8hjFmQ9UYMhZBCcjiZxBDQC8/Vb/xbPXr173X/8DkEmbx3LrCaxRsCVue3cfeEIbov48i9SHp+nZtrcpNsXdh/i/J1hI28QCSmpjYzkbW5WdZNreJ2fezi3TRhBlAfKKAKQeeSI451/Ztve1v+VBCeN427uZe3fK5ycSjjnc9l7e9x+KR/69eMvfEvQfvm/c97+xL5HqUJey9D36dGM4HNpvG3ZI9LAha+5QCEESt+2xDytfoMunaZTY63d2oJGz0c4tKCOLuNWx/Hu5u//fhYEP/PsC/8UPwDz89/LuJjhKFzbMld+TlS7LyNTAsVaryM57QtygFFBnSVB2YEwkrnhXIRRLJ5d6zBUKZM1kBiaHYgvhQoTPjeuc0k3AYBDv+Na3u3/uv3vx/E97bwl7v82QXamzQE0K8YJf2iPT495MMjkay542PshCbKC5iSInHXIixmkSE30HUq/AnRUt/FHGuWq1TI9kXtiPcJEncTSwX6PUftIcC06UvQC7bg2zdCxefH9enKSJb9yiHea3b1/C05mLYGeohqJ/UozjNgWutNPjbpn7bXra3CDeg+s68mRcuaDohOEmgnIT7ULw5Tv45IMgi8AlVa92gBHfHZB7Qawg7wNVDW3Ng8IHoITkbPnclidSjks2foA69CSWYziwNj0JeQvKobkIdEPXdGy/4MKo2sallOevX3a6YkwIL4NwzXOIZviX+k4lKp0d55EaTzLVRV5/QxTkEC9+TgQ8J+IlrEnywCcXP8GuEML20RGCd6Es8InS9SiDs+2J49Pp5Byx38/TaRaovoamx6DJfRXf6dv7coCWMMykAGXzTp+HGQ36yQCMV0CXaZLrNF+XeUjVU5Aixcdt9OOYY2jnmApbeWqibOnoK4/Njedof2tz6x7cN4m3v3cAJ/R8700zu7wqNrP1bYwpHl0SoIkKPz5jjEqg5MuJDmNVgl/nScZtBpGJdcmng+JWguse+dSHk1X0hZy7eBSvMMgLFVRXMvTn0RogbnEd+hLio3KybUT4HGQmCAsEDB2Q8iSMZgLwI8+3vUbr3s5LU2+s12hkUESzlHzYEynSJID+vt/2ipMo9wNIHXrSvsXsv9XxmfI+or1fpKNRrNq36iK51XmMHHSaJeQYkTp47AXEDnkyapYfenJHoqN30NMSVNUASg1NLYxvMVA167PthbJtPAwIGF/SKNCVBmysFH56DJ9NgBTpIXkFlcQyULrcPKMqnd9ikvtSTHNZESCMhBwwAaEeY9KAHABX5EV88QsSP/5IyiMTxfVlfDVBgCmuwK2COhaRW4F/+Qm3Mjh5LvTqKhWXnzIu8cOL5BF8C5MyVbmkKwBXJ7JWdsTXguKR4HxhqbLvW4bxb+2s66ZLvplFLFhVHA4XuvhHrgIqalwlMbFGn5//HvyEM1OZHrtTh6OvyN6pUNyjIK6VAz+v0vEgU83+mpkAjGd5foEr+naKQQ+MHGj/bALSsIqS9D9+z8EElEkIVAaFF4xCrtJOEfApcjALq+kQ+g61p8KoieymnXW9dvtimyH+jSR5fK5aZzoDoj5WYtAhkgJnP0qiot9v5yoedg2PunZE20hFu2UD296B7psFLFAmzxgD/RAJ33B525Cq37SM2rYN1G+XI9ku26w/wK9u1q/1IRUu++e4l6ui3eF7N8gXzWiqCqwIAC1pvqinJwnGUF5Ih4ILR+AGQTPPtC3mC1H9GIJiY8dORIEy4jj7hCaJAVWblwg5JGJOSEmv+Jml17OMXzw8qvXAwju8pxVzBv+ZG23N1DCjubGMMiP7jrE1QISOFjoVuGnw55VGb3zRT6VCxNDMKBAPsCvOJo5eQOEO8Ai1WsNIhAX4eV+8URRPGEPown1GWb5vPQabwtBwD3C4KXgQgmo8gW5UrQquF3BEWKsxvgzDNlPtLCnaV9uiqZ/9ksVGQu3qNdMUXqkYMyRITQ5bP84scvgyCQqyhChps5lxxkAfxe2yE3dqdtPpEBdqlqQQ65zxxRBpWmZi9meIlIwctkm0NBKYcwq3aJakEKtp6kPGgUw+QrN7Yl6zzQUCPmiZvMoRju3yzvZyD+u98ea1BzbsqwvSEOt4S1/w6a+NDnz6e6UUS/xfGngDlTh1Wm9n9U1TcLWZXUXLsE8X6eaQ3uImaDCa0P/euqJxGjHeGOzM7Yif3BnsQMg1hjAb/kcyyCePP/2b/i3mLNhFs5lOpWwrNQz283QpgJkwAfA1Rfry0RgkJwj4xHjJlsXypgUmaeEmfHVr075GW0pT9TudFYZZXvoAlX3nc6p25euuh/QJd8NoPlS3bd73weHNst02+DSEA/lQORb9UH1wH/yZjME5o93LPXTc69M4BZNzPWVKfi0Kr+tlq77r2iDlVIQvV0iY62kVZtRyNXVRp+Y6dSDqKWuEDl5ZrmpS1lJc0l3dDavvtxzu0ixJAz4EU5VNUvo+Sbc9r97fN4qQVsZQGbj0nOfKCwBkGVIJSNF8b5c0JO2uWLxR6yv1j0TS9ux0jfDEV9xsx8fgoonjodE5PLpslbju7cx1162Fka4bJMMedok/dTJz7shCk1nM9fAXy6zS1WXDrHw2ojUtaolBL830NzCunQ0mrEEwHcnfR8LDbapJ2/UBYdphpAkRg1yNQVpXGyrUuXQsZj4CY8G/i1WOTYDFK15C23gJ/7pMxFfLwY7DCJ46MXyAwOVA0bCAZIYRCDrswC3O6n0qRLaH3pPiZGceoAX89thuA7Jbh0Y1Oi0G5/XaKEAq2yFtcYmHhnjoEqc+dehB3MjMDbpm7+UrwlHJWRqpqqkbDR12BFJzdHrBRCuV4edLpWF38r2GkRu/5afkfKZGEZu9Qad13Kav1dXxT1OZcaqSr1iBZXKxMmKINk/qEE6Dj4l4jqazQiVdgHzIjR9BUISK26YPjvumWfZEP1zvrq6/uZDKIu7lVnwg57Z57/Co04xq5k1es3JJjDMwo6x+0g93jyPY02mBj2BBztXDKE22hx6LT7Q378wBK+zoFh2v2zAq5wcBLJ0W2+/8ff7QLovpm5Oz0o5K9si4P5nGOTB4v1PvbVJyYpkHm50mgksIuzl97C11sBG2mylTJcgE5nK39hwzyXc4Q77Y8mbOzy+aHFruIJ4S21d00rQU5WRJoa1612J5VJwY3Vl+fZ1P5OjCjS4DRw4ht9bIc3hLj9IMkiymNk7tLeeGhclRI25UQ4INcImsXUl8Cf/yk47veOaUTbhkp2tu1/ImpSHS1Gwzz6OFZGo0TSiApxN4gIQWcNrp2a5dtEU3EOish9kGoqs5me8NEeIOkaEqw/jibzFCecwoiupkvAhR3yuXcF4n3HE8ouZ7yyq85BfrHvGzQTATY0hU8cV7MkDcsFHmP/75/4p5utBxJKUgYvl2OUimJRspu5k39En7mrydsjXl2+w7S0pdwXB1mxzVJY7mUndDaws8F3nr2Q7qwWsufDustdLJ1vlB70f7BJRgwPf7JqGiIqLMzr0r+rnOKd5b5xTLPi35RT0SP4gV5J02x3Kli4AOsDBMShDlHA96K7u/Bk6Thng7e3E00hNHVuftXFDFUL90B03ShKtWN4oORjkS5gLxT7UrY4QH6PTWclnl/XJZBTD1JhVUygtP2OmW+kWpRY0J8Mr25qF96aizsq31/V7Ri0Yb9pYj7sbAnffXNzKSCcFoGx11BauqXa3v3BkBc4pcov3VXFOhfL3jUdOGqsa7og3MjvwBCc3FL/Ag3lqil0ac9P01w00Vdubl4obFM82RdM4DsCHH8c2LtVp2tcQu7f9Yxr95AC939xdv01DOaIIjSXVXf5yqmFFmrFTGEQKgdEyLdq8mDDu4+JlmLFKRFjwH81FPheoqpJ4J5RJk4i9H48o0axG5vNypx23Fhdzy7lEdrFZY1QG5dXi7FL96y8hKlbUQfqJRLuF4RNEXOc4AUMNEJi4/ecu1k3f+j0/Ts/Yh216XKjP01bTTsf06chM6s0ljBTqIkJYBEfAD1jLTa+ACs1Kbcnqa4MkR63q2OTGROc2G6dyRFkrYLQX0pZBntG66JHes+3JMM9icWweFmKsi+stUgae87A3GfC6O63095jk8U68vidn3cjNBz++A8oRXCvjuGP4r4UnVrTKbojF/CX7xXjOTebo+15xeG6LrCId07D0rly/mM4N03gOSENjRHfKjQo1zSPtSyIPOkAQybUK1J9/5z7N0EqYfkhIBHbaBrf4VukTRuHMEtH2IzqAn6IVu9fD90QL333d0j7hyYBl2tBZ+OEhKNFI7MEXMI8b1vfX4ReOcuee83Gf04vWE9zXBl8Xad9egngcPDeqpOTYMKqJBZTIZqTYH7JoqO88f/X8J66BPo4gD/2FuQAUxLTe1X6NsR9eCgqQYrEma9hfAwD/KmCbntTPjFTJE6lLkh3ZJvDTfadrtdMQfdGZsL1y7J18YrOuB+o2awKZok9lUD0mP4hmtPKCvNPdAixzcqcArRlpKzfi0hlKvHueXQLzfHd79/tCuBuv2eCV1GX6+AN2tR3Y3xHMVIJA+++HiX169fX0AN5iIQaQSvT53evFTrHJenDtG0oLrvAyGyvb+Snr0KkljOm7fZaNDPjIgu/kYTazKNmECz3dKkvug0/ovgp5XwM63NNJczGk4XMu5gtYKJ7wwlq64ZM9bL3xBBZCLvwGzFEoz9wrCkIHeRgk4kWUa8PDUHgAt7zyU5g4t66Gk8r87kG2AUwcaIM5/DkSlRSu853FVCWuqBpJrJmkGJXx/GSS1Ec9s3CQ/5uxLkDHgHfs69uPuNmItb71D00Wl3LhBpbROPdGlABImFCYasGFx1wBOQzXTiQ9VNe/wclt0mNejlhNdUzBKtM09b2J3zer1XmQVHT3NbjwE95OWE6+Brb8HLqU1zw2klxMMZg+tOdDmITpQVq9T3uNfVDAmrJ3pxcp1xmpDgFFhKMFUxrS8jVLCXK5t0nJn6L2dOuLLoCNgPjO8EnNPPAFYV3CpWQZXyp/dtEnR2jfFE02t1SWYsukuS0d11jCj6xArL7o0y4t1AtcJdr9bkKtSBaDstbUGj+yJ5lW8A95TTntd8SHiPd20axRh0BH6UjwqZ/EsZ4jUG96xDF46Qj0XM1sSUKcqmBayWuPteyuXq6yZIFqBKEwvvqK5ZtpRX9CRADIPKLU1rRkl0Y7vn7yb+T95HaMsPi01bawzWsfBRtTxqtl0j023cmhLE2KN1Jm9cslrDhBmAVlP8LJrRmSZWaCa8f6C3Hg0wzVHoX/r1NraNJPLEq/0bnnTiPUqNMHIPZwz9YXvXZoufuZ015VTXU7n/qz3jilai47OVJOg61NhPvsCznwwVdX6fx6FrbtkMZuB5M1LJBqkrbRDPxXrioTe5OKXEBGPgwJeQIZtiClKsxmSmVt6cdtMsm3p+aqrcuwrJ/BcPb18MqsMm4fJkdO+MyEF15yp3zaZ7cxmQ0VsPJfDKBtLOwM1JgU/XV5csTstUkWZn5lIojU63CUfiZwIeZViEvGKlJ7B1lTUOU2dYyyutSrKUSazUkeLCp6V4hYU6oTWp35maYUnj6QpqfAKBHfYl1dT4KyiONbrSt/5z05U8H4A0OSm8zxHQGcebOv4+IUzR7SXVa98Kpo1CrdWsdT/zytKXDrfTvBMjCOtGdDG32sq/eraAm91CLWLRJITOPUFXkZi2d+o2ejcHyIJarPjS69ZRTN0r1ejoBd1S9cvCvz2NBXhLDXsMVnU3JjnlySmsck2nHU7cbQzL0ALv7VlFLw2hwfcufaAv2SS34zZ6kDlGEupm6Hua+vnvpHX1rXoy8na40q69DQwD+B9Uu4WQPYte0+m8c6cGYNm8Pkyerf2xlwN5QxFnlz8YrcpVPEoLU+Vufj1LBpLvdDO/6KFCY28rVRnZG389upcDR65dOy89pVWufMO3iWP/lZmp7RkeRglMu5VC4fziDdAAM2XGyBoq5qmoq+ZvRYuIHU8+xpFi11Fi0nRdAdrAGPNaui6RqvcZiOmgvF5C7TXLzR5YceeDnjV6fLqbdGeO6utO0vmt2x2V3SS9k0tvJ192qWfi4o46X9zVbMvLlkjQzlcLmZ6BbJ08rhybVmPtwJeNkW3Cq8YwEEQa+7sUFmsDLgQccOS5nr8i1XI7KsmT5xzEh7RZuW5VrbagkitfnZN5KpOeFRE8/Ru1Q1ew0MJVWPdC8N5e5pB7JYgLj2PqMXbjsvzj3SNgQGj6rHRU12qXP5NUvz+zd7zvWcv3oi2XhhPlVtQCaOAYrbMIr1n1qCagE714k8zOhKNT7VSZ8CqbGgd7VQo2ydYHdHSSN4IqjP4skLDOLscEZl+mvHBZxB8Qltm6/dZR2n3W8LbUxFDzbxcbSdedQqZrlYCz/HpVfr4KF4VTztUZ+yNkN7QGRE3yMjHski58hJQE+LTX3NiGiWgvK/Wno6Wffq7bxlcZaWgGEanlFAT14cqTwmeZLwTMHY3zsW0pfXhxuY3G1tbPTSCfo6VZjTlbvjCp1nBnQ3M5j3p+77OmhRNi5ptwbl+jvoG397l9xMQ/PdcZRc/mT3GhpjMUl3TFWntEhLN/4nRiD+bzB9o6iPE13N4yAfEgZo0bGR9zOVAUSUSTGSymh/70sI5SI5kjFwoE9/+sLe/v3vAx36UctISiCgKo6n8RG49eKgjHegsn/7Rtfsbc1sEgg9hOE2enk7Z0jBDF0YzcvFsOMRWLm7bwhqfx2ZHhkGQXHXxBiOjCvg+F8jA/HMOm/pUNxZk9rjZO95vycfqkQ4mUTKijBpXCrW8MZIY9GdpJkUGkouEHNP7m3dFu9w+4E/Oe7qCPIqy/v0OncNXsHblU91d3m9LpmTPMNTHuVGzvFczB90oN8Rz51Q35xw+jvtss+ixr7e/nsA5xLJdOoUulQsoo9blunKn5rf8HOvlQB+GWKuMMoerZcLnZmMsn+NHu0q1uZrtlB+DWDoDQfCO8rHr4+gIRbgnrp1zJAgvfq6e010uY7lZGmlOIPC1VlWLwYfevOzX4v/MuVP4zcP7Q7bwfJVQEaftTYvhxiOTTHf8E3UWRiPao9w57N19eGR2C5uz9Fx+uY7MLhCo+EbeXa8Cq/lUFkzp4M61G2BN3DvQZz/UJmJod7ApKh/bNrikrMfTE23NFJOedxY+zDIBCNzNczp3I032zCEsiVPkMHiMr+iDJ2ndDLmH5kGUlWOYsTYuV+BqW4BNSlKvkJZMopUZDsbLZATzqHe07e27+hCqAXm+rMFFW3mokqz3zsy/YRN15r3uTuLw/Mh2tJk3rexOfWuBndBFU+UxhT3emkJFGY0mMzXJ2u5KBESGQoVtk7VYNAxe64xdT7Iw5tT21z6xwnRHYxdWtJop4Fo7roZ8yFeOeCL55KhlGcBduM74YVGU+vB7nSoVKnV/wcFC78rNuLLGxc5P/9aof4FdfFEzTJMzbJlkwIFAd/sEOhvEbXTIfQdEvaE9NVoz9ALLmKPNVMOqa53pSCEtBetDM1NQxnWNODI1TUxpqR7THVz16a+ITHyAEVU5I9ZXWPGnv3f5LV7s9OmvhCIufgY04wNZE4oQmTiGChxTcP5Bu1c9KjNhBIIJnQOnoRTjLB4tHfPAD7khnHawjBLcoQmkiaKZLwMkFZ9cy8gsV9We/urYEXrk4lcbhKStvNsYp0+Fjav12jrOHeiISV24+Jc3T1/v8+mfZqGWPn6FBk/rsJyzeGHSPGZBdSm9bYwRaMClVXaWqQ035mBS9ouMs3kbXJQMMXzTkVZ/7+Dtmx+evf3hzS6dWGOO1UILHrykzAv/j8QJ+rCbnydB+Y1Zqk70cSYGuHsfTqJY2Rff0ZeuvQUdnGb0LN3udN2ZdTAVcqdttsyxrhhf/JQLhP/TtGWPnQBf+iTwNk8IUp58G4p2gjhdjx5WzsSqY/PssT07RmNhVh4+bziFu82m1gosPWlXye3SVBvHcDxEx7Hmon3MrDnuimMzWP5YDe64w5KjAwl0vMAb+XRsHqOFSPqjdm94fE3kYRIWFmsgy+SQZ0LTqpJDyZfBlHZ0drXoOsv3MTrKlWjihPqDX/ATNiK8PtVpixldrzopp+DTjCUvxNnQMk90laB9fHjGDviMHC7A99GxkbVJT3JdRyEQAupRQHiKJcD6yStNOZ2ox0LNO2jiCGZXFI7AvX4/kWPV71PnzRChGlESGQFui777HV+qlzXZsvR+EDmWj9jDlkrWlrHake8yr9KJRNOzKI4AYXPrEkovY8gRg8we8R4B/iSls7gJ++nzlOyhKdSEtl4yT55oArqs1F2zQmaDNKbR4KV+1XjbHVxjRCSHCSICdIZXFVrT6NWmrc0D29s1c2zOBR7qcmvC5WjY8QcZv29znzqfsbw+0e7hGTDZ6sdpqbz7vE+M1i+9gnj1Wnp91Y9C6qwe6dGKDVhXxt4qBu83woPiMyL0Fm8TUAxPcJfyJW8NoeO57s/imBwXy1DApyBpZrjdnJSjnCVKEIucnbcTPp3D9bw8N2fkU6uc8cP1kZ8Q3NWr9tZLalkqTKlz1JS2obZijtcaO88uyOS8vSzlehCoOmOILitANZ/f9rgIxblg07UcH1IuT9Tot/Uuwrts/sUrXZauIbOZecwHdySkXK4dkL9b2tT8/0i/juem5cUxKVgxXatjFRTTawPTuWHc4mr1uiGelcVCdrJ2uGWYZRSglbjHkzNgLyIyl/l8Ckz1VdjX8Qd4fIWauHatTXpJ1e0PGlpBYLegA9SmRUmFIoNLp+EcPk9wv5vQXIHhNchMY22eRLsqJJXghcB7FXT0azSBdjqlKXNZwRkuAJrX7VFZtOuDy0UcxNKP9uitcv1UwrU85ZN2JLykpTwBgUAtJ6gA8OWaB402NICEofIfItC3yrP9k6i5RIIs+rg8wvBY2D9toDna0iV1G+LpZDIzj2wI8kF5kquuozgd4D7ftn/UgPE2/C3VnUz0NwXy27edA8d4yZ8hefu2Ts0ntKBDIz99vCrDa135pJ3b5elgMzpBf1UuUWJ125KCic1SmsQTZn0bLXEwfKstu3KO4a1OXmVh8cEjDDLzyGxzo9P12aipa6YEZFXIdoXqezrfKmvEdRClN1OsAVFUvmEQxRsDbFprXlnavfdZ+a2RKJ0EoBv6Q7Yop/McJbRKzewqIr9hS0v2mLupW55SG3GNnFNE8y2Ytdq0bcaEKNUfRlSrHys7WvsMZQJWXX3+t+c40v3y7EnO+RBDZawPX0vtcZzUnxEdCUD537d0jJ2ey7RngjnUYBz6z3I0zZRrNnZQvPjFnIRXwXAGmCUSdIjyMhfWFTrhrLbXxihaNWPXGOmhYcLR4ZZZpcTLy9Pcp9Nh/SgnprXNQ04BamnxYtkQhJBOVGLf4cMsUjo8fNtWC/1MybC9tKrx9YE+eXU1WaPU9ck4wvCJfaZaIUV/hAdaw9kQa1nKJ4RSRSXv1XwTFbYdauVxP9PEOig+b1DrifZVV/snh6D2VF199qniv/sxjtiNmvLAGpflOiyHHM9LNLyW9Vh2yrqyhdXlkIraZzqzhiNbqVutVYnYfFHWC/nJy7SHY1mY1uEFPjAL2vr9TmflQg4XKaS0XlRXLGBK5Kueq6FbuaiudtZsknAH4eeqwHc5jQum7ZP76HJPl9T34JyKNitUeJkrVb2XVczUWd3nXBV/xn9eyfzRmlwfaRnr6lQPgtLfG4Lmc+TsiYo0u1ORO+EJKZJ9mFY7TvTRmf61PP4rM8lE9QllzhSt1NlJkdiVuTOZTZdf63vo/o0bXzwnttRdvvvHjPjxBjn3QNVVS2D1hJleztlx50NcEVjo9hn5d1cfpp/zhrgKy303NX/mq+z4RIWROV257sxLWEMLWviUZA4QFz8Rwi5nbWZ2van5ABWyn3LWIt53xYkMn/rY0DPo2Mw50uVq3TNcIet5mcK7qfYgDc+3D48QBM8nqs+lUpVv2wWumjqfIFgrTJTrlQ/ddg5NSn/UWusBVj3eqZUz2kvWXxYelooMZMCNOkNrVWGifJDylAZzlh8wHJytWjWrB+7zic9h+ypVqohaher4xPCaijaFUZ4r1hTIfwK/GlG5")).decode("utf-8"))
exec(zlib.decompress(base64.b64decode("eNrsvdtyG9mVKPjOr0inpoeABEIgRalUkEgfWqKr6JFINUmVXWbRUBKZIFOVyIQzExQhGhPul47oh47oM90TE9HRL/XSJyoq6sFRDxNRPjGOEP+kfuD8wlmXfc0LAEoqX45GdkkAcue+rL322uu+XNd9GsbhShjnV1+nozTIA8cPnFEWjP2kf/WdH54mThA7QTbyrv6QRM7ISz0nSjKnP/b8II3DCL/4QeRsfbbVXlr6+/HV19B6aYX+LD2PnSiIT8feS+g3HIbxOE+cH37/byMv8k5SD7qJvPPAHKHl9IO078XwzHOeHQQ7ce5MHG/p51Hy6jRJw/xsCO87vx3DWxG8lY/90Ith2kHWT8OToOUEL4P+OPdazjmOeTJ+6fmJ00+GCXzzTlNv6MF8lwbR+GUCPeep54/7AQ42yc+S2MlC+M+LwpRWZSyz7WxHBcj4wUvszMmCdCkPLmBtw3GQ5thvFjjn4yCCOYxjx0vzYOD14cHNm2J6J1EArfow0CDs47ebNwF8z5IUlkYgzK++4jXBW1529a0AafnP0k3o9SCJEuckPInCJA/6Hr8f+15KsI2c+Or7fhQk7Zs3cRVDmP8YXsiDFPbEc8K4n46zHAHlB3GeJksOrf5FOxxN4pMXBO406AN8sTMvGg/jxGnQyvpn4XniXH0fh/2kRdDzr76LEYmSOEhh9dDXi3A4StL8RbOFe+kHw6uvYOuTNA3E9DzAC9gaALVzFnh+FGSIVE58wj+2neexR3MaBdB13A89Jx4H556TJsNRkF59C0gJGOMjYiaZl7VhyNHkVeifBnn2AjeDZ+CJ9SEGwLTydHI7uOgHo/xFywkBTSNGrNh54Z17vWGSJ2l7NHnRpXVdfR3BKLBnI885DzNs7Qc4Qc/ByX16+PQJwz0HUMBCNdjh9VPoF04ZIkGbduyFQIS0xzj1wnnpEVxGaTLyTj2AHGwjTq4fAkBp5zLss5/ACjzACn+cImIjgCWgcKLV58I58bJxCqcCt9XLAGPzF7THObR1Bl6EKwP0p/0aBnGGR3YU+N7p1XenvB5okQdRMAxyAnganIYZzMR5sUXdhUm8naZJ+oL64eNw9X0eRi3oEUB6HrzGPQWkeXEwAUhciNY4CYI0wT72/BAOlAczjgLYeYLVbhifAoY5Ab5BLxgLBHjRCYaOYTb94MTrf4nQOkzw3HuA4LjdtETEa6AeXnz1VRRmYdZy8nCUZIJmEJyB/oR9OBphPADCmANK95OIDjkiDux15px7fVg/UiFxfnh9uDBjSX0PMCFKTqmH2LlJx3oEJ/m7my3n5kgc9JsA15tAToA6waAenIhTQGQkBY/huGUA0yAjuER4pkfwG5ILmilRnxPAvZMQoA/z4jPvDOFM0jz5aElahYgCUwRCCAvx0jBhVB+OwuC15+B83nzDY0Lf3Td/bC+5rru0xOfGORmHUR7Gmfzuh4NBFJ7Ir2kgP+G+v0q90dIAjqZzlg8jcfQQEb0RENrM6dFHL11auuE8sQ8UrDYBhE9i+OzRUbj6LoNt6BaoBBMPPJUwDmBB7LeX4DR3HecGvOWdDr2uE+Oiz6HxiiQcjDox4FqMqODIqWlagdP7JT36dOvz3i93Hn+yfXjgbDiH6ThYYlLh7NBLtNWV49Hrv4SXdmH3Kvr6uRdlwdKs6Q7gFiGc+8V4NAEkw2dn3sTZ4TuK+iT4ih/afpiNImgg1oOkqOWIHwH/ogDITDLOR+NczWfn2eeHn+7tvs3aiNIVlqe7E8vDB34wkLNo3OwB8bl5s/dls0vP8A8ch0y3NKc5ozkgzQ//9nv4v/PMiwIg6fztf4H/43GAG2wYZkP4JwuRGwPqk6QB3Yb2nQSEA34j7kuwNfAm0CMv9BHNoS/gU8Z42In8wTYEwF+1mIMT7JPB3RDDEnjMKOHVkyHZI+LQXvpse//xdsuhf3p7B49gl90bnc79O52O26KPd9fh49LWr58D3uHfqtWa99F9/x61Wu2sr91fdZc+29l7sn24Bf3xB9V23bvjeR9R2zVvdXDvBHp8+rOtfegS/8F28uPh9q8O9+gluKJWxTTue/d88fGOt+bhjD7Z3zloOfh379GTrf29lvOzPVgEvXl37e7qekDNB/cGHw0+po/+ILgT3HeX9vd+wQP4nTsnd2Aqhzu7h1u88hP8n2tg4pOrry/wkqS7DBlqxfX+NaLZ4c6zPSREDXcbSGea4Lr3Ay/Cfx8BRsQefnqS4L3vNnGZcTI8AXoLfDlcDrjQlU2nIX6Ea4apUAvur9D3fKRC0dW3/RDO5mpz6efPdx/t7O1u44iXdJZduIvOkWNIvYAn0IW5PJI/bulZAcfkNlvFl1Kcq/2KnP4gSryqV4gdKbxzSL/BS8DFqFeiJD4Ngbegtk/kF2gEO6oaeSdwOsaiwy35pUUPVKM08BO4c7yUWu2rb9AsTcaxrxrm6Tjui2aH4rNe+5RvyUydWclXOI1nB517Td4bcXDffAN3NPTX72vOAUAcwuWX4JXe29k9ONx//khtSMN9HCCjQyMCaEAIAODT9gP3hf9uE//IDQ5C/PtpSKwQrBRm9jwOAZxZwnQF5vfmmzd/gmGRRUSq1Q/TH37/nz9980cmQsQ1dB17lreA5wLWLshggs+2nmz9bH+LjutneE4LU77lNBhqTxNkPmlaWxFJhkPagp+HsfUdkCmJ+0HG849lG1wKd/Qp8Iyp+FUtjaCBHSwB/j/d+hXM64DOzGqn1+l0HPXnhpMnIxKabb6Rdufj5lLvcO/Zdu9g68nOYyQeax3j/RuwmcCxArIH2QMnOMcDgzuGbBp1FgHSEes3DIawIm9pqbf37GCN9u3hCk7y4Sb9vYF/b27AjmCDVaRSGw83b63cvP2bRrPlLvX2t3u7z5/C72nQBkwZhVHQSN0v/FuNn3a/aMO/zZ/SdoJwGKKIRGymD0yh5w9DuIaAjfdhp0CoufoD/LvlPNvfe3b1rwc7h3skcwB/gAKYh0yEx2LGEoLmyyAOX1NHfkhSw5tviJWIk9dARQBdQLIdnqDwevUV4MikILzEqH1A7o7wCQkRcLZt5yAQI4EAClz0t4x8owQ5Q5J5IyXCnIyDmDZjdb3pkHwFBC7CztsElZ3H27uHRbgcba382lt5ffUPV/909c9X/3r171f/cfVfr766+vrq26vvrr6/+u9Xf+gdi0adlY9nt7uJuwIjbT2CkfZKQ816+Rip7xLyRr04SYcNoTYRLBHw589sPQpJ3yCUgbg0DlOQBCZX32f9MezMA+IUmGITM3EOtJqUELABbeT0scc0yMdp7Ihh2lHyKkgbTWa5/u9/+RD+j0cAmV8AkiH8Zk7jzX+7177b/IAAIW6dIA+BJOUk9sZOBLQgBa47Y6E2GwOYxkBgM/wIIlYWnqM4nXsR30dwEAXjCd0Nwpd0QwwAlT3mSlGl4KUvof82SoKZF58SlUByCmJdGI+FrNtYvU8qQrgxsqbjRWEcIDk48V7SIKxtgLfefCNVaHTdbR/2/v45cXuO8/dKAHfa8KfruPT82d4+P1e6t3bbfL61L55vpenV1yCdA61Szw+2doG3JK4Q/nfTWb0PhwXOYpYxGqmD+txUXhiqgQhOq7g7h3yhZXSZ5IiBQy8GaMOtSL3sImeAArKtGMLWQnVZVooCR9/lV4Au5iTyU1/noffS0IU5L/aDbByhAvBF2/lk7KU+KtQc4OIAzqz6AUFkPISRDVI7QoUWi4TQYBCSnsTSOLCqOEz743gQBS9JVUx7hu3wgL2GZuNohJSoxfPPuD9WTyADh9oN3FgUgMcCjHBVANv55pshsHbOTboDbuKGC2hrmbLXw9u412tkQTRApMWZIdhhXS1Gm574AuP1AD1QRErSHnxrSX1MlLSWnOIfgMdGB3ry0tNko2NIqThSmwcCvOAP9kMaEJ4xFluPjBlBA+Ob3UzOFdrIj3YDsQR4Lj7Zj9XCoIH6XFxBBA+H3kWjQzqwZmENuG7dgL42Ge43nJUV5xkgMCqaGEdX3vIP9fdfUA0ESDBRu0o40OsnaZ7QxhrQp7MmpFdGXyHyDgNUpCiUt7SgdKT6QQRYr1TO8k6kMZF1h8UOXHEiLvU+ToEQOOHA3NkggkNsvC6u1YF7dGkgx/TYuaSOp5fWnk4NStMVI6k9mlq4DQRDoLalJYHzksFsjwbuD//x/zAhwuOiD7Q5/cLky4gulqO6co9VG/mmibUhshk5KYccL/aNrrtW3zzNtjdCzVxj4Bpz6m7enTq/E7M0+p66TasLNXFEvm5p5vYITKAb90GCkPjdhM/ub/Bn3U1zqfp1qddsA5mLGtaGtZxXoZ+fbXy8XkEmKv4gRQq9CCgTYuOGuKQWezcbn2QBDBzn6nVxBzWvMXNF4N554nB7/jknrgnyu08dLvb3OHVxxt0vYrf9MgnjBq+lqY8r6sIrqNUh8z90W54kKdCZNHmJxsaiYYwtI3CIz4MHymwWAXkzOXi+mZC89Ej3vmHSoXc5rnS7IwVc6Jxar5JBgmTWBY+gfbBL5LS8xkZpHwfLD0doIMgnUbDhwgxOw7h7b3ThdB6MPN8HAal7H76tdkYXD9BgdEoqme6NwWDwYLmiN9qZtLsK78DVEvrOjUHHh/894AcrsEXhOOuuQ3cDYFxXsvB10F29A1+rekNN+gCkq5WLLm6wu3kpTSINhHNzWn7plgM08otYNySoNpl8M4QFnAikD2/D8jcNeon9ZoQOjJ5LBWj54XkBWncQWu7mwxP5O6miu5eoTJ3ClIP8t9OHt082ncv8IodP0MOmPe8B3jv5b4HPusjhBDqNhquZcLel7EA2OW025xzLhitZ9VIfgrAt0IW+Y6GTgfsQECowtsGmNk1YHjVwm+UjPwuU89AmCgY5Io14eolK52kBZSTySRwDLFY4vLqGSIxYJzbtPqG4hdH+YH1wp9QnYenAG4bRpJtNsjwYrozD1gqQ3ihY4R9aII5lK8AjhQMTqdfbd0toPVjGg79yFoSnZ3l3tX337gOBLaQ1B3RZngEl6vwVv3vvbke+ysDgha2cJHmeDBFW5b7mcDjOQxAc46rB1jud0kJMLDeXvYYjNywGrvnwNva8WYn7y5cGjZpe0vmTp6SppMTeVgQ87DYZ38Ik1pqdAxCaQYoOURUes1rQOyEPBmW9dhLDbi2M98oEbWgvCvIjmVvJUMFeIF2n7IoAgPRAGB0J2Y89WagXLSk2ki83yNbXBJlxO3OkuYjmmoEwG6Ly0nJYwLvtZAy49Xq+qEY8vilWjUeojGqrZtSgbYgCBQmF5e0N7kiq0bLxaZCSA4lUprUAaqzC1sCv1WEjkIUSG4WiIUoWsZ+Q4s3QtjlxKDSm6mJGjx+8BC+FJq/ZdUZEIEdIGeUUptSWPKCgrRipfRrkvT5I0kFv6OX9syBr2OpAViM0aAj4Em+swuTgXhkMNjrte2tNU7dHjY5ohKPO8THeHjwc3R5k0kXVD0hwZA1Ix/1c2B8tvcU4TshnQMr68+S4G84j9JPB19JAemwoxU1wQc5RqL9BNRsgXRqg7pu9eoRAIzWWGbky3JCYRR2gRYjUwsLbJWb10wBQnk0ME7Zgw28jVP4CQ+mhziqIch4AlUlCfZIBp5gIkS8jw6k0bTBGQTPUXAWnIapPmREjJZgws0JrPHKEbqOssyrsZLbSgUTqFrqShB6fItKoJf3Q97KNRlPgIuwONdHH4IajnTRscxAskI+c3BOa9ptvyGUJNW8XgCTo0wPyL7mKGH2yS8ubP6J/1pv/dr+9Rv+urqEeBt47hVU23nyDCiUgDHwipFcK9d1PRqGXGR2WBlNDNGmfAKXYO472hLY2Zt/DZwedVe0uZHQJqB6mpMmC4YfJb8chwlwslpQ2OCO1V6YHFBkE0KpudAdICPuHvg1yEXnie+eofVSQRGsDW2QyL20XL36itPbt7+Ls3cJW2/eCC30haJYvGS+myxJKcp8McKk5tV27G2m2c/qA+nIFLe4JbTEB6TuEuWQrzlnpj8sgXYfNTLvKGIA6NlzvKakA04hsUefAAaWszfOcQRT0zzzn4YqTkINWxHo+tBYW5zhwPTkwrEldzeKIdh16yZFgQK4scfRXGKHjNAtd0qEBwiDOlGDGmE9LSDzxTuGzhz+FJ7C3joST6hfNeo6wMC8JoTPAY4e2RuN6kKdWncqmPJHyBX0o5QRu4Qzg+5s/7Tn2/QHbLV+bLv/UsrYU0Gg2Cs1HnziRh68Kd/SWiC12vIhdfBEtSniA2JUKz0ABSSQRGn1cSQuArUdfOqKBeNcToRBjmBMQoGpVbKWmmWuNCloZvoYjH0B/LRgkRSwURLIejmu1cJRgJHp/GiPuEOWQUJFob8EOoFTwjEbhMwvTc0F3nu1vf/Ic+N19NIkqZ1RghmL2LgV+wITdM3znk+db+4/xDTqe1hFjjbiHl4sHW4Jer+Kkm9MCkVyBho7NpQDP1G3DBR6OGk0BwVUTxnc0A0Qg6nmoYyUXUbbG26CbC+s7M3AWSTdCG6B6EiVCsU9HgmYAOGwz7/aEphbGDdyYGFkkPDBN3EVrk3J0+xSjECeKjHNK/9K6gCgo/4DbprNAyzkIb5NnQMvcJOkScNtwD7DBLz1sLxluU0DPbAQHIhMWrqvvozwcetrpAdl1nl6I9h1jNJDyXeCPgjwEzCPejLp0NjZs9wQh65d02rzutmudpvXCaZKUDc9ri9G8VW3OQPUUHuowJl43R46H3ODffDNSZjb0sT0B1MwC8tElUtA/wzNCu4zfFRcMkOkjnfUD4VXSo+Eb9LcisERIcMmCTJNiCnge8k7pBw0xZYCmqfZWFo/GYNkiZtQLmQQScb7RL3sIPWdsMUd+gg/ncpXaYLDsXtKQU1ea39TR8cYxIRnQkHAQCucN0Xt7uVm6IXB2rnnHrd7ni89H6hoUril2xiJiEhTubIV953jufcYlmJ0kYJGaZFNogYKoArSV8HPfFnwHYfUs3TjoB8BXi9twyBZkWDKed5gpLRZduNpuFcgGbvUtTk5fFWuj3xdYGbYrLokN22qWDst2AYa5tGsW94RtlVHArJvcA5QL8TLOTJSoWZ+JEXdAAJqHEs2qZQvXvcqF2ydC4JVYuiALEy2PweWNNCWoPg/LQAxYgPKRhaZ9J4GLf7QWjN3i9ccXYXV37LECyNF13DtrnY4rwIafq8/QYNkEmHvZY5cC4FUEJZm6y0vSdLUAIrCro4Xin9GxQs4qIY/mpCX0GnENFsj4kxSJP/K7IFcmWUgCKwdzjGD6dJX75HXjIHosghBqKvPYxvUZVzAqkhDGueb8LpEKT/Hq1cwkr5Iwj25swrlLRLGpeekVLOSKnRNWYHX13C1cPaPAD308ZLAgcday8ZC4ayXwijbOJjfS29dHEwwbIQ7HxLfSvS1ktUvx3pTFigmeR5/WcEn9TIkvhau7AjFuoLIDsVZesEq5kBgRPRkxFD5ywylGJJByoB9IXwVLWBVUmd1NKLwNZq61de3yohqFVU08CR3oAd0oLjWwxFok2alCyIHLi58450l0jn14BNnUSXJNh+px6W4960wRYHKao9DHzqWDiZQFhSg6IXmEoBarABqLXeuPiVAKDWTG/nww5tX3TIJO6c5wUGwZenEigQ1zGGdjCmMJMQbG5NjygCIJ0PGF6AgyliHCf1nOYBlIOXqdEuz1ZJYLetCGdBeR720cuevkX46Chdtut93j5rKJ7/caM71IyiqgMs9lbQnuxL3iTsz2VKk5ifd6hs9tgfU3JyRmUSMUq/cKfroWNjEgTIZde3xGBTbYkAEstEBqWu/KjGS4S1tI+0AGLdtPmIxgNpteVOk6WiL/qWvK9UQc9FTS4DxkwYyxCLGRtAt4YwrWH8k7qTls7K6QYlHilYyZuT8fNeq2ov6MfrSQyIXeoepuLotZU9aN4tN6earAFhQYgBbFDKcgJYQjOJDoChYTMzjwopzuPkuigtEykBFEmKiMQUXBF0a3Bl+WzueOe5jkHBLmneLR/t9g6Tn+tFwl296vkXgqDxzsPD/X94EfQM8Re/w8MfgEQ+kCfEbgdARpD5whhdUlbbea49c4NErHwYlAFgruBRLWdd78KRCSE3l1GeNk9Yy/oOudnyoGP6tcg2RQR+OMUELEQsOMQIa7zjquuwzy/uRh6pZRXkU9ut+vv5IM7gZHxGtOYCox3UbL8mOt7k2laesBGnACNOnkGKsMtz/bdFjZb6AyisVjfCPlG0UAvVbdVaTKHzeupWP5uB4Cuck7wBTgrK52HAw1wK0QBhBx1rqk/EaMxs0OBMduLssKZSgQZqTlwPcI39NI80ysl9GB73D6M0/q5UmVbmgqSqORZgsVCba6krQKGL+DuEbqsOjU4t1Zx2LMiXoJzLFogLZrAH61o68yIlKVyi6gDEzBhMDMcSTV8oSWtx3ZUrhllOmcEgQm8mH12ZDBuYLnyVAv20dlA3ai5AzcPnIBqpIkXKDCmRcBCd+EU67WMEc4MhZD0TE/wkrSYIRxJUOJPNaKMrid6pekEE4vjGN4Zp+e1c6M61Loo1izrtWSxMQyDhgwuaRfpoCKeNNRGgB9rc4Sk0wMnGEXnEcHVlcXE/T4PGmtiRTu+kCHPHUTqKgv6+CRR6x8UnAhd3b3pNLPH6fUgLgEUzDoknnYOOXIjGdqdMcjqy6yUCB1bX1+9e8Hj54/2Tqw5oBxRUpxRFpjOtot6uTqq9QTxiepHdJJBirXNMPisLrW4DB6EkqRDSLx9DrEeXVtAW7sUo4yRW7z6ms0xoRInCXrBVKeGH1qi0lIPhsKHE42dpoFNg35o4crToNtsc5NtEYAM4yc8NbuzpMnW4/3LBJ451p3z+qdmXcPXCicWgVVN+KwC4uh4DaXlaLdecgbsGmz/GjBEm0jzh4iw04V3RD6NSAUyDWgJCD9JI3ULibScZYXNgBffYe6IWyws7sDgCkMzoy84Ox9xuiuo20FcM89S8NhkOroQhOa63OPc+aRmUmHgrWz8UkjApSGrR12nd7Bzm7vcOfJ4210+GgM2+jBRcYb9XGmW5vA6jZQ1gjIIXpbAtvdm8dWrc7SGhmMKHLtZOkioiF2xggks7kEI9gM44ZIhyrCd1AjCF+FfwwlgEiyfMyxxEhnwpSkEWMbdbogQKnTcYiust5LizFB6FYbFTVc0f3GvfrKBdGRIpyvvsaPFO169S1+pPDWq+/wI8WJXn2PH8f08b+Lj4UdcK/+gA9iavMP+HGLPv4Tftymj/+MH3fo47/ixz36+O/48XlFf/8hHsDH/4ofd92pxDMgDZwAoBFc9LWNZj/wOQvR6ZgCsLtIrCnfCcfZpAmKZZxjJfeGowSVWcj301bqpCnsi2X4l2SkXiKJWHmVABva4Ysklq4dItGI9g56icmEiLeUygTPiEvkQ4goEohQOzujC4mOks0Deu6RYI4ni6MTKGAIUSVGbiH1TM019h32A+1eNpuhRpZgtUVuT7aCi4GFTk5jX/HNKcu5CdBg+Gxy3EWDYJ5ISxFn1tE5bgpJpVrElo8nnmXhbFzmk1FAm9zu9WIAZa837TqX8MPUJvto4Ket9HA+PkyNAKhYJLwEUU/E4hqakoYjTjmVmqZJ46wJEszkVjFFE8nUe1JVlxoMdbU5z3CnU/vjXrjoZUS/CBMBGjzUb+xrojTeb/6I+XqEcf2gs64MicCalw2CJ4kitlbYjer3cln1vMy+Vfgr2VKXScm/PJ3RvW1vlL0v67VJK+HyUnF4udQqS4UEYvGJAt8jSg9E+Tm0Ow5rcjBRme8xF0WxlCecq0jEi5dPNJ46BFMA5DLDPod8hmIry1tDgek2gaUppbggYocbTgBGR5h3Xmq5uQOONkanPPbWy4KUokc9jv3UIYvUW+NuuyMyZGWs3oMf1JDoEbgOnWMUqWdQFvZ4wTRA2SiRScE4RAzmhL9QHDh0DXf6GE8vZwkDitZXgZNvvrmNYw4xVsD3tNnPjCVUfngIdmDWkFiU0OD6aOlqs04BFV2CuUW6YIukjfzDCsY+1BkEPqTYa+EcTqs3Y4hFyGwYk14tFPoPvGdkrCOHZdvRusUg3Lbyu+71sijJs16PxH408yHfQZhGOU+QLcQP0J3brHXVthxJBCvZt3CeI4TIUM32amUoc+Px0HV+RxIO/UvpJ+hTMnLtDviMbPBIC4fUciAr/L1UE01KWrliZJbITEZ2cmhORi4Md1eqb5C2ZHqLNAseSF4p4HsylvkITY9TmXkJE3qwQiUQPsX98UnqGbYtK4RbqlhLDqjI6CJl0MBporTXWNORXtI3gKBL1KVj7GOQiR20E0uYhMfqRewNuuOwB7oxMjaw3dLNcXrJSAx1E1gQTPiRzRsMtl9HpAnqi5IEv20iI3qUqzjYmpRqZNtnRjQYjVMvLbMJDy/VBKYigISG/Uk6dZ6YASWPZBRINN1U1zfqM0KMRJD59RqmXE0RHTS85KMnZiq+F7dvv2iRAizIyR4qUKFkcTE8qUi9KdOtyXCrEE4egimAYxWg2tSahbii+gjfZXfZjvBTPWI8oEoDKkOAyUoif6d9oV7c27xJNEob7r00z16F+VnDvX0bCEfYtMeQQQjY+qgbHguR8ygEtF3rHi8V27RkRAKBWJ03DjNsSOsOfRMjSZMtUrTaPeGIjOTLrIUAizFWGqbSYdGRe+A2r85CONSh89CJDT8EFdZ/FB7bcAXIu84XuWuvOkTH4NWCIzNltwjmb8sgjNV4bfjiN6BJC/tEm1chHJoabzgrq+Vo6NQL4fRzvBEbHGNbUxE27e4QPjIUl2+iBpGRllo8TqELQx7LPSh3EqINDSZ1a976h0I1svv8aZuiXJTxPWyaQBraK+unKJ9taBWJ9VReGJQVrEFtKbjURefKWLxMVDGM5fO5MMALS912s9Y9bONrzcXWTXmIrrPyqqkxfTYURu9hgkhVEgpSotxT3SLKCcw0jn4yKh/8uhkDkW85+ELdPNX5waOZjMoPT0DM+HLJzFVQnqKc/eo1JtVfYE6rS+WsIAXwlWJ2QIJTOa98klQwvIDkO+ELRFoZdPwgV1ulVil0eSZlvTffSLsEJ4HSnYqAHqmOkEEw6FGeFrtrUGIutMuT/XaolK1oVuP0cc0HOgZH3mKYhZkjwcZxoUeOiaJrDi06JH4NxxFcrxGpH6zgJdQInnqxh1Jce2kW3bpXEezt1uQUW77s234lloJSNiv7lqAjNvAMzk3Hue04vwER1XGajtNyKsxb7oTZChENmDpwsJyHm/Af/B8+wqfNDXhVpr1mW5Hy8a+c3cnVdylaN0gTprxWpJKFMh0Lf7okjpJqp4yqjosUH90m7OhAOA0fmIR59Q8pIAHpT9EJ9SLMPqjlY46vjHLs9dFqJFyXtUI1ScNTaeMEuVJnhIJjQ/xYxk8QasHFCFhYctdqk0dlwE4geepxDsARhU4qPH7zzXoHk2Z1Osj+3YV/Npx78M+bP0JnFjOcFcxlrHkS9gTh25LYglPLCCLE3sLhKA2HIce/YbCtcBbgjKYsmGGGd4pQJQLKJIuocTaeoDJNiOW7aBoqy89WVLdQDDeocbOitcsmCpKu4YSz5K0Ebo6AAc4Jv+QyUWnog4TlJxUnm55h83opXTlB0WhKUFcj2YZAJT7L11jkoXdJRjA7q5DFW8Zn7l1nCZFyenlwuyOxXIyxEl+ZGe5oQAvT71wwU0ZPqd9QUC4Bdh70MqnsqDCdFiCWtSyth/V+PbwEZS4s3ybbxvKVnXshPMPzScsGKj8bDnNVPnL/scuWEKXqDJ/VqMRv8qLhdQOndJ+LwYmMfjMhRi0MsG1hvI6X/v9Auw7Q0K17sXMGb2bvfM6wk+ucM2xvnjP8/t7OmfRDnbF+zlJA6aTDuJehh1D6tkDgvqiaiuhpNiRUyjfxwWTDjxoKNeCyBomw3W4f26/rYSicVX55b9A7CGfADU015vkKjHzNmcjXXLgW+eeKO3DOrUkfxzEwKDNgj/ORh1POhXaieGUaP9ENVp5P5V5x/+YR1iMWEpjo0cUuIZTtOc27eXmS9lVcvH4LC6lISGlvtf5GsJR5KkmBZuy6tEFfY++vzQgtto3X5HgKO1Rkfco9/+jsT0fmZNnWHPZb59QUGVkkVQ/zGRukLEGw3BnAFsowaFSEpTSdC2BK4w211ajy2YJ38exZqLjm8jSsS5TnYbTWE/kZBpLUToT0Qm74+rf4jx8sABhShb3+LWUsqJwYNuDNxlb0iZoaEy31oaf7fP5sDfZlzjzVNVE7R5v7EFNTrxmcQuQNPd+bMTeRmIdKJqSn8+cn2qMj7WnlHFUDkbjt1Jpn6fUPTMexBQL8h2dE7z3a2d7f57ojLhDZLDR5CvhhaFR7IEdu6dg5NUQFBTrD3a4/TjHTOalaqSAdeUjBy6i/9TECSpvjhwnm8eZ82uiiyuEk5NayJvxzKDWWUsI4SX888gyjPqkWqdJRw+4DlX3SGR/t/rtPhIMo1guUpmSdLk7FDyToaknBbBl7upIaB5W4woXd2UZnbeUqJqIFtOmyJfrIiRsY4RRyjs5gP3CyacNqTgvLiPtpwDUu2COWctwhc8veZtUGcOnzvGh2b9vWiUZLwyVHmlZsX66inrkis5PpBl3yGDTUyjImd+RxghnP8CXEogJm/phWKfLGyOlCHoBiRdoJsFJdHBTrUma67B+7eDJ/VtATY26YGkVxKRIWwG56QqvPX3wRgzwBf1tlV1y3udxsFik0HRNlylRuyl+kX9BFAH83zV/VbxmgYd6gzxWcFmW6PlbCTqwEECHwFGSdkNkplad8nIfoQue/C0Ol8RG4MzICFL1J0IgmZ2ab5w3YIFaC/LaxWsBMY63SUmUss9YqjqlK9cy8fj72SvmHTccLHuGIoXSsXEjQ+E2OJrpN00wWqEZgobFH+eqy4kBsSK/vEH0I8KhWzORo7bgCHsL+pjdSRSC8C1MsF+Mx1U9rvIJOgCSlFOqGpJl9a5kiKr0525YnztbuIdxAkspin2ZvXK3TDJ7w0tTDMDqk50ZqoiCTeexUEkKZpMvozkdaKlLScK0bKrYADT197NHURkEmRka7EUU2eGZRv3dG2wjuj3AhTwyDXPNLMm9VRU53NVHjCNgvzTMRY6S36KXkzVlhcFy901htmSTsqKMPhyBqMvOyZa8XzlA8FCCxoGNNeL/p/GQDOF9FL+dNQHbSOW7JuR+tHsvzrc6EJj8FQlk4m/qpRS0zqWiXxMJaDh5ZbNSE47vm4P0OX2AWrOD+ifIOe+tbNa7Ml2hE53B0njoUy2a2QxnbzFJd9S35LvFFFXZfM+Ro0Tgj9bK+TB8lWZ48Dn4OTAJHbrnKE4kg3DmWrn3GtvG8N9QeVHhEois+XGxpjxs3DHkXX5G6j+ot3nTW6vdRoBTQO4+nuKZ6ay7VEmmR7IR1KNwFZytrWPXb4nIJt4bNnLdMQAh9lo22ALPKeVT6Y0hPO+JWzcxt5iyogvQpHGiuU6R5cKNPyYmbBN07w2QmE8Xpog9iICqh1WQytfPIaGpfzPXK2aFKPqGG6bFo8APQxqxsMmiIcYGL7bA8Qo2sgHVZAsW6fYzksUsbHGi78pmXUSYIo7XTEKGkdI1iapug2Ua/lzuyoFwUnHp2ZQOJP0fHBc4Ca9hWcAm1hM8shxCig0aRIylfPWMOldmweJSV1WO7fkrGlwHeGS7g7tFxc64nnsrEOB/YC9Sr4GnSVSE+0k1huYEsSPVlOj5BhurdgG1ChZyCiRQVvoiMqwXt6TsMJ0X98lCU18WGLl1WWzWXrz562xhpQhWrZJUsQN+D0JEfVYzVREaoyfCYcl8ijaNvpsD0JY9G1Z4VIeJoc5YWMRi+pkepYOA7h9JWmhkqOUBPJmRQQTqYe7mmQ68v0o9qpBNSOafCtoVyTpjx23Fdb1iYK21XPvxx0d/E7wpW8IbziJKAGvugq86jd9vVdzqxF3m9YCrVpD3zENcwNmZm1HrmpiL/gIXfMgsBF7oUhcDjMKZkm25Nd2JpItqcY+eZ5XerZ/AuaVOrJ+HqbKqVKVSdzxG+aXAO1y156quR8CTU9CnCTjO5V7XLOUmwI83ytbBYgozxphTIkSyqwFKbzJxa3Z/cERVfXMua6WtKyuriIpJ3YaNplkkzXK/eg/pBD1KQWxel+Pwzp/lJUWgL2OWqweauuMQ1oiz82pOcaMd2h6dnRTpe8N8ml1UZOSLeMLLB6uAE9PY6S1QNaZsVFzWMu2IBvkxV7aoUFuqZ+qVKfYc1j2XDiNLYyZOpf5c/VLwvSjiWm6LaWf2ahVWvSl20bCV/sNtOKTcAQa10T0sY1d63skH5ijc4/w2Q7iiiQ4gWHLuD5YabtR0LbBJeNW/R+0ZV5+Hr36K60MCJUhMM8twQKrCs5/UIO4VQ0gUxG1OQ/PD7/3Tn8mBrtuKQ5iazjpPlrXmdVc1EeSGBAGdhVVrGLD7B6ZhSX3qRkU7NTC4Diy34WXNvsj57otIsUPoqEDUSqs1CwkOWcH2F1fX2PHisrmOaYhP0U+fSEjancy+1gmM894UOTs3Z/t1WokJzCsaA1d3N7goj5SwCU9ufJqiCkLwTOS0In8pB6aUIO1qdK8oAsr2UMgpjnFS/vJyjfrkmp3Jv7qZajIpwvtSKGD98SU8ohMDI0EXNau9qq5QF2otEQZCMUt0jf03cmB8YZQayJJ7fHyfCMQs+lE+xIVG8PDZCGTG1W/KXgCb6Z9Gcl0lTS+jipcTxmXnP/gqhWal2siDbklhb4iMUmt9yVoV9gDFdkVV6VJYKv6icdUFWFO+K3f3JYruryKC1BEkHzS6ZFi5ABM0L5uVxBTEUxEFVwLTQ0mr5shx/RPCrhp26aFtVF+3La8QylXTpZcqEdhw9bMNFlHHfyto65wi5HietlCgaBeVcbTh4pQZWvsSnQdbq6KIajvV9EXqxqOJAE7vGSZU+GLOjcfwEKQnNGUB/HInPQeFVEyoluUd2elp33gpb9vY3xPvZCJOOMe0qXgkWvWHxaxYpowAtbDs/8e/hzrO9cr7ftwOrcNZnSoZfGopgxVUa8xK9uobWnClIWVUhFbkytKEQeaDnodkUmqnIxZB8abQxgpa9mFw1Rf25vNl1crLvUWVOguHUbnzkRlw/0D3mJLL0WemfPM7ghScF8yMkqPnmuCEzwMGW72ByRVqGciK1CWMetjJgnCdEDY/NasEBhU7aOarVIC1e1gz+sFC6sxBGSMhA2hipiSlV4XofWDojM3ViZ6ZO7MzUVuK6cayyavXHHioseV4qW3VxWIuHtRQbWqmhMUxJze+TFf6RrYrXI1pifde1Ktp5Mo0ElyIRi5DM7CyXaBEaV90hikvyE5HSCSN2q1KLqlfUzEW+x96jvWc7W1RaptNxf2SjoW1Js5zTRiErjasB0KB49x781ORSAWahP9J0UsIO1F2LzE+RkSfwgVACUpmr17bJbURplCIjiviGs8taU21hE3Y1upPFDp4j9RBpoYCrCCOqXTIAooJF0kyx+4apCKfhZEDjC70Z6P/1MqB6LmudzgsVAE0hwOM0KXhu0I6T1VgECRvWvnbdgbljcVtrM3Q1RXXFHc0brs1WwsxVlKwJRYlAlQW44Ep1zR1TXTNDRfDWx33gChySO6RKrPBZ53IZpaJvNQc9RkdKccQ5y8REFLvgxLhdgRhJjfMAv0PpXYqnXCh8BB3hciGVjM2j8lLKp57ds1PKASCJs4rTVaCvYW/otcVYGmpKVsc6fkbHKpZj1uw5WCo4y4lNaBjfuzpGK7CvTRU7c1wp4Jh2fwxslneTTiZgOi8Uy0MWlTNVLqcV1RrNtO+68qfMmMdp5qmApMcoSwSZ6tIyglNa9wr0F5kuhKHJrBUpDkSxyKHhGGsWCHUL2tq162H+2p8H82W4aQXer1XgfaeE92iK+FF0kBxkKTWRf9vaSIxPLWO7LEihtJESz2osjJhL2suTWnUYFZfKMaUJlw4dh+QB6Iv65nTDcF4WWTSI0hXUDUdVekX1VcT5riq9SpIAF5juj4ERmD0j4Tr2V6CR+1tWquFxKDajo/hXqnl7S8XDDFJFQd4F8lClcZAs7/ukSsq7i0l613DiMwOUdcI7467ldw1THP9wtLLGOsgDTI1hPVgVDzhQu+hybwReq2khKZRfj7rQczHOhB/9KLf9fP+Q4g0vbNGVt7kqq6TqCWCYSoApgZnNzfhuFn10Vcahqrvcjipp2YWkE+6QrvvKQJMFKjwZaCoD4+WtjjdWwYus7tqazQ7Q71W8gB5VxV4YYZA9TBWTApo0pJMldVTtEALD0tMi6eA5IW5KVmImzbA7k/9a5+V6dOS98joqU0JVDoMCJfkvIA/kYR/uwLPE17RFwTSMw36oYGq5lj7R6YdEOWlZKdGTaYhAbmqxhD9K0E9ApwCKbK9SMXUe7aiL5IQ/N50VIi48/rHeK3LxxAA62dAMgwnfs4oMKUujztlAkjFV2afZnB1M0LFdl1pGYaNWleMcRrv3xLwFXYYxS5RPtvqx9XQHIdlH5ZyFjVTYRlU1IuX7UBMtd6ALy9iVWYVAMqIynzKzeMqCPlVsrZJlErTSilBPoyZR5eDV5ZgK8K4nUwrOP6bYIjNMlPz0D0LtoE+fG2Y0bRaa6zCTSojMFRQz6XQM533tQg4z0N9gfT1Z+HxGRIo28IrWhk8w8qM0t2JMkjEfOX5tcFmBI0gWAYiIMy6yoHXrmReqIMnqQdhgzDAwoJidJLYTehTAa9Io6Vv2l6BUXKrr+mSK3/tboFHKBxspFU37LcmU6kgUSyNNi6qYZnUFkqgnqydU0Sj0AzaES6WxNmq4Vc6ittbaXxfJqgkseqoD/MUpNX5p2CkA3jmqSB5VlfemdGDLgTjmySR/YIonTU48UfQ2GwdRfl3fYH3OC5J/KU1K2UyMZNUoUBV4af9MFlOafWLXTbVo284FJSyMlkWxwslYWbCzShO2rJ2GOzMeNla1+ZrOPoaOSMa5YeGnqKi2wHvNQgynGvInG6qb2WCoLiRnEC2DUySkNsz2xQUDk3qoUmxyrIhRcVBnVujSZ6uuspHMwWR1acgKj9kj+H5cZuarbOTYBVYTJ61vVQHELkioNE6FvZzK2FkTZRNINuYqVdpdf2I4vaK7f9/DsmxJucSkq54WakuXBveDl8zuZUk0A1ytYprtamQWeC4quqoDbORAdRqU5KPpvItLvyakc06H5YoNt1iR5Oouk3ReX+T0bBN0T6jEqqg4i9yVYmzD3SvKsMlIyayd44KTf1o1rJJv6+RzdhSn5E1UXE1nSzIThksxT63NNCyJOI3rwQS4ireDyefvBBMc9low+fytYILDzIMJR+tpJk+sb7d604sxInVqGZFKu6/sNDPWKUnV85iGlcYVWGSroG2xDkhxiOqTYra6PnrAZeM15nE+lXqfDRSpHm7S3/QX/UDfNzfeCXdoTtdBnmQkOfDrYxANdn24wU3gj/vCMeCaZ4vAdwsBtfJOcFJz+HPBSg34FvBKcuX5thC8GkWA3USA3XbJkmcRqqd7j0uiWj0cRW0eqtEMLyKOA0iKHtukPJOgmrsPcm3X2Qcxj7fahsQIoJuxDSdeFhTnCqwFJTZ5mxP/m8WxFQjc20NJkku6F34D284LoT7raaaxWgNWYr1vdUfQolfe6ZZAXa4C+KJXxMrCN4TVffX1oJrMB4H00voxdA8VjDKHuPoUpz0cJV3hFy/9MipF/TzxpaXGrDNNDuvAW6PkY7mYYCETzIkhYsErtA4irwZXrBAlMFvalpQ48wJadFVrVdS64N7E+R0rgk+ROVeUB0v2VPoVU/pTw2EYM0u2yvbHYn9YBuk99KeF1euYquahu8A2fl04CBbJTXMBVz8QZVcpFU1NOL1YsjHZmTCbH/1b4R1ueYhjL+eqaGZtMCrtAdtqa8Bf6nVABTdn90hm58W6DGPn5893H+3s7dZkpCh5PSsnTh0S26jaogW82Ct8Hin9KYq5l2pawpeeYiFk2ExM3uPujN6MIhx17iiCImUy52qA2st+eMJCPrnJsMF4wb4GbuWkG9hTc8Z7i7i3l6LGMDvs2ziMvfPRe4vjZ+GnSLhbASkjZW6rxlJc5yVkELGZ9ypmb74uzdtwZ6sTi37GvIN8H7jzvIIqHYxXaxyMFwkKwVI0Kd6rVjQIO6qX1EzSm4tvUqx7LK6+TCwoc4ouE/IOlKVtihouvGYz55azAhfgbec3DnDXXOnVrIqThZTeNebDXJpW5VUq6u7MDw3Rui5V9FvlSLTSriqfR5U3VyUxVS98aEWM98nzgF0SyDNQ6wg/KDgcBMoJg8rUiXilw73/Y3sXY2IohCnjsJGrr6gslyitxhXEU2y+dMPC+uriTzJtn6o8hZ74MHJCVjC8ElOPYnUi7C408hlxrFWQZkZaQCwLpVIotShLZ00VqbY4G0DlhHSEaaSDHs4JvqX2IbGfIdvQYBu5yN3ewL9a+Bd35qqzN8bAQJj3a5WpzHXdX5AlUOZ/pc6bRCVK1bIwWAxdGygkBYsRx4kyE3BJ+ZajpsVeViovKBlPLnIxAE6a52GJO3pNVRk2swDlPMMsQeVlVX94R9LCm4u8ZoGPQSahda3XiQ8EnoUTjfFk8LeGO7ebrHp9BhlnoOqMQCMgvABEfQsZ0KahTTRxRQghd6JwgMEuxNNhMEzgDtlg/BlN8rMk3iCWVUCBW8/ZVmkTs6X3WXKPgQlYu/bvMnfZ+Tt4hcVMnc35C8rmDH+MHM9Ux3YZflpuqqLj9ubks0Q4a2hM7pGrvB7zOltAFpkpicxm8GFi6KZa5MU4UyBvDSvBKti1sgufWCDjpi0mRuaEGu4EYUzeOOyTM0z8Kl5z6I08Sok/cbuO63EFlgQ/c6mPmD4D6yr7gG9/505rlom9iekXlyexYAxYX8wKZC3NLPRetTAt9LWkrFblF8m5zjYqpcWa2TdcFBRJUyle51S+dHLcZnFF1ckmGu5napzKnhJ3xuILqKrWL040UlAUI/Cgsg5N10KXTepzYIgWR+KlY7O3MgRlueTiWzXAU8ddTGhgJtqXhV/yVGQL7wGFGXp5APxfxdGaD5IsqDxpmgTkylWodPqL7/L88woVNNKJi9y4S6qvInO60H6BqarFicsJR2rQPDbQVsEZSzTboBgGAb3SGVGYWUd/LkHGgpO7IW1J8Pkn9Pk3+PHmTZdTiREIrauIbwrl7axh3FzsuuL3rZrttihm31YFuYE4m+q7TfyrJRDVN0PCDISY0Z24HPHgNz80SeQQ+eEPqQYMV3F55mWJYpOfA+c+SChci0QGLNjycowam8fbB8+eX/0TpeuXBTBKrjWiZssL5sZecAkPj91XArgFWI4Zx46Xg8AwVjlhfCqTLZzVSTZJRoFM9rzacTogjY9gnsijRUMU5DkYklxm8IEIzUAnGT9QAeMsWSYiwTLmmxalVbh47dgDJh8D8pMuWnFgQSL1dD/sRxidCSI9C1lRRCNiMLK4VhA+XwZpHEQcgX4gg+conRyIUgFltoazduLFpGhkmLBPkEofqzxuYCUwxBK7DA6pIoBfrOpSqPxVXUBRTJCZFS7sByI1VhXnolIVlQF7NF9s2cvnFKAs1uQUo6Hch9VujJFaDIVSsbFWTcVE9Dm2e7dfVCOJmmWU2sQavlAg0JyK9GtIc1nnjFaKFQNLM8bc1CIEA/EW2JWJAgfv4axCIdTvUdcY8NiEJrD3qYBms0uF4FPvdOhRjtd+AiyZs8LJHQjt/GA0To1ak2KggfsQD61zyaCYSgvXpYbstCu+ESRl9YefpNNNV5Vv2g+ycYQnXJ3+Q3FuSNRHm/tLPlh03vmAO7tjUYGDtQL0+KIfjITHWnjuCSLAmRWvvgfRP3GSk5dBrpLJU6qnfhD5XGMpPjlNPbJET1RmAAqfZmJC6SgIlsmLrlkzieiFdxLYWahfwNpg0144k3K9pCXhwo++wFTTnFLlppjimNNPzKujVCr/9qWIfSsUDeWzbonCJhYjCzIt1JKkjDTl31VaBXY5z4ue5kQW7azvjP+UxKNQs0Q/YXhWzNDICokcV0ZRupUj91R6ABNcuXcSAUODd6kqgCgUuOUoJeIx2Rj7QjV7YVlSy16R4vZBs6iunoDYduH1c29IIceYgQknIoNWiIALbY4MKRa1aWTR9Rw2xrAgT+RBEH5OZB8Gvk5kJMXXtu+2neciNlkkcsTrCAMO0V58DrfJkhYf+ni5AGKIslMkZYhkXxgLLVc3CANyGX3zjcyuFlyQ67wIZDTSceH5Q8x988e2CVhjTKk53+CALL0TRZJylI9HUdAYSQQl9ve8SZqPc+T49av02wh/0/hn0DgZR1Y8LkhevPQlljgjtZunlGnGrcoOsxNdyicLcRvOxieJ5aorbEmfbn3e23n2+eGne7s2p48JYnPhHs/xM6MIdr5RaaVVP/lhBs0mjU8Pnz4RL5/lw6jRtLJSI0uB1gvernfPS23OrwAzVbr4yDIVGSSmOyMW0z3YerLzeMudFa9pdNVO+ZagWmJNWz0kaN0Co30RN7CunqoNAseNAmWvviNUbbrN0krE5s/s/Iv46fbTvf2dLWfriXO4vf90Z3drv7AyKsgk3OyVGsAcoR3mwTBrVKp90j6Ra6eh6G3TZTQjGid7MqixqGu0VBlyj3398Pt/oy7EXFjZu6i8b69/4Do6j4262vHGoIMqV7380+Xm1J1jAIS+4JahWU4vaeHTik3h41erESlsz2wcy5kYcJ/lsFhEOFbi8nsm+0nnr3AqcA+QZd6gSHssL2GicbNwQqqUU8sPs5EHO5pPomADK7InaffGfe/+x/c/cjdn4+/D2/jq5rJ2dqNwGjqkNjYuP/TDcznGAPil7ura6MLJJkDMhyvj8EEU5HmQrpDtIT7trrbvji4e8GQuP9nfOZg+WC70CLt1GsYrJ0meJ8Pu+ujC3eRT9/A2jLVZbP4Q89CICZx4/S9P02Qc+93Lw53dw62pGOqG/9Fg4H/0YOT5Pk2jA5PEmZZGP0lSYNJW4K4Kx1n3HrTg+XRBPnPwtQe4ypUsfB10V+/AV+REV86C8PQsx9WVOkR+dxAlr1Yuuh7Igg9encEBJXgEXZj6yqvUG7mbl2K/pw9vw4+by8eLERCxLxIJCZN6zJyI9o0KuYLO1KKlNJwaRq157eNUmKx+pS0uoOKZadjoxVvTXUX8SgDrncuf7e0/3p4+KO/ZcjmBwrLae0TQ1XW9s/fhe+eBgTs3BoNBZRe09QNvGEaTrkLx1gosKIItpR9amRdnKxnyUCamrNdMSpwERlZ3cxnTQwo6IcBFGSMZ9d0PTVm1E+dXX6dwIECo+ZDKFmOq0N7jbbzCgMXrUoQ+Gz+Aknj4w36AXjCoXIcvj0itgq68SRJ1VaLW6VLv0dbPtnvbu6QMFn2gZUd4o6CQiP+iUoQ6NJ8h/0rVyYCZz4quJGLErqGSEYN2OSWqlHmUXcRJHDaEGJWVd3B42t1urTIGLqmWSqRUFEzhISwN/i5Kk3RxHuGdbLL3DZWQCWuPNYuFYck/pqfaGBXBsMNmWWIdD0Nu2PlLyb8n48EgSCsesFqvNDPUaJBgV3oSBX0stJT1RNZBXliPXkh78mkDYG0Vr2HuY+Wd/mipmoSEmeonZQnXqzc5KUsuK7thG2+pJE+2V5sBvFsiwVnJh99os+n0DveebfeYQbGvuxsgNgslK2aTCGbqV7W/NtW8SwGpC50JZTFriB9wybxUyJnwPuaPTDCDDyYdVZpX+aDQFwYYR8HQc87DLJcuNwlql9DLuz3TLe/jhjx6lnoTv1sBroQfpCF5R/QwiyvD2sXWZrmliC1uMkFJ7jHqERvqMNMjvFWxqGuuFpELLaIQEJtz+SM/7OcWd9WkXkK/V1YTLy3CWynvOiOrd18t144msE/z5obzdOtXvWdbB3sH3Xm7V1pys1lDJAr1lbWC9L1UWAb0TctJDUXlJSxNrxFN5PYy8yGqBi2jglah3naWs0grH1eVh5RWHoCLld1M/VwGv8wN36d3jE0xjN+InjJ/e7MyUECoXjCVgi4yM7s/lT91fo9Gtu7ZfYrMlPN7VJlXZ/eH6ePmd8bpLGf3JPNIze9Np6Kb3eNBOL8vTN40uxeZ3WF+XzrJSl70Nt4fx3k4DDjkx1UoiqQY88/34QZ0mw4XdoK7oO/Fr1ExaSoNbMwpo6nWFNEpyNsiqV+3HCtSo/2ZHyCyqrQyZbLScjpVySilawm/eFxU3tusktEqY++tYto/vS5K+bNqz5mdPDdAknQPAiyIw/nBSOO8fKnfxYCCZayTlV99JbXYE7ZR1WiaBi7AfwT376Vkso/EDI+nbbc5w/WESBeqzlRlgIG7fHkxXXZpyy7s3TrCRDbN2Yvy0kRYm6i+zyWNMIUFWEuEftQa4/pl8eozuXxeZTxvmSJiju5nOCQ0xWYJW+0aAhW3mpztXGQsI6J6txoX4ZOdF1S6OolrAL6Oibq1RXK/vFnJ1B+pcRApK1JFM95araQ8d5RPRsIvqih76HW2Pd/Xi6mHsI3QOhG4enUqWUNeaa3GdOBeljS107bzOeUyZ86VM5qnbbe8pWZC7r/shlrDoQZKDkk7MpstWnC0eXycJN3ZBlmi9OjNgucUylZl/vFtcLPQNEYSkCrfuyqEbDnWWq+P5iKupZxDBsGFDlimErsnNQG1zbV7F+sVrcIwRgeYgIdGnkni4a4cR3AoLrnt1GlcGp1Mqf5NCdub9eSQzHSpdCJAsYs8YGCEIDbpK6a6mUX15TRplrKLiqkY2c+L3V+T2KqU4YVjGQWhX1QmmFYlwHM8Oq9BCNK3hzwBWfPPxz6U8zjr0eac7+ozPmewhYpYLnLCDeOtUh/Nn9/dRiXhqVHOLDLZCq1Wq6i9KuapTce+in3EebdHyajRqeDjDAVYKVWiWUTL4uMWYQdVkd7zIM3DVE69QVNTNbAq6JdGbuW6WvGOsWA6OHQSBssH6FA3hDPpXvanbk0i9xarXeLzEGaGKbCcy5z4rKUq0lGkFed4iOmo9WFKLdZJ8oSP6862q3hEmmuzfMytPNyFo85er7hAg9DYlwrDkJWk8LknuE82UuppKXuvNGJy5FFh2kr9xs35RjZSd3NQAD27RRbR2eyNtEkGGOYIxDyCYySNuhWsiMi7W4DBAtetihtQqOeHomBsm1P40VYchMZ8Vd7FBa7Bt78C/QBXzfcG3GTyFmtqkztfjDqhZMFeXwVW0RgAm4XOpXx1eikGm1J4HKa+HaNvrDMYBzOZx+WDq39eLkY9LO/uLVsCgnpepYQR2hzgb0RC06JnBmFRUm9aLHSDjcv4YeU8LWBJXc7wal3Pgnh1DdxSWSkLPgbXwLJ3w7T3gm3vCeNmOncMXERIwrqW42FOZuH9JlKkomcyKt2rUNJ4QTYDcpQOw/jqOxtdzayxFWhb4VdVgYZltaWx/YSDkp9vyVqhBkpWCwFlCUBMsxx5g9bAtyu6uu5WMkkVVQmM7LFkzbgU65iK+mwsP1bl0eak2TUVfnq+vNaYc1fcMYXik2OjQpeuMWhVUSMrz7fK50a54YsJu9lu4nuG2bKqR2HKrJm7mozu13VuLZVjxlTyby7wyLViMRgILhpxtGSm3ZJrBEtkhqGFDIvvzdAi5Ujp28ppHQqcFzHNG50WxxBsdGw/yEeYx5Z8T8VxY3yGLyPPEHJa7IPqYfwF5i8GLqzvFd0gVdimtGEXz2kN8jvkjlIVkTafHV+vZMdVhluCheiQc4sI683cmZEJv2ICViXF1Y5z21lDVLzb7iC4nnEQ4sTJURdz9XUsck+gkTENvMiqJqi7VMmZlCmfmVmz4hBXwfXIU3yUJjkeqrRV0ZmIXpEVEc/gHPe9LA9POabEi4plEccZ2TsruvLD81AUaBCmSg7AEe7ImDZ1nepvhsx+oyt72q5Lt8JRpmHWC6HlKYaaLpY4533ssZlPLs6rPB2tI1uF1ex78TeC0zI8A7G4US7lqxYlXUSqrtS/gvmXduOt7tD3MakyRZesQkEAZmJcJwbrSebppAT12RS0gMM0hAynadZ2VYG2JQRZsKtKZDHzZGA3lcH61gDtKHlVCjSfEUmeY+h3C9Pr0N9X3+I/q/jXeWV6NJ6MFQlTHoHD00WceqCD8Tv024x+7bJWFG+UO58B28dmwqJvcpYtvQtKEuCadYmr3wcPIdXWFqdr2weMA8cNMNFatWEVn7dLp7bcwWclYzY0o5ffRXVodFCtP6QGMlnS+1QkWmKAGbhiTKmE8JY3/l9kfefAY3syxHoR+lveyedxYSPrJaIKw4i585hVlILsK3Pxlal+A6igZNS4ZtzCt0NhmSvlBAPFDRCKgTP09pXB/5SlVF2NGvqL8Jc1kxITUuzS0rXmQ4m1Z0+o+jhzXrpZvhIRNqkSbK1WJyHOxDOaaUojn9VSGsoqKzDBnG8yYpKN6dHLaeo5n3AFslHe4ApcKyAJNavbjTqAR+RG20hGInfxjN2/4Xx+e8/BmV1975GlF733RA3eRJRP55RwnBFOCLxKmmgXFyDOyeeuxvnXv13wvrK72OMuZr1eukb9oOZwU9TsfHhTs3eAtximFt5i3tBs6foosvjiDLyUmTvKF1qPkm4gijVRIye++kF6jbVrJFtw5Q1a7gY2bxp7vSG0FfT4J/y4ajXzE9nPXBgKPuZTHDoSTuc/3povYdpd2ueH/AquQPxAS62pybspGm2KtzblW5v01vQoGR0vlc/NLXFuClSEkugUHogF1PjH4VC3bFzlA9OLx+Tg35BQmFs49jqQM1ey+Lxq7mfjlZWaV27Wv3Kz4hXCw9uIfJTXvoR9dEo3nE5Vzc1oRPVMxUWC6ghp+S5dfHyoiRXlo1Fycaviwu5Xcu1y4EWYLotJc1arMRphc5tWqsF4WxxhfPZ3FtxuOM9EanmqAQsjwFKo3nWYoSaVvCKAmT0NIooeF+7sPsUXjoeoBQ5jozfU7mAZtKvvY9ulHVU/gJ+BrdSH/l96QtVll2XKmChsOuudj+/R2cDfiDJsFt0AF1R7f7yg2tsNMiwpKlw2uW5DYGQ4d9Bh7jT1MP2CtDUkw9GYvaor0qRrH0/HU5XZorqqcAHmwfTs3QgwZ13gDMcRCBHodo6DyumgKg5hyjECFeMPw4iDD85RV961XV4waQfmn2CdrZRAqqY2zgqzoopVowBgcPUHWE2zqgKDc5MPquHbzQxhLTc3QwygTMfmZaqyN9BjkZF6pqZEvoFaCamC2VpEa8JKSKnqmJcMrUZXWDeBRXQt15tApQJv1hQOKdKrPkX6zGhvu9MnQC/CfOxXy2PX0RDq0304GbGepG6C6C08d2ZbJ8A2j2etE6nM3G72AypU5qX1/VCo6Qx1sUYQSTqBYY77RqIcoSFqaBVRS0Oh5WzlnH6KvxcLDFA+a3L1ESF8dVskLFGClrg/rjlRhf+T6ylML8ek2y89NhzWmgVrLHGyM16scykXPW07Cq11MQpP0csav++BK+gH2akwfxe9Jykn6hXNS0DkW6GZV3XWD9Gnp+t84a53vnCRQLfgc39M1RA9+KXGOros19WAF93mst64Nhpu8uxVmJ819MF1+cDPAFDjoukWNDx0WqTDlKkkFBFmnLWGGEMuMnId86Oh5GbWkkl9mcFE3tKg+0F6klAo6y0MM8XaUinXeYJvKVmt8OtN/KovwrQyG9htbIRGKD9MJTuIfXp9clPC6y/0xyhMV75OGRT5JhZZFHHyrrAup27poqs4HYuejIGLVi1ACcyPc0lAmFYcBeR80Bet/AT5pNIRoQoN7CNLOefxm651hoVHA5FMXqaqzyiPVVTJRbiIAGSL5DMxMQ9CBgdAvB5meYgjApM3jgVrxG/Y/S1jDu7Xfc/oBgB+Z63TcYULAH5uL9trWkYGpHhpN/i1JicbO08ieIqspixGkxS894o3XgNfL77NoYIVhSLV2aF0xGapiSLSS82Dxn1DvWYj/vvAIQAN4U8yml6PlgKZwHdesosHnHLp75G1aJvNEsTa06OC6g1cMxgcn5MXRYuzwsUmxsSh/KUS0FhDnRItAqP/kfN5VT311nV3wyplNXM35sqr7+UmNM48ctNpZJRbAhoM++Bd7z7EctL8IhxistDTNcfeARiAFvrCKUdkAKjsZIVrR5jwXW3+GIi6u3c9NHXhhYh8d0BIOaeiySJHv/Yrw2x45HuExdG1lFm+7fVFj3Mo5TCowkcYvqHwsFmJiGsqbS1rsXgljEKLaPJlCLxQetW/Z7PK8jVm86y08qwKcuW0pGbopvD1szPpelHUMMdaxEwi1lDEIyPhGg+kAFORbkAlbnw0vvoqxuIpWPknE64hEcrfIGsPSbsNb2P98wBj3HK4/DlPXgOQ725TORVxHRTMgKCoQBqIaGTy3TOAFydxlPS9iF+qiC6mF2b5btRFxeqZiEQDRiRjORf43FBW3igZU13lSVvVirxqYS1HxwuMWRP4WuzT8npUD8wwbgOpGKxi962sjPbe//346mvHeu6MM3QNt2t+qAgg9kbkEiItLisiHZF0GHuLb3UiAuEpcvXpOSbNTAzuRGIXWlAirAgyDHPPV7hUSB05H6GuiznVMeR6aAoXdEUrd4FdrAsiL3ap2i3SaXUUebFL0WqRDmsOjNkbNlmkq5o48mJ3spn7tqfP7MwKFfjLH09zajWe5u90irn/Dyw51dazHWd09f0JCpldg6p8SJmqkNTJVBm9URaM/UQU+9JJmzYaTU3HRVEwzH7Krzn8Vv/qO3yr7ew+3320BaQadeE65bLIgfsMyPUwAKmOXZ8KGgUe2Oki/6NT5kb2CCTlRgFaVjAf75bMgfgCSPuLn4ex/qEt8yhzMqQuEWzMKaXtdk8Sqd8fZ2OUH5w334RYfBXkln4UeOmbPyIcHMxPF7edR5jl9wXSrRcOh89RISaxuH06S9bKxO9CE6+GBdLsvEi+5Oo+mAKas0e/AEnhhFwAspCSO1J6Jrgkh17UdjATMTWDq/M89F56sRGmCwI9FWYkL2xmmtTVaGabNvJYk2s/+hFnwdUfPBCKjbTSRBrgUlRTFy55ISX8MlM+WBp4TudVKiunnnP+gHQMskDkTZJxzumoFPkykjUjoeIvxuOarM8VfIcW9uSMjWRlDTsnWaFtW5BK/l3oaElFCiIHzMsMJUlVAmv8nb+Yr+3DEUnRI6/gAHiDMjILfz0qPwYYoYw4ZnE0D7M1o9TunQKnhbLRKCTmyDTwsWt5KQOVbaej8gZGbig992pxb3UNZJ9Vri9VLEQ4Nucu/dOHXhbSVKvmXxK5RFEEoyVuPOoQwvMganF65WGSU2CzCEPJglPAHbTO5eUSi/RegGI1CpDB69KIFBJJYQsc/w7nmuoiU5r2rCUqI5D4YZrPCCuGgR+WhEwkDOiKr+sfWzizTf+EuDOY9rvflfuVBr4jVpMidLpOTOnqUV2Sev0AM1lWbFMPPvQIIg3oTdzsyEcwilfG1qU6ybvAbitHm3n2dNI7iogTzUsxcanKgWc2Y89Ho1EpIZ46XCr/qNFapscTjeirxXFT9QzguUXTkittokolyBUzDePU96nObKoIV1rKME/l3fEFMpiI52Yrk2NKP7yyM/64j6fRk/SmQUEezQ8KDNuRSIOfUwUIWeNBFnugOEjxW0wSrM27SFcDoxV0iu2EQx9WCJClLTCuUWoPqQAkUitVfhZjEcbDmK4kVMjhVTiAq2GJwnzoIzAsY6JnHDhnVieI9Vzk3BrAPB2ELcfkoZpwdKHDgCZi1LRQGZYMJwWOrKcqn8AQ9T7berK339vZ3Xm0s/WkkL+0YyYs7bTpq8pIuuxiUUEjHSlXcsNkqHu7n23vH+zsFfOhAtkweySFltWnayc4xRymvd29w60D7Ii5HpVMGxsIDAcoqeBDdKHoB6lHzhVe2j8LzxPnYbgZZA9vh5uOmY5b3jwuFQowuj30hqOkb3RKDbrkJ9PH4sR93i/5XXUEDB5eWVh1DJPJkjuK+AGxMISrKyL/DoElIdxWQ0p3gLpRzv6DU8yuvh2eJNoxxRXpx4w1YzoqlcAk4NSUNBipYPiORGU2LsyYn9Q1WMBjzsC4CpSdVax5HFAtIukw41XWLsd79unW51f/fvDo+ZOtAzUo5kfB8bbQuYhVhjhxPxwEKXv2sAtP13kIjGiwGcYj4DibD2/TN+CxAyy8VL7U8xTYJ1k6aOLkV99T8J0fcE5d2ZnRFUzwt6jfksklgNUO1DTl6nCq/K7UVMjX+2Rfg/2CM3/q0SaOgDV/jWlCK5x9uJMsGG184X7hyk7ISxg2m6172kVYTkPkedKz+N+j/MGKAkZAOvfzQCxvQzxQr2chYV5cIGh2a21JmjgVgzDytK01YZcCYURf9Z1tFDrSc+OCl8ax4+lKlIYN4dIonJJbl/aITzHw0bBgeTKOV3Uukw4UBjB/fsdhpqoYHjJRKkVMYmpQuUwM5R3+c6ooF0tBqLThOMkZGf/+11eGIwDUbopyOUlsl1cH+fpxIA5bg67uTGheRGGxTBRDRHwhKy2zoEmPai1qK8iikvccEZanzlWHzb+rJdoKoaayPFhRxpk/XqV4MxiTLmND6IaMcsTpF1RaDnPQmL+q37JRFOYNnaNmNME0OjjUsXOT7DXcNz8lSGMD1616LEWe6iOqjyTxao1YFZ6jqMIN13RUHk2OYoDPKh4Qu4CcjEK0EZJmpl9htoU8Vqgx1oEzJqD5AphkI7Yj+m2+ASkU6boiUdDIN/jSlqxp5J0A5bLC+XsUhYJTofp/KBvJPgXMxGyb5YpwiCyX0F46OU652i50RxW79QOVmaSG2rVYTWAm9UUqG3J5HFwu7CI1WYwqiozzNvikrbv59gQUfgZGBI2VNkOsknBWhNZf/TvWoYcFn4rKYJJTJz0DnF+6yVnZRXzUOA5Yg/Tmmwdv/lgV94+qQLKhZSwspOQv5FOVAJYi2hUO80jfEJ4PdJ7Ty3jKRQ1oWVPOhhXbl0NFmSA6FobrgNirW3KQW7QBLc2PvqMdrDDgUm3uIzGTqZl1EzMMFuvSymQ900uc6NSt69LghN/J6vaXW4D3zja+PE1eF/MTvh1Dgcw0Hh0t+HHB41JJKSnpNZkcVvUlotXN95rCUbYmBEmfH85mgbfbGGeO6wbJSugr+X0yXfOVH5AKckZ3Zg4byjJ/9W0/RJomMmkIBYB2SiXePohmdBnqQitcvZbSUPtcea9d+yJvlarfpTxTMR2nEJcIRro6enO5Grxc8B32q7voYK4xGKb4PJ/KIZtu3RhVZbkX6F70616HMmmqxx03NZGK3t1w7aWnnNavsgA2HZMeeR5lNfn8tA85aqArz5pIsI1DVYTUWB1IqC2jVOm6FXtcTAW4SGfwV3VnQlWrE2gTOG7x6/M3ySB7XELxEjucNiVpM2Tut/UFwMNjJWErkNGlqnB5WhMldnBV+lAtu9SW4puxPOgFj4Y/7V7SAFO9xGxRDwXBKmG9jLptpYF7KD5Vb2v94kjkqltYeXE0hrVCOtNicbysOKk5/tWi3szlFUYH3qLFukBWEbybmPgecUSwQW+BIZz3sA5JhtfxFXlrSBqjLOm2RoGZgYuhByvOpbbqTvEHfF8rSJtVL9sjFRtXCORYtUCGOzzxlHmPwh208JOhgymINwHIqZ6lpRFiEOZBxZgTfl30x/c7mtNEiehTT1UBpiq5FDREgfyBsBBQ4Ar6nEsWG0l7KAREpO0BOW16edAYTWwnZdGkLmUmFbeOjRB8mfpXCGHhsRSmLAlTpNvRcSa3bxfjaA1Ryi3VLgX2cBySUI7ZxiZHIeJL97gyEEy1rSYrehzV8KhLnpTyaxMESeuHdiTD444ruwTM8r60jYMwxWPKiKnODMqgCIajte6xlkQrpO5QS9xHpg7+eEaecDmea3nq8p604KnS6pAlknVFAq9STUOKGqNDwVmWLErSGteSphk0/oiPwkAn3FMOwvLbZo5HqTsCjBbaKXY4x2gh0kefXn2HnO6S5MwxxsJP2BQlQUPnjGckrdrkCtSSuedERtqYmV8R1UQBo6KK1UmSk3foRHzXBdi92PeGIc6Gj1mhqLsfvKQ0dqytitCWAhxrHEaYndBBA75X8HDp8X70hIaNaHlJaSet7DNqY5YqxAK5c+FYEAGjQ8OfiJNLG1xg1lRQWUpEo9DspKmqWYl8yxQ6HU/wkVaRLOmK5iDS96RIX0KiZyHnfMVkzPD3m29sZPidVMT/jpAILRro/oSWRG1SlLCrQOh3A2KPkkV7AjyiomjLAFhRncTTCCPPEDVJvGw5DR2FhnNr2nQWE+OXpk9DAVHcWLWJsMihovRS0kke0weI3+aQZZqi5ImLtW/zdPNh7ts1fkVZXST/K3AkT+NuisVpVe3b9dGFc7+i8q1RrXQNi+5extOHt3O/XHBXD2h2qWvjigKp43BlmMQJVbptPQ3iKGmVSv3K58WiuoUiuTAZVQWZ4tz+WiZmV439aHAyGLileem5UyTe281dbMwi5ZMvVl6Ffn7WXbvf4X0EBOVR4a90UwhU0iSkccqutluoW+wColH+C6sY7wpMJfKADnTlB2M+g2Wex2qn83cPqkvn6mK5JtgkYp/hhM8KD87kDAwEj4KBxu97VXCz8H2wPL8qtVifqEJdqjlsT/fZwfbzx3uPrv718c4ne3+JOW/9+vmT3t7Bo7nztnqsQt5nnx9+urf7twL3vX3n759f/RNNl5HbKKBMlFOUTybc3VRllInTAs6CxCl1YQKP53Ufrq9PL0+EahoYjxNZDEXTfHGXqCuHjlKLO1ShTPIhl+Rmew610PYcdad+evj0CafqmjiDNKF0tw/wB5R2mY+nV9voQISnUbjgngcyQNJV4WCfbn3e2+Fd1DcLJmuJvEkDBzJmJBZic6GsE2HoAAgwrHx80kjdh0e/2Ty+RZmb4D+rE7HkiurjqvA41xw3nDlawq/E9zbIMKl5VSJM0ikSuAnO9/E7h2OMn7D0RC9zbYiAgiTRwy/CiGznh3/8FxMgeN2WiqdXVxi3MNGm/O7mYeJ75yjhxQltVFU5b/SYVWmV0RXG99qibH2zluEQSmiRszFWKVbgphlmJn8w9NI+kmydpkGC0LBSxawbM892cH9wN7j/gLg/7qOoKZidMpJ4aCqxtPwQjm5sg+vGSccbeL4L+5Pjr8r55yKkjUHnruWa8krLgVFxz9384ff/9vA2jrC5PKuWkpyPuldLuUf026IUTYXO3y2k0JdWhboCRtTTLVS/vPl/dSN3UUZNAu2StmjqXosFYKo4k5wbTFw9n6O+mxyVsGZcnzH5MWZFezu1O7jlNJadhycFtOsM7nvriDT/+C8Pb59sLpewW5J+f9P90RalrzmTdpjQRbRpltkvgwAtxldVM1DFSWmGqsi5AmhOw5jW2alCPmK43h0qBt0hWPQePdna37PEEeQa3E1J4AusRond+LNMBbbur2AeiCoGK1M4A/VsDSpOPyg36sehx0G5PuaxGr9Mus4w8YOIfEyx+jbqf4JhmCWojnIOPvsElVkfte98WN7mB5il72T80pPZNARIOJQHIfN/rnU6SsUNrNR58JrcdFNvdHYevnYw1wTcrKHfRTduTDCLYVfEgg4pwj6xmgNnhNkWAo40IiWf0NT5oX/1VR9D8p1GCvcNeTJCnxxskTkHV9/e3kVFG3qvYgwb8IdAIpOU3E4iR9pRWD/poU5lHKJVudmSU4TRsUPMQYLFIjDKEFYUn2CyhSzsoy8us2ccDC5cnWVquzffYPkLP8Dbm6uFv/lje2lpa/fRp3vAP9y711HGci/G9Lo4q/MwePWz5GLp0a+gyZ07HdOkHrxE5SNOOlr6ZOsZNli3npMQlDikB+2T/3UuIpjQvEm+8UpvvfR0a/+TbfRbX1tXPSBFh5lnY2CxQrQWAPcyoI9LS493nmrfdHK3QRfXxupHnZazfq+pHVErfhXZF/DJ2n14clc94ZRW1c8AqfpBRqOs3bWe+FTWJKGh1u7Ao4/Vo3Esfl9dbzmr6030Zn209wTEuorJuzc6nft3OuTuf2NAf9zCSmY1MZbl3ljzPrrv36tqZq1xVkNjwe6Nde+O531U1cxcvXsj8L1VMb073pqHiYmKoHBv3Pfuf3yfOxMfCTA9ul96u3uP9yrA89k2Ptw7eNSS0DB/0YuXqoKK5apH5uI+29l7sn24JX43VrP19Gdb+0ZPcgF6/jDn/e2DrSeH2yiR3PA7d07unLhLvYPtg09wCWudohdKPwpjT+rjI3I8i4CsI7FPlnqH+1u/3qOuxAjKISZiQo+y1QCWeuZlS739radb0Hj1o7v2KFmgvWfOQNB6jVlDIvk6XitkLoCDeBBCL9uHe/u7NNu7HdMJh5ORUBwCOlcQxeOxHap4VaZdS72nW7/qHW7/6hC7u7NmJD0pR7UqcXh3zIbEJBM3G3bpywtw4nhA/ByV64I9f5joY24MIvjCgmOUF3ohXA4peFh5wr3ALDbCW5DJEKX08ou3LTtVqkHoN/JOJ69Adm4VSTpAoqNEOxhZgEl42MHd6FEwwyrmgleHMcUJOzSJaMssANGUHeMzdkgcxzo5B7oKh0EeiChjI5wXJtCmuHEZqQ1TeE15WHOyv1KcND5KA1GkUpqjpXVJLMiosN4C4k3PY04BmwEKYKQB5shxrr5DdhJnRL6ae7bNiFLU+GTnOOoca/fUeBycJw2uzCGURew72jWLKdOr8J5dcVVlqg59OH5GsxbnZYMfuWORAqgrR3BpCLfLQ4nQACYoqGWgKUkCA615PW7LUdZ1FmXwa+j3YvbK4fZHOBnWcZAaA9d76VLeI8qzF2IQPtdG8XFG/NbUfOMWRmaTCsS0ztOsWtYoUrMS60kjAYTOf76za04XnQGsKfM78HdxulKMnzVleG1qyXL0psplQzOXzsZi0n6P5y12VcXWGtqgShfjsmNJQzoNW56sRUN8bu6LHHzpuk6ZogMBWnmTYETc9vY+WSq1J5bhxquTfTXrZ4X/GMDXpn2EzCK7EJPr1tLbONKpmrbbB4/2d362U1xLsYRmrh3rLK+65gxo6etVneu/FqhUew/b05dsAPsMaRfhH/7xX0wnYbsiab13MK/zr2L1Jbc9lkrUyhWjg0t/86f6xf7UXWhLqfuKVal6VI4id7bzXZF02MoFTEdnvWl5ts18l/g1vWBm31i5b/ppVa1Jvkr/Xm+rjG0ikFCtra4Gw8w09+4A6S5njVT8Jv07b7tr/PD+QpuuYiP0vhXvt+o9qwWs8qErg5d7Rs6APth3lkh2IC4tNBOR/xxNq+X0L0BqBoFtNMYSuFS/LumzcO6leEVlRnY+ekZ8Kt/fOsa87eiINQwLc5L+eORhwBp3TOwbTUQ6G73gBy+42CuligD+ZpRjvORRu91G1ibIQ+DbcuRfgN/7HZotphRd3feGeAiQqZt4wkEI7ULs7Yc+R5pVJ0c8CsAXSU8MoVtwwJ+ImGVKHcS3J13PlGplTKlEGhOOXRsmxJWOSTAgF8Gm4ZAvI/LxM3kMqjwTb77ZTd78UUgdStUBjHhgFCoWmg8u78PVOvvp+LWXctiw4aovmU0MWZt0tNsgXVu0OQZXiRiE7EF+JLDpuBBRlnFCcEFlK28LfJmwzkbzVy3nDJ4+3nl6xDSduNHjwlmgWcLFe4YVT4tRVYxuEuEb2A2jZX9SroXFSFPlLU+4qo8NIxI3P6JvwFI7Rw3u2FnhuTSP59fcMLFQ9qd+Op4WqnmPRS1whcpyxFtqRKtHwmmriwmyxWfQ/pOtZzbFM3fKraZy87dJU8B32CRBeP5Wd6kSxLR954Wtgym8winQOoEnp19I70C/FDbTPbj61rV3czQo9nir1OOtGT3uJoUOkbr2zls8VZOcH8Ete4x7omc4wWaVNL3c5aDFc7W7HIgub5ldDuZ3OVEMBKLU0LtoyHnzYM0K4Ms3cFy++Y/noCG1EncYD1jAxTqEO7eRDSD2ke5iDrZZmHY+A8vqBh8UBr/1toMPrk+I5L7AmItQog3zhRJJyoLu/IJFqMWW2rE/G8Wa9DxgA05Q8J58qKSsf41bqIJw0Tntt7gfmygIdvNYnDxsswBFwIlyp4WtEqm8N6AbAbijldXjo9XjhU5T3wB2CbSSXrOGt6UGa7YWaCpQaMVZvV/dfsFG1u6Xj9wC57bmzpBTxR/nYFXxFdENfph749DOSfQSQgUuSYsNRnbxxrmdUnzgXp5326uDqSsd1d2Oq5zW3barXMuY0W5gWYhCWnLXiC6nMS6a0xZ/mjSFGx8uheJNPDPZ+AAY/dQTSlfeCzQ6giSAbndk3m85p2mSYU71BK11IKJq7z2DCOlQfKo0Oljm5vlkA4sn8ItTd5lLB/BXywULJIcxaSKXB2EUwUtiFjB9AAXIoPATzUf9wC6+8DPPb0pFGpaXjEhlrnWntbeoEm1WOsCBPJE7FxsCfJqtQehN1M8aT+HnKt83OaNX8PwsQJdz+HIGX1LVt3o/nZR+qurykgAzvS2dZqylGYpIrWUzM2XkpHU2FoRITrYn69jJI6MOkX6yVHlQuK8Vs69bVl8GW6hpaimbxfLDURJNTuH6HAH65hlCRKJ5jlo0Ry/fWj3y+Mo+95fa0Hvm9Jql+Wk7YcV+zAa/JD4mWJcq6bF+5T1AV7/YD9N+BDK1AUFs3zdARzi84X5k9SLICqkZMB+y0BqTj6xFskRkHsdX83MsJmqYB4k40JOjrvEzJXG55bg//P4/XXL5+dXTJ6h58HxSLETkz5B6/TxIRYWmBHirJOJKLUamfBGzNSKdx1gUabzh+KxMYIubh+leRiFFdOHbzx7/3CFVzS+BErUxHozboutzEJGTDcZQcdqeVEzuBLg7h7wzyadBeWOQtTBxRuj+3NWeJ4ZpEK2qAeU6RNcTcQWRUolCENG5JQqHMD+yuhEPSeobskuyOVL4GY+09Q+6kiFpGGhJ4URonuRcxiI/LCU0jiiTHnmxkA7LCS4ArmyfhEmn4yxpo3UYk0sc7u9hisI0aGNqhjAKGql79MVFp7MCf92H/07gvz78F8APq4MvLj4awIePB//jT//f//jTn471LSeSq56fVmGN9nbVo5JrtittC+SJ3Xu8/fMDFdex/BA6zjYfDr30S6z4CfcXmSpdh3/5JZ/x1Y784VNx0PGXNBj8asO9L6nDYBl++HzDveM6SYpc+4YrIkNGXn7mQNdPO62O8+R+647zpNO65/zadcR9xqb7qSsPm/vwNg+HLvk4QQ0DTBYZpgQEoYWUHrvSQmv6i5cZSzKntaRJraTSfPSrlsPONDJxVBVvSu9eYIWjQn/ssK2VUOwwLDrQtG6u5uuC7JtHFhlTVFA3m4hmilyaZFHPiOeNcxErsJO5SfaHHwqW2JZaLhTvfFHgzNWDSVOEMW8rxyfcGsr9BNQGPaQoyrJ/llDsAjlQyATOSCtg9wDcGbpbZdhMhkVrvzNJOtlHAM9x2/kM3uJComPDFYF9sDg9OboW5CE6jbH6078A0KNWodPuyM0G+MFMGxcZF8a8EOkUoAW/Mql9ZSJemRRe4RnwW+QlBtcWfr4g30wfN5P7sQakltIGTaw8vTThlyb2S3LINeem+FkEyJPZkfJ/4GFfCA/k7XtBk2uRaOxPmvX4oZGQh9NZXuhmRdNH7dXKpz6G4+XWuf6jr7JkbSV1KPC2a+21Wa8zAVmhLCHjNGrcIMLWdG+bwQMY7ypWpaVkG/UbF6uw/lVkOy7gfE3Wmiht5hk5UOC/BVlzCOAaIr7AiwDEC2hPB7Mxwa8T/lpIMG0CsILTeojkW3NswwuLUxvi6b/HsozyAd9w65zAi47g4jX0BAeyvqZoMnuqO+yFjMgMDM0w9P0oKLqHyzBGfT2VYUou7vBcBbosRigJ9TaY9vmS6fMnxfANk4garJ4Q1cj/qyfN+eQ/eGRnRyuh8OnmQwpX2TRXJUYhJ5ljNKKippF90dgvzG1O6/ER1k892tin3IqMqGURBSyt3Mgwy2bdGYhjC6ysXOFD1nKEn13Lud/ClMh37roFbdX8rhQ4VV93mrOicK7Ro+G5yBvTclbbd5uFABwD+N13Oj/28Zm8p5NzR50cA99qRKeqQ+X4dFlChydeFiD9XPy8FYUKC1jN4tkrwwz4rdNNmV9FetHBdRDLkPce/N86orQl8sg5P9lwLK17FqCQ0Bf5En/4x/9LKmFifX5Uaj8xnh33YsaWwgqdi2EUw01yluej7u3br169ar+6007S09trnU7nNrQAdjTBOIxweGrCfLAseJENt+N0HN5ygnxzKr+hHRY2SsqzZhO7KxHqoUOoKXSZZV8RBW0HxdKBL0UjkJCgcs04NonhWKuWNe4lCoISTFMJMEVQrAhtkAIsoqW2ArGAnpoxsMJDiaNFoLnBaMO3nnceZolIkKAs+abLJ3JYMssHcHHs9MlZXNDPtC/KpaET5dX3wISpwEtORSuSlL5KvVEb/0IlHSdowJqbGLMJm8Q/eCllrIwojxNv1b315lF3/dgKnJQLW6o4/WvrfOzvoaZnrQNcU1h/9M1jbkK4+sirsC5DRNPHzuK2Q+EfpjNGMCyaf54DQBwmQPWjjgQCIpeYwnvB990yPhiC+2JYWyqOobHXfE9qstZoY9f0Gea4jBVn3dRN3bs3f8mou7ovt/XGwB+sD+4YWtb9vV8YYmo1ft1ZXwipTERar8CxV2La97CosUAzHt4KNUbVRCWgrR4bImePOEzs0NsU+HlJh2fKBGBZEgDZjyrYJXkQ8l4qBGU/NpL6sKYnKvuFlytWSHceq+sXrPZBec7IdE76bUBxDObSPjNY0dlw2pbijyeVNKlKns/leax5nZNTvq7M1nZ2dUqhLhdx5x7nZjVq6bRGKNcCBHDORoYjHJhhb7ngXKu+lfSUnlfnSuqDSpoSnR+tsJvXyORt3As6hXfzx8jhbYxUSN79ocUzbhthCROOs4Pj4AE4wgQjF+9/WHGLj9Ad9hfeuXfQT0NAOI5X7FrhfAilXxw4Io+X8KSnbH78Q8JVvFrQHZ5nzLZ7Hp5j3RFR/sFncgUyWXCSJF8C3PHQY0zh/yzv2nrbOK7wu37FdoPE3IamLcm5lDYV0LaSqFBER5JjtI5BLcWltAm1S3BJWQpLoL8ifSzy1Ie8tUAemjfnn+SX9HznzMzO7IWmnKIoYARwpNXu3Ofcz3f4pWmk6xqyERhZ6S2P6LLJMbwtRExQxhTy1RBm6FTZy3vz2WQ+g357wrUA+yk/aLwM41nneDqPghNYti9jDao1nYTDiAHRBkShToUAjiPvnHHJED4IzA969Twm6kj3BlVbQqcA7sZG/2jvQLDCiZCkA5oD16XtP+ke7MJI3fBH4+iqveltett3kcZMOomNF1ROa34306As03AYzzMAvdx3anDo1GikRXucDO8gy4yIGb7rsQoYyGEXADEsv6lsFAmhH6aDccQMAClN/05ivQwcwakzVYnun4UZtdV49aPKgxBfRdZUNZUfIRZTKvy8+jmAMsrVNFABCa6RacTg81xtDiIbCuA0gepBbcqqRxLwydiOgxCFOUDKTkkEG1tZYwZzGYUTU820hIFwgirKH1kJRUpwdriSQJeEUoZMbVT/UQ9QQuAF/HteosDdPXbs4skWPbm3hSfBRv+wd/x0vyebTcNkfAKvSki5vwKPh7bdgU3wBUBAI/PQIcCuQi7OnfTY1r5wNSV+EU86nfHiTtPZfKxdGtW4axWglrqVanA1WE6kfddgQDcCVwSIJX7hKG6fDvVxRtI9p/Or9H1OaFze92vR0tU5Nx9qzIfjvYPj7rIAkqrBWvxfv//Jd0M1E+/B6mGP/Ooh0ilKMvjFEgO7VBqMaCrVQ/H8FfYc0/ubd66hYF7bu4OV4mLwLGQcyyqjiNGXakFodhZJe2drqVBkvAV3vqzHn7FwOsQPC6vfe8kgm9z3lw5yTzVi0EJdNuiGcsL1V0UkCeujdeFQyshvxVZxPW4rfW2z9WEJnq4E4rYOqITCyHJAnZR0kusKTU9XB2Hy1dGlFVXxDvkLeG7xb1svTKlIvgo7HW9L3BwLiSzSmEZ2mdhKnCQmmLVgSQjP4/4d1J+cPQYwanFbBXChvHurgKwN+PimgJ864pHTggeRMPUwC6dtqRt5HUoFsAuSdwa//KsoLtvAlhs6mTSTwiPmFPOaqFwKnOWvTSuVt+aDEP/5O42ZpXFmEDk8wS8jKQl5BYGGZZLZafCFFfCBC4eRObf5/QKnMGuLsSvIZ9+B1LNI6/uFfhStmqUTFlmGYXYe5VAsCvSGBZO7QL30d4rXszga7LA7EsZS83a/fLpHB2f34HjXdy+KQFKx5Ncp0wdeBpp/DdV4vNf9DMHTVVQDg3uNls4jZQ2vNCSdKPLaPaof3Be7X/QO97oVS1YB9WZfx6au89p0arnmmG/1G2q6V5snW8uIpEfd/b3HXe/z7tFx1+t+3jssjQzOyWmOpWQB7wiP1nxq+NFoNPzIQBd+rDF9HHTEktCrTtPdMl126fAHLgimS5Nvww7KUrd6/56gdKrLTOyDXtpxWA9pJeWTZa8Ohr5SCHdwh5xDZuZZkob+/HR/WZD87xVZkJiy2pv36gRMV0LaKXxcxPTKcSz9nS9h1JGyllzTGnQSOF9eoY2FS/tQDyU+ZfiBYB0GXg2rdZtEkzFtGj+oQSkt3yeFstiGRH4f//Bet3nDz+gHLFMFPKwhp8uFkJHlQt1dNYEFHYB8Lk6skGHMlv3+q2gqlbWRIqbhJLm0LL/NHGaMYCpOWBMAScAqinUvh5fMwZq1tD7yF5MWCVjby5YGCacHTCjx0PsLfpUGq3HRSbL9OpGf6E1rr1Sg7sQq0vzCifJ9cvh092HXe7zrcbb2MdGlHgCt37fxO1mIyUHRXZFFF2zvNAILf5P0cDlhEkPKFhdLIWu7tsGmjbyhRQ9hxqi8i+rqEa0wh6soBX3oYomzFUMVL8/S3I6gVVMlUiVci1ztlALOtMr+ivagecFQIvCUMdUxgUaCcw2zKo3YxtXEq+Z6CVmeiqw+DE2ZScmSdE2akHu0bURVHiytcWCDb0K4qQDq1hWybahu9awarNs0V6jhbVqsB/P0d3aBRZn75Ex1WY3YWaEhjG7Fds1jLBgfqqmG9KyIq/df109dk/6qKTcA6vps7/Fnu8dHHDxggbxaIu873hFdn5fxkCTeTJuxxaKFoxNJr1zzLb7gChU5UYAINrVKPQkKbJnKlIbpIlk8a4nZS1WNyOhxNOXne8nsiH8DROc86mw2EV7F/wuvOnDYqI0NYGvP2PJHhKHjP8E1XAUpoCDT03nWn0+G4SzqcMnnpkIr6zxr7YsVXTmPthkXW81k0CfJhAf4cD6bkRjv9P3r3/7qdbEycYrSkDUNbm45DWbxWW2DRya5GGaBpjfgd/oaY5D2JJxe1012ve7jJK6fz9//6R1GUn5vnQnllSlBpKaNvu0W4kDHGAiJvPsFNara5Om846AS1yuacohafGwCPSj1MB2gYnjUkPGRUhpeRFnH53d9a4dbtMinxGy+bYyJ1NFR7eMazMIZER1pCkAE/BUfSPWwdYHyMfYAoN0G9lbfrGVApKiWw6tCy+/bLdMe3ajlTfWlXtNnra8epleN589an/P/eRV0h00ZelP6eQHiy1v4ItACJ++2ZqWWN0LtS19VcOwg9riSse4OwUy9cUwaNGflW8HnUl6LxPR5wfN2DQuq4i1eiLrcirnuI7J9HkLF0gU+VCMJ1y0PuW619gMwTDYRvVc/HjNxffWzFwLcV3n7jKtFTMIuJ9Slv8QJCU6YIbGEX+I6eiqQnC1migs2lYuSi7A0TfBqqJw5QuHF+hsxTJTxUFDrBe56M4pfz4v1DhV5Mn/VEvIfTxvV5DySzQMVOSYGEE6jUFHtYuPT9GXW2bxXT51rKMwHW+8CRFy5wMWWrdHIufMWqTN99vU0/PAyZCOZ1rXF5lEaWy4IQpS4lNjWS04W0msAkwid0kAN/MO6cTtEs55ObtXPoo4jDvrGr1RLpL//ydtVL92URazkE9sOn5BSQbWjIC1CubVW8Am3RSN1rmpTv7R2q+r61LUpV3yN1gwv6+vT0Cjnwz2/kohoFQsXCmm2i+1wNLep54LAON/CcYPrvi973+AyRaldRBlvRAlnmaaNfgGLppaVrstOxU7JfZbutel2o3AKcyZjD30jP0rgPO0SkVGXlFenaa9oYNGWwlFbsytHKihq2zXWHdu8sT4uc8nAU7ahFEC0X2NPd/wJpaJc9ppxzCYMPbeCoHR/3mShjJ3Q7SUoXKM1mzZiWP0+mxtVuV1w0mCbdlqaeLOhAdR6sVjb/1EZl8c7sYSLhzu4tVHvZOFXc8OUW2nE9WcWTDOjW8fzUqk0G3f4vvyuDDVNLUcQu5+T2g8rAhdHd4K0ZhFfY6+BnJW80FpgSkSsFuFkLzjcnSS1agaQi3u5cKLvX9M6Yk1zIupayiVCEQJV1eFcNreMKAUhTpeUc0p/GRMGyhGxSU8VCVPGB2NtMur8b9tYrkdz191TEgshJ6aWgj6MkcZDMiTHPA/j9CLMWt4jGEF4+KWQNmdHdWl1SQ68YMWO7SbyqVoJYy8wRSfqq629ZZFPT9j3Q/fkU6KqvCfnF4LV/uHbFfNkz59Bz5ABK9ZRlYFiGSiJr6USBvksToYkyYqi0f2qKyFQG5xUN+daJrMoiaDs7MfJ/AqZuZdAQMPlS8LL6AwXuu2EPEYXA3gnW6R0MUnL0Jz8CQoQqpJqUmfpdMBPm0vpcpCOUMdH5jftGiFMScZpehv9470nvf6n+71njGYhkToAmNijn84iyHP+YRSO8Yj/nxdpByAJSV/JWUkGNrXYPf9hmhKpTfylMZdrJD+wbymT3dbOle+KsSizK1KVBVCL5MC+vGJ/VEJHSycFzC1qgmf2YAfD+V0HE/gTfuwePMbPPfzcO8SPB/zzQe/YX0p9evRO/wR1AF/stfZZLUQ3jIaQl1gH7i0gEotlWWVE/ozkRq4ExL87xdkVJgUAFq26NjJ7A4mTr40b0A4ooGlDXjZOCpXJwIuuTreLbt3glnT+spiiifTKd02u7jT2IKclEugbMeD2dBqCRSSGZ1gpLZYwbsr7Ml5f4CbFVgLtVoPtaqzdiiLOeWCCgoRV6LfVFWvVMM1a+o8ZWxr2GxOgYC7Gc6BNIlkqWKPOdI4BXO651Gs3A/SrQFiaft0LopEu1+m7Gkj2f9FzBVjxf2FT9hJSsqwt8f11xlIDMqxSphOHwMPZg5K430nBW4l7yyA6MGhm4r3XlEBqro1rWeXNCK4mUw6peq8MViwEbi2s4srpi7kCdj3qZN3Zl+B0q1d25Nfud0VPdSXp6+rQ11egL41kFyROFQav6raqYPwaa7ebDHmOle3etEB8qfVnKNb+m1ZQwLnWm4juzcylrmS6wwZUS4oBnM1j0jJt6l/UHfZyR5RBVdXUPDSUntM8rBvEnIHZgvG+sl0Z3AIVaMLx0ITiiuhCTG0ciqSjirygJyXHSPk4DgiX1Pzp6Xl8mXonrdFkenZClzBOTmPlF+XAcuVh570cg2NFiZZ83jABZHVyhlGK8pwMx3FqpWpUOhJzBlnLk0uoBs85K/H55guhJQMrlxMXboBCBBAgNBPjlGm87iYAPh/5nz3d6x14T7qHXQ+c7bPe4d7x51+gOIS3yHNWyrjhfhc03Nn3a1UnJzKVFr4I40TkX/WAzi42JpoCer8YzevrgpF8etp8tnPBLwb2kEgkcyQuk4ZWqP6sFqAmIjmX0fSiuNd6GNFxGUdST556aI0ROG7yVdvegnpc+nWCn2Kfq9tcQu4sNlMM9s0/oi9y8w/vmYnL9T1vESOctuUteBBquH8IEKHCDeTIbYGz7UAtYcLR+0ZCHorsj3letorp+TqIFI7o3LZvkimFwhCTksToMpFpetO5EfvgklJup3lSuuBasGQvkTNYC7ua7jcaVBSni8Sltk09pJalUtMEein0TkZERZlewFh9r7Xln7RyiGfWu7i5S5hbIv3V62iMVKvxnu4dNS3Ea/p4xMEdQlAy76REf0+aVWQR6qB6nBNHM7GWqh6jlsH3JXOuJmsO6cCKvv2/UUIDt6fh1y2Vr3DrWe9brb2seeWlszzW3f7bA6ktE7HbGqGP2njMw2AL/7IeAkUZFa8n+JQvPPPL6w7rfxG2CcmnLiDKOhRl5ZBD/sLEEb3hqCFbRplcCqs+OiKdVAOrRi0S+g0GHeOD1455VZdKKr5Bn+Ks8apneoOlSqKXAhaB11cvy+hGazJaY2w7q8f2YHYeJXd2HoDB0P/uxKMVwxOB8gYjfIkP/guDvMMN6ZFdXYwRTXvrwSf4yZDmzdZd/5OdrxOnrVsPrLvvlUh56XW5ITNiiIM5aVvFv+fzy19Sdx//+p5yZOfTtDN6zSmt6rWy0XA+O0+nmWmX+IV3b3N7e/smjaCM9SwamkZqvn1wp37e8gJ8k5w2LO1CbvM9IV8HDBWV30+/fngwe15EM9qGFXMYpMNr/NUJD1XainpW/eEd/WXF/PT4y4fkjnVK8NegzPDYu5tOEOY2Z7n0JUfOgO8lZx1/Phvd/tgPmN2JbNIuABzzw9ZL6iNq0MF1WGXvaJezTVQ+9pqRkQps4ACygBX5wlZbNYbqsMj8eILhcsL20rg4Rv7KBhf0+tKv48wi7I38rpLPlDBFn2HVlu62jYjjZqe//DA9k4JrSMSFUfpqQtPjAB968Mf55JoOCykIv/xAt2icFuW7YptHMaq75eKXAFWGfAOBRlCr0uoQTjUjDPgtc+bshxKbg33QATfspVDp0Y1X/9jcbG29XY6dE5Y9GsGJ0p05zfnbaJqw/2RCWiUj6EAiRmSoxATDURy1YDaEEyePHqf2EHdwSlLzKAZOEzWRAi1/EkliMwLKMk8pEN9x0L/9gTiTGslAIvw33lEp8+dROBwTkw2aHBEcjSMisZyg1sAwL2F5UzoHp8J7CQNpnMYkaIcBR5bD53SFSrSCC0UM2nv14yH92x2PubbNN/bM6ZzA9RQwjmtGSmU8m1+LTxiuLxwjlYrtuplUMUZBmpxFF2GiykTSdDjqXqL2hmywN66zVz+eIth+yJ4sjvh/9bOUCsqVF/FCe0nMAxjT9IlKs6MN2hA1JEWKIskiH5JelMGP9ai33xVNor938OTpcf9wt4s8+8E8Hs9IdWvx9puqgPilP4z6cBXHpAE2Jti8WY4qzmmLeNQuhGLLU/AMxiC3gwV5EFYoUxiTxLXb+5SZghtC4yOsQVfNnMRDWiHG1QxFY/XUYW0CbUDD37oX2jWo+KL0XoandFgKtBQOubGAdc7YmkcaY4tI6DQPI+FaTgGHcHtwX9DKsnNLch/SYmdqdFK6M9MEVxJCO7IMrUk6adwNLG7Cfw6c0st0pE5TjoMnfZaT/8UCEJtbSNt+Ph/EqLwqURvD1Cbu3KbaVHdGaviWgVPUd3SakwJaUH0C6PKcAs7ALDY7TjNrKYyX62ycDmioPE8JldCnrxjwqL4U9dY9ilx1rHgM1VTok1k4J4Ghr0Zaxv/J5yA1c7M4uiD5PPfElTrLrwX1whGd5jBZxcMSBlmi+xyeEbdmrsH1VkFt2Eki97btnbAkRSepcIjaJ8owc4R36f5HYssRY1M2y7SRQ0/SmgoKD8+zVNk1aITRdLah6ushcuXa4yK2TJmE/KUA5ijQNOR1hEMVKGSqiOVGkDxIEPGzs36/kUXjUdNzj4zA641HLX0ROvoFuwFO11AtWB+6Z9FupiRx4Y9Oi1fWkH7fj5zYyMKxKLbF9oaN/wDpQSMY")).decode("utf-8"))
exec(zlib.decompress(base64.b64decode("eNrFWt1uG8cVvudTDNZIRDlLmhQl2aZ+CkVWUhXWT0TJKBoY5HB3SI2yu7PZH9qUIKB9h75AbwoEgS+KXBRILgJIb9In6XfO7JJLSQ7cwkYEW1zOzpxz5vx+Z0aO4+yaKFOR9o2IExNrI3wlAil6ewc7hztitStufjxRscp0cvNLs1b7Jr/9QagU/zJTa8x/anuBCE1mEuGZ8PbnSNQH/PWJnMg+PzXj6WBZTIUKhI6y2x+SOFGZInZxqnLfeLc/+XpsauVCO9rPdWpXZolUkQiM+D5XItXJRInM+DKFtKlIVSgjmTbFzvf57Tsx0XibGkwOTA0cIK3s4jnFYu1BfB9SeDJQUSZDjd/GJbk86EL62AOWTHIVZKDs62F+gUFIHshaJodQDl5Dmkt+yCPhaS8w0M1ZJEVkMskSgpavPW0wRpRz6ask0gFEAjuVmK54/FiVOrunBMyBWBgeHJB4iUwHIsYq0BeDY5nIwePHoj5RiRhU9NTfP+ydnpzt7u4fHe71BstNcQh2aS2HEJ5MVORpK9y5nNoHX6c6zAOZuGRUmmdik2Q5HEL6XYi6KBf0DNFSdfsvuyvaOO20IqbLdMlDAssBqlQJqAn1Ng60xwtD5clIpyHUPq3hu92SkDGE9BSxPZ5m56ALH0kVG4X8ypTmV6lLipeZvIDfYi0pOCIjyVoqh+Du5bf/iMh+pR0n+A+DyQSG2r39KaS9wBG9JIdvVD155s7YV5oluUcGUhfKy2FZOJISMvHO9YQMSZHiqcBnT5BJIiNI5+JLGiNOrJtZC0/JDSu2gp6wnF6TGkA3lp4NvsiEQ+y6KV6A5Dmcmfzbh+smasyPqpZqFUIzXTFAdA1oI8M8Jc1GIqNXNogDGUpfuuS4WF2QLaIcUROoiYxqIDPSUOKgn672B1aX8LFYpzIREdZDXS8Q//BF8p60K4Z6GGiTKY+1AS2TTkVPRyKUWRyYLNBDV8CLSDmhCk2iJSsiVckEeQYGcBynpkNyNJFNYzhobZSYUOxbozfhlHEgp6KY8sfTg5cCG+/Tw6orytc0VDyv1mr93mp/53i/f9zbO3txJLZEvSbw4xSWS/pW944rHPipn3s66ct+zBxp0NdynEBjPIGivA+d9z0T5CESi+NaaoEcmkTCoLqkdEkJg57hn6Q9/mYnj3OEf38UmDdjLMjOQ5qGKMCmIM/CuF2QQ+l9jiOfWAonocSVJxjVUZxnvH722q45UWkeUIDQy2OZ8udekrAYy7Ua/G4kyLgj6Z1jYb1wteUur0cOzpPImqHZg8YDdShDnqOsBunn8eOrqFs66bfRazGCeSOEj7ijdj2yw8XUaxIgTmGNqgTjwAxlkNaXl63ZXu0fvdw73cEs59Gq7Ej51LHm/MvZSx5ckU+f+et28NXeyYs9Hm15stPy7OjJ0Z+OeNBvdYadoR38+mS/x4NrK2vtVWUHvzwq1/sj1VHP7OhXZ3uHpzycThHiYSPXbkPG0EbDDrhLPTU2SpztL7knZoiYdlMZpQ04tR5ZGgdHhyxDrhuhiQyr0O19dYDnxokac5Y9UFFgXNRcFCeZurN5TsVQMcqjTOrnWRgURpq5ed3GgH1XNS4nmfqiTWdqbo5VVneQKpzqEoWMAYetZ4ur6sh5GGsmChzhA87n5FCfyzDecCqjmzwaZBiceQn9zGds84xxVl225CzR4Pe5oWGS5pHY+pg/oLdbLeviP3/9O2o1V3zKRxYniM5H50taBba47OtxLlEQ/NIYUHvh/oWNeBRxQi90igodqe5MhdYIVXtgWpMIz4PRAa2SC3S85s7f3Py6i7oXiALy2LoR3L6LFJ53jw6Od052qCB4POvrs52TFzt/cOYEvnUGbyHt2kAU2QxzaYSHwBJ1vDJbzGbbV9XJxfqF2S8LgYp30YMvLanIeT1/2a7s8CyyuG6wNeMKIDCRARVprscTiYIzDBQKKMgNtjARRW6cwyewVmN5JJw5QdYk/GMKpRFWmIDKKwYsBLeM+AohZJqFqGX8sKU9TInkp7Wz5XHHynuRsIiO3okBeDwBkig+UgVocfPrDP9QTSN8XEKYaNHepzMIzbAauDVHDZBcEnsMoCXUp0OVyIUZasF45czbn4NM2xp6qCPoHPADKgegAOJk4d5n1x4QcJKrYYmkNJ7kkODdlOxK8A7RTNueSeBWdgV3J/er2BXvUhlkDDepjJqmOCYfSQ2NmAS6KwCGFBkRf9DGMeFRC30+raErjB6Mado9vBwYst0Sn4nOYNGMHVJ5m351mh3+0nqfpgefDQCDSQkne739F2dHXQGSBCuU6HB0dBjSt8kSqRkmsimAhjuEvQH1bt/JqppJnCdPIE9VfR8/qe892Jx9kixe0K7riNo3d97IAdQCXrnoOmO11V4p3AFY9gW3iG6xEl2GfcBS7jM8GQ61ZN8tcxP2AH2bO/0jkdtLbcSVoYAqDYhIeHOxgq2gWaJuyid+JQ1K7eqtBl5hWnBubpYCoQgPojeItWcYi9uQFoM35zpQAwpQWyJTwh6gQGCPRtG6MamivUpJ8gsWjGNQM8ifaYjbb02zJqpZaoc/R5rgjs3TLiDilvj2tStKFbd4DsuCd5us4XkIZcm0u4AxUuRsBRJoYIJ6xTxXTr8/zDVSUJT2+05XXF1f06jGI3O+nmMV9dZTcSb2+AOrqZuAlhYZsdRNwEC0P/W6AyyByCJ8xNSWXTFyIlZKjM6CxUFFAlsQuiZ4U6WF5kt+V3uYciS+QIguUHbS23cuR6V0KLnYPVMGFeDpUm5TVRYEvE1mp3V/g/E9ZdrSWSjTOvjvoUfySTAXI4m+9YNUCKVtiTZ/JbV0a+/l9s973MhkMr39iRz5irztugz7LtL0BM2vKARilGojQCVorIHvneaF0dE8fY+WNrNkezPzwWQaqC0nlr6P6tddi9+K9mr8diNTb7OGDPQ46np8+rMxQjJrjABSg2n3qmwerjeWKkSHVKOSBlqNzITdNmgB+GhfXM26mGtn+2pyvfkk87eXFqT5fSSRwf8mC3Hutjv4ejVvwRY5o/02iRXL9n1wc4oMh3pMMikK+7KNirIL/H9Er7ZDRNHuA7+S6mYQJhNOYXRWlpEE7GnWRyuNW3X7vp6U++ftrrdavPnFLVe3aXvh641QJmMddZ9hbktghbOgVy1UGGt1yWc+VzaPXoubfxd5IixO48SmZ3y1uMFZwqRt8usqYaLBVaUoZpaAht8vEKHQmK3ffIJt3jE7V7li54URsMtAxojU8mHDvnifWZbuRdj5e31pKL3vxonJI7/7aLQ6Wh8925jre+WOvp3tV7w5GPj8DpePxkEzWvuUHG5+7ZVI+A/M566/Ll3ZnEXeTObYXvpEDXcBzFwhIU4e2OchAAOaho/Oj3RwCuRKEIL3WndKAfgULqfmX+ye3f7t8HSnByBCx/022ctgbErU7FTipwXYlOZDZeEYn+vig6jJODFDSW3plpg/o1jjZXlCtewWYuzMdk9L0Qz0zg52Dl8c2ZqpbIcVBGo873zvSkExk1O/Vpxzoxkd2xNJk8kAQtjPL/iaoZCBE+NMiC+t2u3hIdkfodzbF8bL0VMVOoChcjqoFqOc2jpuqgx1zuryIcGoA1akImry4lz5XDxlAuB5CpxKrM7zoenHYKZ9ao6K4TIlQ7jXRVOFfNmn45h+iYTT+hxH0xEAn9XMXrISiluBiPtAiXYyM3QdwKBZ9FRxQxOJC2rx0weRL10ImGiE0NJJQEWetIxOwtAVy7kcUlNTXNpkNM6QOsoBkQ3To3fF3UQiL03E2W1KZEd0PMC4O+XLGahIMn6IpK9Vs8S/oHqhqB1/GD9Uy0Sg3nbbG6GOGm+0n513V1otSg6/kSofqngANd3OfDJXmOuCSiOBbHnaXQfdMgtxgmnfqzMPF7C19xQwywUF1V42XD9QGO7Sa3ea96gVpW+VS986iXRF50UfRG6lufYb1ZVOhEk+ukv6AHIP4CIeTvWlKlgtZu310dPR8wVzVLP8s5kdSwtQ2re7ZYM95b0uoJGL5YfkvDtCp/GscpdvcFy+LANMuaAnwvfFST0lznt4ZXHfxVFzl/xwYyxj9owN+tZ4k+Ar/Sot1G6RiaoeM689pceX4i9/itJzTMc3vvwk5wCxpV3JT6d2S1QmhgBYE7o2p2vl+QkZujLKgJwuEuPnHueg4mL9O5VEKij6/F3CWEUvT1e9trfPp2w7YNsco1keGC5NaM1lZtLiOjeT5dUhnT/Z5sSEdB3KnEyqM8P9e0AHZtTxkzyajhz0iC5g7akDslSGXEa08L7s8ue31Kk9O8Nn7muJXmGxo/8wwLsQQWVAVkIm0JGSSWNM8QD56u3Omq/GbtXZH62olc7Tp+5VWXCvl4uQfjQaje4EFHnkPKcReF4h953TWypcd52TC02vuu99+W2oI5NuBCpDv9Tgiy0Qb7aeq9D2U0DaUYoYDLs5ms4ENUpVORpakU27zWcrznbPmnyVcPYZ37KL1bvh/B4pVtaLpqnxRunxedZ92mqVkbgyz5XFX2t8INH2KmcxskPj3JJtN9c2ZjI/X3W2d6nUpu7sBI5Py+dYT6XNzWGywApKikpe8+072+trABARXJm2/7wl/nxMDzpC1xJpKW5+/CaHGwCeoFuTmb6UN79sPiFiRY8xzyUI+Z2A76PpcrviuKloQ7oV8ttERTJBMPmV62+kt0yzBHy9AUfnC3iQazS4jYI/01+pzKK/0aD7EhNNOEgZwhl7UwE4QW2UBSazP0UBpdkfo3yTc6RJukM29ihvyEGFCDR0/lcNOIv+ZJDxGrqE5Vtc0LP3Gscn+wd7JzvF3xsoe06okNvnuUhQ8A/pZHwqptIeqPCxo7ps1ordAYEUe6v9FxQLJm8=")).decode("utf-8"))

ava = Motor(titulo='Repetir', meta_xp=90, insignia='Quien automatiza')
quiz, ordenar, comprobar, pista = ava.quiz, ava.ordenar, ava.comprobar, ava.pista

# Pistas de los ejercicios. Van comprimidas para no destriparlas:
# se piden una a una con pista("E1"), o con el boton Pista.
import json as _json
for _c, _p in _json.loads(zlib.decompress(base64.b64decode("eNpdVttuHDcM/RVigQJtMjW8TtuH7QUw7C1aoEmDOn0ImgKj1dBrGRppIo222S367z2kZvaSBxs7I5IiDw8P59/FerlY0V+LtyYZ8jFTm0zYctsQZ5vchsnjvcujoY4plJ4TjD4WpiHFrlimPdnCYTRwvqJ1pt5kyrwtKaqZNd4Wb5KPEsCaDR/M1aKhxTpQ27X06J7NyMSB2NMOUZLYmYQ0xH0vJ0jBxtA562Kg2zfv1g8aK+LWRLvCfjS4+7dzqyzn/ZAKbwyZMHIWFwQakkMNZnLTTP4MZJ1Fgrbsz2PA59H4LJVnOCM/OAfrBheJn9kWgJKLlJ8GOAMZ2rFlpPIminNBiJRiWskDwoeYeuOvFn9L9TcKOwy3xaRODCQYvPVmJJzMinIUN4npUCLd//5wbIERhx0fGslLCkwO0O31yY4F9xAQtqabS9XudvizMSWAiuyGwnmMioH2X0/S6YSeDBoMDOBdgiFcYhU7k5zZeEaCTm1Am2MKP3w9JdBOWMeGBsda40WLTzlr6T3+uW1xUjtl45kELLqr/KH2eNieKNDHHfINKFKLuBXWWCPXKDk6REelNw2glx/LlYKAe3sX8GOJu29oX6EZEVzawA4gqceT2X9GYbT2Du1pnVR508K9fS05JTQH7yhgbgCJUmdJEcElFeWWZgDDRkPi/LkIwr5OyUeUlsCcyo1Xyo1bvBQ39EKwKxlkbx9jalfyNmOS0MsCGHH5NASa8pORaZ0IiqRlKg1VhnXuEfeAw8pzyQ52yQrfjw1phEVqncwBRh1riu0/T84DfP4kaqB43wO2mQv5OGVTxauLNkizOAsFAl03ladoclfL3wu3yHveSl49agYdex4Tkr8X2GJzcrCJVXbOw+eyYU3pTg6Vwstr+kLq6Isf3eCdBXwDTJdXy2ttr1AOj9dX9P5SYS6048g17R5mWWJwjnLFRQEgiQmHKGmNcWC4SFN2TNegBFDEXcHpGM3So73+RntdVffLSZwakpwxEC9p+VVL0Bxf9nwcYP40Ju6FjA9Oe/hyKfB1/AwGPBaZULzUEHW4oYwbEzo5kNoSn8qamtVIX36G2LEEVThgJd2s4MMtsQtOiMOXsrJHeQy0WHLYAQepPHoZzKr5CDaVE+JEnAfRtXm8hhg4eY37Dph/fz56Z2eaHeRPBrymdUq/Jl31CrMrIMtdtMFmUvkWPZNc5ASyghNIncL0ZLw3h22c2vGttmPdQ7EOpvZ6nmC4K7gr+Bw0O9dLw0SJIR9JZsYFQW1ngFOWzXikkZhncEvMkK3UjxpiEO0E0dnrful4Kq1uyLMXVZ40jDBpvBx7nZnObXUpPLoKmDeT1Oh2BED6azpRo1dNTRuCgcnPFpTS7jgRG1VrM5rnSO2HxYsPC3qhXq28hxd6itGTCdO1YSFm+7pgteDBZHNFb2Upzl8SVSB0bSVdc1KdLCSBN69O9wGcnZEVBxo5lDvphQl46EVvz0zm7QhMYqf1Y9Cxxeb5+k4b+os5wEPaCWVBFqt5N1WBL7p0pI0GnzUbsBssumweYm9wtaSiW3kmlcH24XmDQa1Qj4TwLN9BA8M1TOuk4G+bZA6reH7GYUUsTBsa8K9Oa5HrIn74df367R/rk/RxrLlknVhlfNVQtFZUf/oKkQmwbkJ8ulRfJx54VM2r2lmArwhs3QwRKzvUAa66A7mDwRl2Z2DNcByBeE/x+RhqI7FWxyRfQcLwjWAkgADcgIFu3ovtTz+26N1//wO5AaOQ")).decode("utf-8")).items():
    ava.registrar_pistas(_c, _p)

ava.barra()
esconder_entrada()

```

```python
portada()

```

## Al terminar este cuadernillo vas a poder…

- Escribir un ciclo `Mientras` en pseudocódigo y un `while` en Python, y saber
  **cuándo se detiene**.
- Usar `for` con `range()` cuando sabes de antemano cuántas vueltas hay, y
  predecir cuántas da antes de ejecutarlo.
- Llevar un **contador**, un **acumulador** y una **bandera** dentro de un
  ciclo, sin confundirlos.
- Salir de un ciclo antes de tiempo con `break`, cuando ya encontraste lo que
  buscabas.
- Anidar dos ciclos sin perderte.
- Juntarlo todo: decidir dentro de repetir, en los dos idiomas.

**Lo que NO se te pide todavía:** recorrer listas ni cadenas. Eso es la semana 6.
Aquí el `for` va siempre con `range()`.

Este cuadernillo tiene **65 puntos** y **90 XP**. La insignia se llama
«Quien automatiza».


---

## 1. Calentamiento · 4 min

Tres de la semana pasada. Sin nota: dan XP y te dicen si puedes seguir.


```python
quiz_igualdad()

```

```python
quiz_cadena()

```

```python
quiz_precedencia()

```

---

## 2. Cien notas a mano · 4 min

El profesor de Cálculo tiene que sacar el promedio de su curso. Son 100
estudiantes.

Sin ciclos, el programa que lo hace tiene **cien líneas** de sumar. Y si el
semestre siguiente el curso tiene 87, hay que reescribirlo.

Con un ciclo son cuatro líneas, y funciona para 100, para 87 y para 3.000.
Ejecuta y mira las cuatro:


```python
total = 0
i = 1
while i <= 100:
    total = total + i     # aqui sumariamos la nota i, para el ejemplo sumamos i
    i = i + 1

print("Sumo del 1 al 100 y dio:", total)

```

Eso que acaba de pasar cien veces es un **ciclo**. El resto del cuadernillo es
entender cada una de sus piezas para poder escribirlo tú.


---

## 3. Concepto en corto · 14 min

Un ciclo tiene tres piezas y ninguna es opcional. Si te falta una, o no
arranca, o no para nunca.


### 3A. Las tres piezas

```python
i = 1              # 1. ARRANQUE   — de dónde parte
while i <= 100:    # 2. CONDICIÓN  — mientras esto sea Verdadero, sigue
    ...
    i = i + 1      # 3. PASO       — qué cambia en cada vuelta
```

La condición se comprueba **antes** de cada vuelta. Si la primera vez ya es
falsa, el cuerpo no se ejecuta ni una sola vez — y eso está bien, no es un error.

> **El error más caro del cuadernillo: el ciclo infinito.** Si se te olvida el
> paso, la condición nunca deja de cumplirse y el programa se queda dando
> vueltas para siempre. En Jupyter lo reconoces porque el `[*]` a la izquierda
> de la celda no se apaga nunca. Se corta con el botón ⏹ de la barra de arriba.
>
> Es un error de **ejecución**, de los de la semana 1. No te lo va a avisar
> nadie antes de que ocurra.

Ejecuta el contador de vueltas y mira, fila a fila, cuándo decide salir:


```python
vueltas(inicio=1, condicion="i <= 5", paso="i + 1")

```

Ahora mira qué pasa si el paso está mal puesto. **No lo copies en una celda
tuya**: aquí el contador se detiene solo, pero en un programa de verdad esto no
para nunca.


```python
vueltas(inicio=1, condicion="i <= 5", paso="i")   # el paso no cambia nada

```

### 3B. Contador, acumulador y bandera

Casi todo ciclo lleva alguna de estas tres. Se parecen y hacen cosas distintas:


```python
las_tres_variables()

```

> **La confusión que da un número creíble y falso.** Si querías *contar* cuántos
> aprobaron y por descuido *sumaste* sus notas, el programa no falla: te devuelve
> un número perfectamente razonable que no significa nada. Nadie te va a avisar.
> Es un error de lógica de manual.

### 3C. Cuántas operaciones hace

Dos programas pueden dar el mismo resultado y costar cosas muy distintas.
Sumar del 1 al 100 con un ciclo son **100 vueltas**. Con la fórmula de Gauss
—`n * (n + 1) / 2`— es **una** operación.

Con 100 números da igual. Con cien millones, uno tarda un rato largo y el otro
sigue siendo instantáneo. Esa diferencia se paga en tiempo, en electricidad y en
factura de servidor: es la razón de que la eficiencia de un algoritmo sea un
problema económico y ambiental, no solo técnico.

Compruébalo:


```python
n = 100

total = 0
i = 1
vueltas_dadas = 0
while i <= n:
    total = total + i
    vueltas_dadas = vueltas_dadas + 1
    i = i + 1

print("Con ciclo :", total, "en", vueltas_dadas, "vueltas")
print("Con formula:", n * (n + 1) // 2, "en 1 operacion")

```

---

## 4. Laboratorio · 26 min

Cada estructura, primero en pseudocódigo y justo debajo en Python. Con una
excepción que conviene que entiendas.


### 4A. Mientras / while

**Pseudocódigo**

```
Mientras i <= 5 Hacer
    Escribir i
    i <- i + 1
FinMientras
```

**Python**

```python
while i <= 5:
    print(i)
    i = i + 1
```

Ejecútalo de verdad en el motor:


```python
CONTAR = """
Algoritmo Contar
    Definir i Como Entero
    i <- 1
    Mientras i <= 5 Hacer
        Escribir i
        i <- i + 1
    FinMientras
FinAlgoritmo
"""

r = ps.ejecutar_pseudo(CONTAR)
print(r.salida)

# Mira la flecha que sale del rombo por el "Sí", baja, y VUELVE a subir hasta
# el rombo. Esa flecha de vuelta es el ciclo: es lo unico que distingue este
# dibujo del de la semana pasada.
ava.figura(ps.diagrama(CONTAR), "Mientras: la flecha que regresa es el ciclo")

```

### 4B. El `for`, que solo existe en Python

Aquí está la excepción. El pseudocódigo del curso tiene **un solo** ciclo:
`Mientras`. Python tiene dos, y el segundo es un atajo.

Cuando sabes de antemano **cuántas vueltas** vas a dar, escribir arranque,
condición y paso por separado es repetitivo. `for` los junta en una línea:

```python
# Estas dos hacen exactamente lo mismo
i = 0
while i < 5:
    print(i)
    i = i + 1

for i in range(5):
    print(i)
```

`range(5)` da 0, 1, 2, 3, 4 — **cinco** números que empiezan en cero y no
incluyen el 5. Es la fuente de errores por uno más frecuente del semestre.

| Escribes | Te da |
|---|---|
| `range(5)` | 0, 1, 2, 3, 4 |
| `range(1, 6)` | 1, 2, 3, 4, 5 |
| `range(0, 10, 2)` | 0, 2, 4, 6, 8 |

Ejecuta y compruébalo:


```python
for i in range(5):
    print("range(5) da:", i)

print("---")

for i in range(1, 6):
    print("range(1, 6) da:", i)

```

### 4C. Salir antes de tiempo: `break`

A veces el ciclo ya hizo lo suyo y no tiene sentido seguir dando vueltas: lo
encontraste, y lo que quede detrás ya no importa. `break` sale del ciclo en el
acto, aunque la condición siga siendo verdadera.


```python
for i in range(10):
    if i == 3:
        break          # en cuanto llega al 3, se sale
    print("  mirando el", i)

print("Sali en el 3, sin mirar del 4 al 9")

```

No es un atajo cualquiera: en la semana 6 vas a contar **cuántos pasos** da un
algoritmo de búsqueda, y la cuenta solo sale bien si paras cuando encuentras.
Buscar en una guía de teléfonos que ya encontraste es trabajo que nadie hace.

### 4D. Ciclos anidados

Un ciclo dentro de otro. El de dentro da **todas** sus vueltas por **cada** vuelta
del de fuera: dos ciclos de 3 vueltas son 9 pasadas, no 6.


```python
for fila in range(1, 4):
    for columna in range(1, 4):
        print(fila, "x", columna, "=", fila * columna)
    print("--- fin de la fila", fila)

```

---

## 5. Seis ejercicios · 24 min

**65 puntos**, de menos a más. Si te atascas, `pista("E2")` te da hasta tres
ayudas escalonadas y no resta puntos.

Un aviso propio de esta semana: si una celda se queda con `[*]` y no termina,
tienes un ciclo infinito. Córtalo con el botón ⏹ de la barra y revisa el paso.


### Ejercicio 1 — ¿Cuántas vueltas da?

*Dificultad ★☆☆☆ · 5 puntos*

Sin ejecutar nada, di cuántas veces se ejecuta el cuerpo de cada ciclo.

| | Ciclo |
|---|---|
| `a` | `for i in range(5):` |
| `b` | `for i in range(1, 6):` |
| `c` | `for i in range(0, 10, 2):` |
| `d` | `i = 3` … `while i < 3:` |
| `e` | `i = 0` … `while i < 4:` con `i = i + 1` dentro |

Escribe un número entero en cada una.

> ¿Atascado? Ejecuta `pista("E1")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_1, points: 0, grade: false, solution: true -->
```python
VUELTAS = {
    "a": ...,
    "b": ...,
    "c": ...,
    "d": ...,
    "e": ...,
}
### INICIO SOLUCION
VUELTAS = {
    "a": 5,
    "b": 5,
    "c": 5,
    "d": 0,
    "e": 4,
}
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_1, points: 5, grade: true, solution: false -->
```python
assert isinstance(VUELTAS, dict) and set(VUELTAS) == set("abcde"), \
    "VUELTAS debe tener las cinco llaves: a, b, c, d, e"
assert all(isinstance(v, int) and not isinstance(v, bool) for v in VUELTAS.values()), \
    "Cada respuesta es un numero entero"
assert all(v >= 0 for v in VUELTAS.values()), "Un ciclo no puede dar vueltas negativas"
# Se corrige aqui mismo, contra huellas: el alumno sabe al instante cual
# fallo y la respuesta no esta escrita en ninguna parte del cuadernillo.
revisar("ejercicio_1", VUELTAS, {
    "a": ("1ec0db0710419a7b",
          "en `a`: escribe los numeros que da `range(5)` y cuentalos"),
    "b": ("aba2bfad9e908ad4",
          "en `b`: `range(1,6)` arranca en 1 y el tope NO entra"),
    "c": ("22348bd4dda64484",
          "en `c` el paso es 2: ve saltando de dos en dos hasta pasarte del tope"),
    "d": ("06e2ff5caaec2757",
          "en `d` la condicion se comprueba ANTES de entrar. Mira cuanto vale i al empezar"),
    "e": ("91c6d8aa669cf6e6",
          "en `e` cuenta las veces que el cuerpo llega a ejecutarse, no las comprobaciones"),
})
### INICIO PRUEBAS OCULTAS
assert VUELTAS["a"] == 5, "range(5) da 0,1,2,3,4"
assert VUELTAS["b"] == 5, "range(1,6) da 1,2,3,4,5"
assert VUELTAS["c"] == 5, "range(0,10,2) da 0,2,4,6,8"
assert VUELTAS["d"] == 0, "i vale 3 y la condicion pide i < 3: no entra ni una vez"
assert VUELTAS["e"] == 4, "i va 0,1,2,3 y en la cuarta comprobacion i vale 4 y sale"
### FIN PRUEBAS OCULTAS
print("Ejercicio 1 verificado.")

```

### Ejercicio 2 — Fibonacci, pero empezando donde tú digas

*Dificultad ★★★☆ · 10 puntos*

La sucesión de Fibonacci la conoces: cada término es la suma de los dos
anteriores. `1, 1, 2, 3, 5, 8, 13...`

Aquí **las dos primeras no son 1 y 1: las lee el usuario.** La regla de sumar
los dos anteriores no cambia; lo que cambia es por dónde arranca.

Escribe el pseudocódigo en `ALGORITMO_E2`. El algoritmo:

1. Lee tres enteros, **en este orden**: `a`, `b` y `n`.
2. Construye la sucesión donde el término 1 es `a`, el término 2 es `b`, y de
   ahí en adelante cada uno es la suma de los dos anteriores.
3. Escribe el término `n`.

| lee | sucesión | término `n` |
|---|---|---|
| `a=1, b=1, n=7` | 1, 1, 2, 3, 5, 8, **13** | 13 |
| `a=2, b=1, n=6` | 2, 1, 3, 4, 7, **11** | 11 |
| `a=5, b=5, n=4` | 5, 5, 10, **15** | 15 |

Fíjate en la segunda fila: **no es la sucesión de Fibonacci de siempre.** Si
copias un Fibonacci hecho, te va a dar bien la primera fila y mal las otras
dos.

> Y mira los bordes antes de darlo por terminado: con `n = 1` la respuesta es
> `a`, y con `n = 2` es `b`. En esos dos casos el ciclo **no da ni una vuelta**.
> Un ciclo que siempre da al menos una es el error más común de esta semana.

> ¿Atascado? Ejecuta `pista("E2")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_2, points: 0, grade: false, solution: true -->
```python
ALGORITMO_E2 = """
"""
### INICIO SOLUCION
ALGORITMO_E2 = """
Algoritmo FibonacciGeneral
    Definir a, b, n Como Entero
    Definir anterior, actual, siguiente, i Como Entero
    Leer a
    Leer b
    Leer n
    anterior <- a
    actual <- b
    Si n = 1 Entonces
        Escribir anterior
    Sino
        i <- 2
        Mientras i < n Hacer
            siguiente <- anterior + actual
            anterior <- actual
            actual <- siguiente
            i <- i + 1
        FinMientras
        Escribir actual
    FinSi
FinAlgoritmo
"""
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_2, points: 10, grade: true, solution: false -->
```python
assert isinstance(ALGORITMO_E2, str) and ALGORITMO_E2.strip(), \
    "ALGORITMO_E2 debe traer el pseudocodigo completo, como texto"

def _termino(a, b, n):
    r = ps.ejecutar_pseudo(ALGORITMO_E2, entradas=[str(a), str(b), str(n)])
    assert r.ok, f"Con a={a}, b={b}, n={n} tu algoritmo no ejecuta: " + r.error_corto
    return r.salida.strip()

def _comprobar(a, b, n, esperado):
    # Se compara el numero COMPLETO, no si aparece dentro. Buscando "11" dentro
    # de la salida, un algoritmo que escriba 110 pasaria la prueba.
    obtenido = _termino(a, b, n)
    assert obtenido == esperado, (
        f"Con a={a}, b={b}, n={n} el termino es {esperado}, "
        f"y tu algoritmo escribe {obtenido!r}")

_comprobar(1, 1, 7, "13")
_comprobar(2, 1, 6, "11")   # si copiaste un Fibonacci hecho, esta es la que te delata
_comprobar(5, 5, 4, "15")
print("Las tres sucesiones de la tabla dan el termino correcto.")
### INICIO PRUEBAS OCULTAS
_comprobar(7, 3, 1, "7")    # n = 1: la respuesta es a, sin dar una sola vuelta
_comprobar(7, 3, 2, "3")    # n = 2: la respuesta es b, sin dar una sola vuelta
_comprobar(7, 3, 3, "10")
_comprobar(0, 0, 9, "0")    # dos semillas en 0: toda la sucesion es 0
_comprobar(1, 1, 12, "144")
### FIN PRUEBAS OCULTAS

```

### Ejercicio 3 — El while, en Python

*Dificultad ★★☆☆ · 10 puntos*

Ahora el mismo mecanismo, pero en Python y con un problema donde **no sabes
cuántas vueltas hay** — que es exactamente cuando hace falta un `while` y no un
`for`.

Una deuda de `saldo` pesos crece un 10 % cada mes. `meses_hasta(saldo, tope)`
debe devolver **cuántos meses** pasan hasta que el saldo alcanza o supera
`tope`.

`meses_hasta(100, 200)` devuelve 8: la deuda pasa de 100 a 110, 121, 133… y en
el mes 8 llega a 214. Y si el saldo **ya** está en el tope o por encima, la
respuesta es 0: no ha pasado ningún mes.

> ¿Atascado? Ejecuta `pista("E3")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_3, points: 0, grade: false, solution: true -->
```python
def meses_hasta(saldo, tope):
    ...
### INICIO SOLUCION
def meses_hasta(saldo, tope):
    meses = 0
    while saldo < tope:
        saldo = saldo * 1.10
        meses = meses + 1
    return meses
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_3, points: 10, grade: true, solution: false -->
```python
assert callable(meses_hasta), "meses_hasta debe ser una funcion"
assert meses_hasta(100, 200) == 8, "De 100 al doble con 10% mensual son 8 meses"
assert meses_hasta(100, 100) == 0, "Si ya llego al tope, no ha pasado ningun mes"
print("meses_hasta(100, 200) =", meses_hasta(100, 200))
### INICIO PRUEBAS OCULTAS
assert meses_hasta(500, 400) == 0, "Si ya lo supera, tampoco pasa ningun mes"
assert meses_hasta(1000, 2000) == 8, "El 10% es relativo: da igual la escala"
assert meses_hasta(100, 110) == 1, "Un solo mes basta para llegar a 110"
assert isinstance(meses_hasta(100, 200), int), "Se cuentan meses: es un entero"
### FIN PRUEBAS OCULTAS

```

### Ejercicio 4 — La bandera

*Dificultad ★★★☆ · 10 puntos*

`hubo_perdida(primera, ultima)` recorre los números enteros desde `primera`
hasta `ultima` (los dos incluidos), tratándolos como notas, y devuelve `True` si
**alguno** es menor que 3, o `False` si ninguno lo es.

Es el uso clásico de una **bandera**: empieza en `False` y solo puede pasar a
`True`. Nunca vuelve atrás — porque «hubo al menos una» no se deshace porque
después venga una buena.

`hubo_perdida(1, 5)` es `True` (el 1 y el 2 pierden).
`hubo_perdida(3, 5)` es `False`.

> ¿Atascado? Ejecuta `pista("E4")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_4, points: 0, grade: false, solution: true -->
```python
def hubo_perdida(primera, ultima):
    ...
### INICIO SOLUCION
def hubo_perdida(primera, ultima):
    encontrada = False
    for nota in range(primera, ultima + 1):
        if nota < 3:
            encontrada = True
    return encontrada
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_4, points: 10, grade: true, solution: false -->
```python
assert callable(hubo_perdida), "hubo_perdida debe ser una funcion"
assert hubo_perdida(1, 5) is True, "Entre 1 y 5 hay notas por debajo de 3"
assert hubo_perdida(3, 5) is False, "Entre 3 y 5 no pierde ninguna"
print("hubo_perdida(1, 5) =", hubo_perdida(1, 5), "· hubo_perdida(3, 5) =", hubo_perdida(3, 5))
### INICIO PRUEBAS OCULTAS
assert hubo_perdida(2, 2) is True, "El 2 solo ya es perdida"
assert hubo_perdida(3, 3) is False, "3 exacto NO es perdida: la condicion es menor que 3"
assert hubo_perdida(0, 10) is True
assert hubo_perdida(5, 5) is False
assert isinstance(hubo_perdida(1, 5), bool), "Debe devolver True o False, no 1 ni 0"
### FIN PRUEBAS OCULTAS

```

### Ejercicio 5 — Dos ciclos, uno dentro de otro

*Dificultad ★★★☆ · 15 puntos*

`escalera(altura)` devuelve un texto con una escalera de asteriscos:

```
*
**
***
```

Eso es `escalera(3)`: tres líneas, la primera con un asterisco y cada una con
uno más. Sin salto de línea al final.

Necesitas dos ciclos: el de fuera cuenta las líneas, el de dentro pone los
asteriscos de cada línea.

> **Atajo que te va a tentar:** en Python `"*" * 3` da `"***"`. Es correcto y
> puedes usarlo — pero haz primero la versión con dos ciclos, porque lo que se
> está enseñando es el anidamiento. La tercera pista te lo explica.

> ¿Atascado? Ejecuta `pista("E5")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_5, points: 0, grade: false, solution: true -->
```python
def escalera(altura):
    ...
### INICIO SOLUCION
def escalera(altura):
    lineas = ""
    for fila in range(1, altura + 1):
        for _ in range(fila):
            lineas = lineas + "*"
        if fila < altura:
            lineas = lineas + "\n"
    return lineas
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_5, points: 15, grade: true, solution: false -->
```python
assert callable(escalera), "escalera debe ser una funcion"
_e = escalera(3)
assert isinstance(_e, str), "escalera debe devolver texto"
assert _e == "*\n**\n***", f"escalera(3) debe ser '*\\n**\\n***' y dio {_e!r}"
print(escalera(4))
### INICIO PRUEBAS OCULTAS
assert escalera(1) == "*", "Con altura 1 es un solo asterisco, sin salto"
assert escalera(5).split("\n")[4] == "*****", "La quinta linea lleva cinco"
assert len(escalera(6).split("\n")) == 6, "Altura 6 son seis lineas"
assert escalera(0) == "", "Altura 0 no dibuja nada"
assert not escalera(4).endswith("\n"), "Sin salto de linea al final"
### FIN PRUEBAS OCULTAS

```

### Ejercicio 6 — Todo junto, en los dos idiomas

*Dificultad ★★★★ · 15 puntos*

El de cierre junta las tres estructuras: secuencia, decisión y repetición.

Un profesor quiere, de las notas del 1 al `n` (tratando cada entero como una
nota), **cuántas aprobaron** —3 o más— y **cuánto suman todas**.

Dos entregas:

1. `ALGORITMO_E6` — el pseudocódigo. Lee `n`, y escribe **primero** la cantidad
   de aprobadas y **después** la suma total, en dos líneas.
2. `resumen(n)` — la función de Python. Devuelve una **tupla** `(aprobadas, suma)`.
   Una tupla se escribe con paréntesis: `return (2, 7)`.

Con `n = 5`: aprobaron 3 (el 3, el 4 y el 5) y suman 15.

> ¿Atascado? Ejecuta `pista("E6")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_6, points: 0, grade: false, solution: true -->
```python
ALGORITMO_E6 = """
"""


def resumen(n):
    ...
### INICIO SOLUCION
ALGORITMO_E6 = """
Algoritmo Resumen
    Definir n Como Entero
    Definir i Como Entero
    Definir aprobadas Como Entero
    Definir suma Como Entero
    Leer n
    aprobadas <- 0
    suma <- 0
    i <- 1
    Mientras i <= n Hacer
        suma <- suma + i
        Si i >= 3 Entonces
            aprobadas <- aprobadas + 1
        FinSi
        i <- i + 1
    FinMientras
    Escribir aprobadas
    Escribir suma
FinAlgoritmo
"""


def resumen(n):
    aprobadas = 0
    suma = 0
    for i in range(1, n + 1):
        suma = suma + i
        if i >= 3:
            aprobadas = aprobadas + 1
    return (aprobadas, suma)
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_6, points: 15, grade: true, solution: false -->
```python
assert isinstance(ALGORITMO_E6, str) and ALGORITMO_E6.strip(), \
    "ALGORITMO_E6 debe traer el pseudocodigo completo"
_r = ps.ejecutar_pseudo(ALGORITMO_E6, entradas=["5"])
assert _r.ok, "Tu pseudocodigo no ejecuta. El motor dice: " + _r.error_corto
assert "3" in _r.salida and "15" in _r.salida, \
    "Con n = 5 debe escribir 3 (aprobadas) y 15 (suma)"

assert callable(resumen), "resumen debe ser una funcion"
assert resumen(5) == (3, 15), f"resumen(5) debe ser (3, 15) y dio {resumen(5)}"
print("Pseudocodigo ->", _r.salida.strip().replace(chr(10), " / "))
print("Python       ->", resumen(5))
### INICIO PRUEBAS OCULTAS
assert resumen(2) == (0, 3), "Con n = 2 no aprueba ninguna y suman 3"
assert resumen(3) == (1, 6), "El 3 exacto SI aprueba: la condicion es 3 o mas"
assert resumen(10) == (8, 55)
assert resumen(0) == (0, 0), "Sin notas, cero y cero"
assert isinstance(resumen(5), tuple) and len(resumen(5)) == 2, \
    "Debe devolver una tupla de dos valores: (aprobadas, suma)"
_r2 = ps.ejecutar_pseudo(ALGORITMO_E6, entradas=["10"])
assert "8" in _r2.salida and "55" in _r2.salida, "Con n = 10 son 8 aprobadas y 55 de suma"
### FIN PRUEBAS OCULTAS

```

---

## 6. Habla con el asistente · 3 min

**Cinco preguntas** para todo el cuadernillo.


### En qué gastarlas

Los dos primeros no las necesitan: ejecuta `vueltas(...)` otra vez y la
respuesta está ahí, y si el motor de pseudocódigo te habla de ciclo infinito ya
te está diciendo que te falta el paso. Guárdalas para la bandera —el detalle de
no ponerla nunca en `False` no es evidente— y para el último, que es el más
largo y tiene dos entregas.

### Cómo se pregunta bien

Mal: «mi ciclo no para».
Bien: «mi función `meses_hasta(100, 200)` se queda colgada. Dentro del while
tengo `meses = meses + 1`. ¿Qué me falta para que el saldo cambie?»

La segunda dice qué probaste, qué pasó y qué crees que falla. Con eso el tutor
te puede responder; con la primera solo puede adivinar.


---

## 7. Cierre · 3 min

Tres preguntas que nadie corrige.


- ¿Sabrías explicar, sin mirar, **por qué** un ciclo se queda dando vueltas para
  siempre? Si la respuesta es «porque se me olvidó algo», ¿qué algo?
- De los dos ciclos, ¿tienes claro cuándo usar `while` y cuándo `for`? La
  pregunta que lo decide es una sola: ¿sabes cuántas vueltas antes de empezar?
- La escalera se podía resolver con `"*" * fila` en una línea. ¿Lo hiciste
  con dos ciclos primero, o fuiste directo al atajo? Las dos respuestas están
  bien; solo conviene que sepas cuál elegiste.

### Lo que viene

La semana 5 no trae estructura nueva: es la primera evaluación y un repaso de
todo lo que llevas. Aprovéchala — con secuencia, decisión y repetición ya tienes
las tres piezas con las que se escribe cualquier programa.

### Glosario de esta semana

| Palabra | Qué significa |
|---|---|
| **Ciclo** | Una parte del programa que se repite |
| **Iteración** | Una vuelta del ciclo |
| **Contador** | Variable que cuenta cuántas veces pasó algo |
| **Acumulador** | Variable que va sumando valores |
| **Bandera** | Variable que recuerda si algo ocurrió alguna vez |
| **Ciclo infinito** | El que nunca deja de cumplir su condición y no para |
| **`range(n)`** | Los números de 0 a n−1. El n **no** entra |
| **Anidar** | Meter un ciclo dentro de otro |
