# Cuadernillo de Ejercicios — Evaluación Autónoma

Bienvenido al cuadernillo de práctica. Este notebook contiene 2 ejercicios interactivos con evaluación automática a través de **nbgrader** y telemetría en tiempo real.

```python
import os
from IPython.display import display, Markdown

# Obtener datos del estudiante e información del curso inyectados desde LTI
nombre_estudiante = os.environ.get('ALUMNO_NOMBRE', 'Estudiante de Prueba')
correo_estudiante = os.environ.get('ALUMNO_EMAIL', 'prueba@universidad.edu.co')
nombre_curso = os.environ.get('CURSO_NOMBRE', 'Algoritmos y Programación')
id_curso = os.environ.get('CURSO_ID', 'CURSO-101')

display(Markdown(f"""
> ### **Información del Estudiante (LTI)**
> - **Nombre Completo:** {nombre_estudiante}
> - **Correo:** {correo_estudiante}
>
> ### **Información del Curso**
> - **Curso:** {nombre_curso} (`{id_curso}`)
"""))
```

--- 
## Ejercicio 1: Suma de dos números

Crea una función llamada `sumar(a, b)` que retorne la suma de los parámetros `a` y `b`.

<!-- grade_id: ejercicio_1, points: 0, grade: false, solution: true -->
```python
def sumar(a, b):
    ### INICIO DE TU CÓDIGO ###
    return a + b
    ### FIN DE TU CÓDIGO ###
```

<!-- grade_id: test_ejercicio_1, points: 5, grade: true, solution: false -->
```python
assert sumar(2, 3) == 5, "Error: 2 + 3 debe ser 5"
assert sumar(-1, 1) == 0, "Error: -1 + 1 debe ser 0"
print("¡Ejercicio 1 verificado correctamente!")
```

--- 
## Ejercicio 2: Verificación de Número Par

Crea una función llamada `es_par(numero)` que devuelva `True` si el número es par y `False` si es impar.

<!-- grade_id: ejercicio_2, points: 0, grade: false, solution: true -->
```python
def es_par(numero):
    ### INICIO DE TU CÓDIGO ###
    return numero % 2 == 0
    ### FIN DE TU CÓDIGO ###
```

<!-- grade_id: test_ejercicio_2, points: 5, grade: true, solution: false -->
```python
assert es_par(4) == True, "Error: 4 debe ser par"
assert es_par(7) == False, "Error: 7 debe ser impar"
print("¡Ejercicio 2 verificado correctamente!")
```