# Consolidar
### Semana 5 · Unidad 5 · Primera evaluación y consolidación del entorno Python

Media vuelta. Llevas cuatro semanas y ya tienes **las tres estructuras** con las
que se escribe cualquier programa que exista. No es una frase motivadora: es
literal, y en la sección 2 vas a ver por qué.

Esta semana tiene evaluación. Este cuadernillo **no es el examen**: es la guía
con la que llegar a él, y el contenido nuevo de la segunda clase.

**Empieza ejecutando la celda de abajo.**


> ### Antes de empezar, dos cosas
>
> **1. Ejecuta las celdas en orden, de arriba abajo.** Una por una, con
> `Shift+Enter`. Cada celda usa lo que dejaron las de arriba, así que saltarte
> una hace que las siguientes fallen aunque estén bien escritas.
>
> **2. Al terminar, entrega.** Tu trabajo **no le llega a tu profesor** hasta que
> pulses **Guardar y entregar** — el botón está arriba y también al final del
> cuadernillo. Puedes entregar las veces que quieras: siempre cuenta la última.


### Qué mide este cuadernillo

Cada ejercicio calificable está asociado a una microcompetencia del programa. Estas son las de esta sesión:

| Microcompetencia | Qué significa | Ejercicios |
|---|---|---:|
| **mCC87** | Identifica variables, conceptos y aspectos relevantes de un problema para desarrollar algoritmos que permitan solucionarlo. | 4 |

No hace falta que hagas nada con esto: está aquí para que sepas qué se está midiendo y por qué estos ejercicios y no otros.


```python
# Motor del cuadernillo: barra de progreso, pistas y verificadores.
# No necesitas leer este codigo. Ejecutalo una vez (Shift+Enter) y sigue.
# Codigo fuente legible en el repositorio del curso:
#   notebook/cuadernillos/**/ava_motor.py
#   notebook/cuadernillos/**/pseudo_uis.py
#   notebook/cuadernillos/**/contenido.py
import base64, zlib
exec(zlib.decompress(base64.b64decode("eNrlfd1y20iW5j2fIhsOr0kXBcvyT7loSxOyre5yjWzXSK6q7dUoqCSQpCCDABsAaclcRuy8w75AzV1FR13VxUR0X2yE9Sb7JPudk5lAAiQlq6pmYidG3SWRQOIg8/x+5+SPPc97nRZpJuLLv4VRkIpQiTjNRTCVocqSKKYvoYrF7ve7fqv1T9PLn4TKxYX4C3168xZfWhtLP629vFBizIQnUVJIEUsRyIkk8pPs8segiALJdJ0X9UA0ClTeRZN0EKuxpFe30ixUicxwNcoLXFJ5IOM0kaFEy4HMMkM1HWUqT4USUZJHoySSuS9eSpGpIktlHI0V+hFEl78krSgZqzCS6NaFyKMEPcuLaJTiDRLE/jJVAh1TeTFFowQDyZWQRaZmUkgxlHEsM7Di7t19KZK0oF8il7GiXkiw5Wf/7l1h713QyAuFwaAblz9LbpkwlzGUQMWhHmQyGGXECNE+4Q/9KOzhubzoqzOVBVEQpf03J52u7l0OGfGnWIHYSTDNi3Tsn+Un4qJ1IidRX51P0qzwJxcnvjiMhIxHl3/DW6kDqhJNqHLIJFOB6oq6KMCWEcjnkUrCtAV+R0MIDDLxxR7rQzRQ6CZEZgWr9Qdf0ixDDxMi+B7U8IfELOPpGHy6EJc/xa3JVIEBmYIks2gQZbF8KtBfQeJLSAWDCPdJwjzIYArRKbD8W7RhvZtFM37HNBHjy1/CaUyUSQxTFRepviM1d1do53U/rd2yv3kK0hBtHKuRFCcf0uz9PYdNfjS5SAYnXeqIzILTaJYKMBp21KPunMoL8PyXBPKdpAmEi2bRmCUj9mIxUgkzkQQBJmK0rgTu3o2SIINkJcsGw4SKUksSYMsMG5qGwUIGkwwKnpkx++I16S2ogw6rpJiobBxB8ND7QuKytjH6BJWBakc5fw7ScVq+K0mzsYzpBTP1kftId5mzkkxSjKKR1JI5hBWBFx+icKSKnDqsNYxYSGNTPTxPik0mRBRmKsvJFomsYXJQTLX1tcZTjDHjcatkmgQRmAQBQ6ljVrHLX9APGDuYCnb5YjfHCBsKTM8xNagARjCIZQLvRloGIVj/wuxM0NEk9Vue57VatwQUQHwLIyW+C+hhUkShDKFy+RTMuKEm3RKH6SCLaNR45+TylzwqoBvEH1ko3Igt4TxN2PHC5YBPY/JtYB5dUeP0LMr91u7/+G6/K+h3/+3hC7EtvFtb8ssn4WOvi4/3N7e+ehh6re/3Dl7udQX/KdttDp/Ih5LbbcpHwYP7Xuvg7TdvQe31890DbhI8efDV1oCbDB59ubW56bX+dPDqsCvod//F/u4Bmj9/C6rc/NHWo/sPFTcfPh5+OfyKP4ZD9UA9QSdevd3fe7erWw4eDoKHYG3/xeEhXQGbn+XFRax2Wr6cyY1Ansn5gLx81rs/OSd1iEJxu82XOvlTfWsjVsOi99C5Lz9O4+o2dCua5r3Hk/OnLQFHHoZRMurd38ID9/HU07HMRlHSe4Lvm08HMng/ytJpEvZutwP48xSEhmlSbAzlOIovevkFjGy8MY26G3IyidWGvtAF6TuHapQq8d2rO92DdAAt7+YyyTdySHOoaeTRR9W7/9B/hNfGUaI2TlU0Oi169/1Hj54GUPasd+t+QP97uqhY4Kfv585QN3TD221YimZD1edbw62hHD6qPQ1LXfl4lp6lzafD4cPhg9rTHFpXPi/H8BLLBL5SwxoBRNl0dfejlGypSeHL4aNhaCkUUTFnzn3QjHr8aPNp+X5IOc2DKYtIC3EDXEe8I2VoMFCU5GrcK59vMGy5ueHWEmuWW1q+mKaMQzYwiDCdW3FD92qDVl8h0g5XKGyKXg7j9EPvNArhcOo0M4Xgk1RUNzdv33bJkobJbIOcK6J10f5qE562a82j60igQ5ax/HZ4WgCmIkqTHvx3cSr8B4/gdWSubEeAXeR8tXms1v0HIGsZNcqivBJdkU5cuRXAFKXe4YlYTnLVsx+e3uSdjpE/ZiN3XyGK0677LbzK4VjnQWTubxKH1HmxAQw0Snqk203K85XuZKUC17R8s+xjPhvNx/J8g/lvRGzkLadF+jSM8kksL3qDOA3e21FS1wTfNlQmkbqRlLbYQzXkBBVxhhsQ8spc4W1UfLZ2yIqu6ZJb7wFhIE6jV/fuihcW4QqZhOiWPAMUyAnEBBS8gdgs3AAkYkCHHOCji3ig/TPKH9AxjvmEDQzMDPgDsAC6MyV8D1gjCQsQagzVGUMazj1mbMTAzRlwHASEFxC9Els/tZhxCNwL5EHReOuR+Mfn1PPD7//EFBGPJYE3mTA29cXdey0fgCvuncq8vSP8dFpMpkX/Q4aYAbTHQkmDaQzjE7gdJbg7t7JMgAit5HSj5q1n90yYRMgUt8Wceiw81lKvp2Mxgi7pF75qdOCx9uGrG7U9doK4qLGBJkOeDpc0DPDYmxEVwgP0hPYZ9IwO5eZFWpHN6wheGGqOm7Xv4bvCI7Uy/em2FgSwXgMpG5UIIpUV6T2VZcBlLEnwmoEPf2Hkw+oyg3ggGY1GizRMKQPKQSyPSMdlrvF/gHEgZTMoOaI8RnHmAxUYchbByUgiZ1ADAG+/1X/14u2bt/23/whk0uax3HkGaxRsidve/Uee0IaoP88i9eF5er7tbYpNcf8x/u8JFtI2sYCS2thYzsZWZSeZtvfJubdzx7wjiLIAeUUAUk88EVzwn2zb+9Lf8qCEcbzt3c69e2X7iYRjDre91w/9x+KJ/yDe8rcE/YfvGw/9r+xDpDrUpSx9jz7dGg6H9tuGHRI1NmTNHQohSOK2PfZh5QN0+SyNEnv93g40cjbauQNlZBG3OpZ/r3f3/7Mw8JH/UOC/+BGYh/9e398ER+nChrnye7LSZRmZGjjWahXZRU+IW5QC6iwJyg6MicQVzyqEYunkUk+5QoGsmczA5FBsIVyI8PnlOqd0EzAYxA986+vdP/d/ePXyT3vvCHu/y5BdqfNATQrxih/aI9Pj3kwyORrLnjY+yEJs4HUTRU465ESM0yQm+gNIvQF3VrzhjzLOVatleiTzwn6EizyNo4H9GqX2k+ZYcKrsBdh1a5ilY/Hq24viNE184xbtML9+9xqezlwEO0M1FP3TYhy3KXClnR53y9xvU2tzg3gPruvIk3HlgqIThpsIyk20C8GXb+CTD4MsApdUvdoBRnxzSO4FsYK8D1Q1tDUPCh+AEpKz5Qtbnkg5Ltn4AerQk1iO4cDa1BLyFpRDcxHolq7p2H7BhVG1jUspL9++7nTFmBBeBuGadohm+E19pxKVzo7zSI0nmeoir78lCnKIlz8lAp4T8RLWJHngk8sfYVcIYfvoCMG7UBb4ROl6lMHZ9sTJ2XRygdjv5+k0C1RfQ9MT0OS+im/07X05wJswzKQAZfNMn4cZDfrJAIxXQJdpkus0X5d5SNVTkCLFx23044RjaOeEClt5aqJs6egrj80vz/H+rc2tB3DfJN7+3iGc0Mu9g2Z2eV1sZuvbGFM8uiJAExVuPmOMSqDk1xMdxqoEv05Lxm0GkYl1yaeD4laC6x751MeTVfSFnLt4FI8wyAsVVFcy9OfRGiBucR36EuKjcrJtRPgcZCYICwQMHZDyLIxmAvAjz7e9xtu9ndem3liv0cigiGYp+bBnUqRJAP19v+0Vp1HuB5A69KR9h9l/p+Mz5X1Ee79IR6NYte/URXKn8xQ56DRLyDEidfDYC4gd8mT0Wm707J5ER++hpyWoqgGUGppaGN9ioGrWZ9sLZdt4GBAwvqRRoCsN2Fgp/PQYPpsAKdJD8goqiWWgdLl5RlU6v8Uk96WY5rIiQBgJOWACQj3GpAE5AK7Ii/jyZyR+/JGURyaK68v4aoIAU1yBWwV1LCK3Av/yI25lcPJc6NVVKi4/ZVzihxfJI/gWJmWqcklXAK5OZK3siK8FxSPB+cJSZd+3DOO/2lnXTZd8M4tYsKo4HC508Y9cBVTUuEpiYo0+t/8W/IQzU5keu1OHo6/I3qlQ3KMgrpUDP2/S8SBTzf6amQCMZ3l+gSv6dopBD4wcaP98AtKwipL0f/+WgwkokxCoDAovGIVcpZ0i4FPkYBZW0yH0HWpPhVET2c171vXa7Yt9DfFvJMnjc9U60xkQ9bESgw6RFDj7URIV/X47V/Gwa3jUtSPaRiraLV+w7R3qvlnAAmXyjDHQD5HwDZe3Dan6TcuobfuC+u1yJNvlO+sN+NHN+rU+pMJl/xz3clW0O3zvFvmiGU1VgRUBoCXNF/X0JMEYygvpUHDhCNwgaOaZtsV8IaofQ1Bs7NiJKFBGHGef0CQxoGrzEiGHRMwJKekVt1l6PMv4waPjWg8svMNzWjFn8J+50dZMDTOaG8soM7LPGFsDROhooVOBmwZ/UWn0xq/6qVSIGJoZBeIBdsX5xNELKNwhmtBbaxiJsAC398WBonjCGEIX7jPK8n3rMdgUhoZ7gMNNwYMQVOMZdKN6q+B6AUeEtRrjyzBsM9XOkqJ9sS2a+tkvWWwk1K4eM6/CIxVjhgSpyWHr5swihy+ToCBLiJI2mxlnDPRR3C07ca9mN50OcaFmSQqxzhlfDJGmZSZmf4ZIychhm0RLI4E5p3CLZkkKsZqmPmQcyOQjNLsn5jXbXCDgg5bJqxzh2C7vbC/3sN4bb15rsGEfXZCGWMdb+oJPf2104NPfK6VY4v/SwBuoxKnTejurb5qCq83sKlqGfbpIN4f0FrdBg9GE/n3nmpfTiPHEYGduR/zs3mAHQq4xhNnw35JBPnn66d/0XzFnwS6ar+lUyrZSw2A/z5cCmAkTAF9TpC8fjUFygoBPjJdsWSxvWmCSFm7CV7c27Wu0pTRVv9NZYZjlpQ9Q2R98TtWufdz1kD7hbhjNh+q2zfs+OLxZttsGn4ZwIB8qx6Ib1Qf3wZ/JGJwz2r3cQ8e9Po9TMDnXU6bk16Lwpl626ruuDVJORfhyhYS5nlZhRi1XUxd1aq5TB6KesUbo4JXlqiZlLcUl3dXdsPp+x+EuzZI04EMwVdkkpe+TdNvz6v09UIS0MobKwKUXPFdeACDLkEpAiuZ7u6QhaXfF4o1aX6l/JJK2Z6drhCe+4Nd2fAwumjgeGp1D02WrxHVvZ667bi2MdN0gGfawS/ypk5lzRxaazGKuh79YZpWuLhtm5bMRrWlRSwx6baa/gXHtbDBhDYLpSP4+Eh5uU03arg8I0w4jTYgY5GoM0rraUKHOlWMx8xEYC34vVjk2ARaveAjvxkP47TIRXy0HOw4jeOrE8AEClwNFwwKSGUYg6LADtzir96kQ2R56z4rTnXmAN+Cvx3YbkN06NKrRaTE4j9dGAVLZDmmLSzw0xEOXOPWpQw1xIzM36Jq9l68IRyVnaaSqpm40dNgRSM3R6QUTrVSG25dKw+7kWw0jN37LT8n5TI0iNnuDTuu4TV+rq+OfpjLjVCVfsQLL5GJlxBBtntQhnAYfE/EcTWeFSroA+YhffgxBESpumz447ptm2RPduN5dXX9zIZVF3Mtv8YGc2+a5o+NOM6qZJ3nNyhUxzsCMsvpJP9w9jmDPpwU+ggU5Vw+jNNkeeiw+0d68NwessKNbdLxuw6icHwSwdFps/+Dv84d2WUzfnJyXdlSyR8b9yTTOgcH7nXpvk5ITyzzY7DQRXELYzeljb6mDjbDdTJkqQSYwl/u1dswk3+EM+WLLmzm3XzQ5tNxBtBLb13TSvCnKyZJCW/WuxfKoODW6s/z4Op/I0YVfugwcOYTcWSPP4R09SjNIspjaOLW3nBsWJseNuFENCTbAJbJ2JfEl/MstHd/xwimbcMlO19xu5E1KQ6Sp2WaeRwvJ1GiaUABPJ/AACS3gtNOzXbtoi24g0FkPsw1EV3My3xoixB0iQ1WG8eXfYoTymFEU1cl4EaK+Vy7hvEm443hEr+8tq/CSX6x7xM8GwUyMIVHFF+/ZAHHDRpn/+7/+VczThY4jKQURy7erQTIt2UjZzRzQJ+1r8nbK1pRvs+8sKXUFw9VtclRXOJor3Q2tLfBc5K1nO6gHb7nw7bDWSidb5we9720LKMGA7/dNQkVFRJldeNf0c51TfLDOKZZ9WvKLeiR+ECvIO22O5VoXAR1gYZiUIMo5HvRWdn8NnCYN8Xb24mikJ46sztu5oIqhfukOmqRV3OgIPKM1vNWdGcmE0KMNCrpwU5VsVj6jzgmLkrMW7S/mmgKlqB2P+GAoaogn2oCpgMzA8Jc/w2i8KziyxsGm72/oXbWHnZfz+IsXmgPpnDtuvavjhhZXMDS/Uopr+jyW8W/q9Ovd/cW7NJQzqt8nqe7e91MVM4iKlcrYAQJzjWlN6tVE4eQvf6JifCrSgqcXPupZPl1g05N8XF1L/OVAU2ldLdiUlzv1kKS4RlnePa7jsAqGOfitjtyWXHNvGTSoMs3nFo1KALtaCiyA7wNEUeN0ubLiLZcFfvC/f56et4/YbLpUdKCv5j0d269jN1cx+w9WBL4IGQeCHTfoG8tLbxDyzCJkSldp7iKHG+/Z14mJzGmiR6dFtAbArpanL4U8pyXBJbkT3ZcTmpzltDEoxFwV0V+mCjzlFV0w2gtxUu/rCU9PmVJ0Scw+l5u5Z34GlCc8Ce67Y/iPjLxVt8pEgcb8a0Kz95aZzDPRueb02uhTD96kY+9ZuXwxn5kg/h7RluK47pAfFWqcQ9pXRnN0hiSQaROqtfzBf5mlkzD9kJTB/agN2PC/oUsUaDrHAJJH6Ax6gl7otx69P17g/vuO7hEnxZZhx2sjqwMSRCNrAVPEPGLI2lsfmnUIn3vOw30OzF5PeF9SZF6sfXZNQH/02AT0mmPDoCIaVCaTkWoTfq6rstP++L8kYoE+jSIO7ke5wQPEtNyUNY2yHd8I5ZBisCZp2r8C4fxRxjTvrJ0ZL/4gUteBGhIvTeWZ93Y64g866bMXbtyT3yVIH6gJbIr2T031kPQoXtCkOn2lsjrN37uzXNeMtJSa8WkNpf6vC+P2eJFwGX5+VzR3S7xUAQLpi+8u/+XNu7eHcIPInCOV6KWn08sfY5XzutMx8Diu8woPqkj7K+nRoySN6bh9n40OOd+A7OZjNLEq24QJPJUnSe6Dzv8nkPMdjTQXcxoOlymuobXCCS+MpSuuRvOuAl9Qbn/5N2CWQmnmXkMYMtA7BAEnskwDHp61AqDlTXXS3KEVK5Qv/WcHsg1w6kADxPnPgai0HoO3862qzkzVQHI5IM2ghO+vgqQ24pk9ieTHnCX3Mga8Y1/HftzdIavlrTcfuqiUX25QKS3BTnSWS8KEwkQDNizuGsBpqGY66aGC3T1eSYoO81LLcg5nCkaJtrnnTeyGUL2Uiayio2eQjYfgftJK2TWw9ffApbSct4H0coLB7KE1B9o8RAfK6iW4e/yHaqGEtTO9DrfOWG0IMCoMJZjKmFZuUTqYy7WvtNwZeu+mjvgy6AiYzwyvxNwTzwDWFVxqlsGV8mc3bVK0rEvxHIrLtAgmSSlkoNrlq7ssHdVZw4yuQ6y86NIsL9YJ3CTY/W5BrkoVgLLX1hY8sieaMvAOebs0bePEh4i3K9OGSIRBR+hL8aicoLKcIVIHvBkXvHSEeiFmthygzlQwLWS1fNn3Vq7EWDP3sQJRmF58QdOotFm8oN3uMg8otTVvM0qiHd8/e7fzf/Y6Rll8WkXZWEKzjoONqONVE8Uem27l0JbmehqpM3vlktccIMzaqJ7gFcWMyDKz9jLjpfO58WiGa45C/9ZZo7VpJpcl3uiN4OYl1qvQ3Bn3cM7UF753Zbr4mTM5187iOJ37s94WpWiZNTpTze+tT4X5WAc488FUVUvbeRS27pLFbAaS9+WQaJC20ubzVHjriE4ufw4R8Tgo4AFk2IaYojSbIZm5pddtzSTblp6KuS7HvnZuytXTq+dpyrB5lBw773fmWuCaM/Xb5mmdiVqoiI3nchhlY2knV8ak4GfL6wZ2p0WqKPMzcyS0/IS75COREyEvwEsiXmzRM9iaijpnqXNCw40W/DjKZBahaFHBs1LcgkKd0tLLzyyt8LyINCUVnlx3h311NQXOKopjvWTyB//FqQreDwCa3HRe0kJW2s6/rePjr5wUoW2aelFP0axRuLWKpf5/XlHiyqlkgmdiHGnNgDb+XrPE19cWeBV/qF0kkpzAqS/wCgnL/kbNRuf+EElQm/hdeswqmqF7sxoFPajfdPOiwG9PUxHOUsMek0XNjXn+msQ0NtmGsyQljnbmBWjhr7aMgped8IA7Nx7wr5m/NmO2OlA5xlLqZqj72vq5b+S1dS36arL2JI4utQbmAbxPyoXwyL5l79k03pkzY/AafL6K3p29MVdDOUORp5c/2xX4VTxKywNTLn85j8ZSryHzf9WceyNvK9UZWRs/vTpXg0cuHTsv66QF3Lw5dcmjv5PZGa3GHUaJjHvVmtg84rX9QPPl2n7ahaWp6GtmG4ELSB3PvkbRYlfRYlI03cEawFiz0Leu0Sq32YipYHze2uP1ayhe2bGnA15QubwwWbTnzkLizpL5LZvdNZ2kLUELb2efNqDnoiJO+t9csOuLK5Z/UA6Xi5leXCudPK5cNtXjXW5XTdGtwisGcBDEmjubLxYrAy5E3LCkuR7/YhUy+6LJE+cIgCe0D3eula221k+rn13ut6oTHhXRPL0Rc4OXp1BC1VjSwXDebtSP3RLElUfttHhHbXm0j64xMGBUPTZ6qkuVK5tJit8e7L3ce/HqQLT1mm+q3IJKGAUUs2UW6e2gBtUEdGAVf5rRaV98YJM6B1ZlQ+top0LZPsHqiFb98R5HncGXFRrG2eWIyPTTjM/0guAT2g1av886Shu7Et55iRhq5uVqm8yqA7Z0tRJ4jg9m0icj8YJv2nw5Y2+E9IaOP7hFRj6WRcqVl4BeIT79NSemUQLKW0btwV/Zp7/7lsFVVgqKYXRGCTVxfajylOBJxpvcYndPWEy7NR9vbH61sbXVw0vQz7HSjKbcDV/4oCa4s4HZlyZ939dZk6JpUbPjNdftqG/w7V1+PgHB/5Or7PJHs33WEJNZqmu6Iq1dQqL5DxiN+LPJ/IGmPkJ8PYeHfPYZqEnDRtbHXA4UVSLBRCar+bEvLZyD5EjGyIUy8fV3e/v7u4d8okUpJy2BiKIwXpWfyq1Hj3WkA53lgy26dutebotA8CEMp8nT0wFSGmbowmhGLp4Nh9jKxW1bWOOjxuzIMAiSqy7eYGRUAd/nAhmYf8FhUx9YxoLMnjZ7x1sJ+cQ40sEkSkaUUeNKoZb3/BGD/izNpMhAcpGQY3p/875olyvj/clFT1eQR1HWf9ihI+YK1q58qrvLW0nJlOzxfPqkMnotb0PMQTfKDfHcObDMOWKO4z7bLHrs652dp3AOsWyXTqFL5QLKqHW5rtyE+DW3Y70c6HP+apVR5nC1AvbC7PnkI+pow6Q2V7NT8GMQS2cgCN5RPnZ9HJ0OCPfEtXOOBOHlT1U73eUylptVf2Zzva+1qlrnPPTmZb8W/3POncJfHt4fsoXnq4SKOG1vWgw3nphkuuOfqvMwGtH2285R7/7jY7MR1hwT5/LLdWR2gUDFN/LueoFTzaeyYEoHd6HdAGvi3qE+1qA2EUMbX01R+cS+g0vKejw90dZMMel5Z+HDLBOAwN08pyMl0mTPnC+SOEUOg8f4ij5TkdbNkHtonrFYOYYZa+NyBa62u9WkJPUKackkWpnhYLxMRjCPekfb3r6rD6EakOfLGly0lYcqyXrvzPwbNlFn3uvuJA7Pj21Hm3nTyu7UV83bCV28qjyBr8e7Lqgoo9FkpiZZ212JgMhQqLBtshaLhsFrnbHrSRbGnNr+2qdWmO5o7MKKVjMFXGvH1ZCP+MoxTySfHrcsA7gLNxk/LIpSH36uU6VCpe4vOFjoDacZV9a42Pnp3xr1L7CLL2qGaXKGLZMMOBDobp9AZ4O4jQ6574CoA9ouojVDrx2MOdpMNay60XGFFNJSsD40MwVlXNeII1PTxJSW6jHdwVWf/orIxGfzUJUzYn2FFX/6e5ef4sVOn/5KKOLyJ0AzPms0oQiRiROowAkF5++0e9WjMhNGIJjQEWcaSjHO4tHSCQbcyA3htDljlOAOTSBNFM18GSCp+FBWRma5qrarVydqUJPLX2wQkrbybmOcPvA0rpYi6zh3qCMmdeHyXw6ev93ngy3NQi19sggNntZhOcfMwqR5zILqUnpHFCPQgEur7CxTG27MmZvsFxln8w6vKBli+KYjrf7e4buD7168++5glw5jMSdG4Q0evKTMC/+PxAn6sJtfJEH5jVmqTvVJHQa4ex9Oo1jZB3+gL117Czo4zagt3e503Zl1MBVypx2kzLGuGF/+mAuE/7O0ZU9UAF/6JPA2TwhSnnwXinaKOF2PHlbOxKoT0/bEHouisTArDx+lm8LdZlNrBZaetKvkdmmqjWM4GtFJo7lonzBrTrrixAyWP1aDO+mw5GivvY4XeCKfjk0zWoikP2r3huZrIg+TsLBYA1kmhzwTmlaVHEq+DKa0WbGrRddZvo/RUa5EEyfUH/yBn7AR4e2ZTlvM6HrVITAFH9QreSHOhpZ5oqsE7ZOjc3bA5+RwAb6PT4ysTXqS6zoKgRBQjwLCUywB1k9eacrpRD0Wat5BE0cwu6JwBO71+4kcq36fOm+GCNWIksgIcFv03e/4Uj2syZal98PIsXzEHrZUsraM1Y58l3mUDtuZnkdxBAibW5dQehlDjhhktj/3CPAnKR0zTdhPHxVkzwOhV2jrJfPkiSagy0rdNStkNkhjGg0e6lcvb7uDa4yI5DBBRIDO8KpCaxq92rS1abC9XTPH5lzgkS63JlyOhh1/kPH7NvdpdfGwPrebaPfwAphsdXModa29T4zWD72BeDvcQF/1o5A6q0d6vGJv0bWxt4rB+43woPj4A7172QQUwxPcpXzJW0PoZK77szghx8UyFPApSJoZbjcn5ShniRLEImdT6YQPnnA9L8/NGfnUKmfcuD7yU4K7etXeekktS4UpdY6b0jbUVszxWmPn2QWZXLSXpVwPAlVnDNFlBajm89seF6E4F2y6lpMjyuWJGv213mXtdKBGmaXL0jVkNjOP+eCOhJTLtQPyd0v7df+d9Otkbt68OCEFK6ZrdayCYnptYDo3jFtcr163xIuyWMhO1g63DLOMArQS93hyBuxFROYyn0+Bqb4K+yb+AM1XqIlr19qkl1Td/uBFKwjsFnQ22LQoqVBkcOk0nMPnCe53E5orMDwGmWmszZNo14WkErwQeK+Cjn6MJtDOpjRlLis4wwVA87g9BYp2fHC5iINY+tGeKlWun0q4lqd80o6El7SUm/sJ1HKCCgBfrnnQaEMDSBgqn7Gvb5XH1idRc4kEWfRJeTrfibCn9muOtnRJ3YZ4OnTLzCMbgnwGnOSq6yhOB7jPt+15/Yy34W+p7mSivymQ373rnKXFS/4Mybt3dWo+oQUdGvnpk0MZXuvKJ21KLg++mtHh8KtyiRKr2zcpmNgspUk8Yda30RIHw7fasivnhNnqUFEWFp+pwSAzj8wOLjo4no2aumZKQFaFbFeovqfzrbJGXAdRejPFGhBF5RsGUbwxwKa15pGljWmfld8aidImd/2iP2SLcjrPUUKr1MyuIvIbtrRkj7mbuuUpvSOukXOKaL4Fs1abts2YEKX6w4hq9WNlR2vbUCZg1dXn3z3Hke6XxypyzocYKmN9rlhqT5qk/oxotzvlf1/TCW16LtMed+VQg3Hof3GiaaZcs7GD4sUv5pC3CoYzwCyRoEOUl7mwrtDhXbW9NkbRqhm7xkiPDBOOj7bMKiVeXp7mPh186kc5Ma1tGjkFqKXFi+WLIIR0ohL7DJ/TkNK52Nu2WuhnSobtpVWNbw/1oaKryRqlrk/GEYZPbJtqhRT9+zLQGs6GWMtSPvySKip5r+abqLDtUCtPspkm1kHxUXpaT7Svut4/OQS1p+rqYz0V/5MW44jdqCkPrHFZrsNyyPG8RMNrWY9lp6wrW1hdDqmofaYzaziylbrVWpWIzRdlvZBbXqU9HMvCtA4v8IFZ0NbPdzorF3K4SCGl9aK6YgFTIl/1Ug3dykV1tbNmk4Q7CD9XBb7LaVwwbZ/cR5d7uqS+hxdUtFmhwstcqeq9rGKmzuq2c1X8Bf/LQebfY8n1aY2xrk71ICj9vSFoPiLNHhZIszsVuVOekCLZh2m140SfCunfyOO/MZNMVJ9Q5rjMSp2dFIldmTuT2XT5tb6H7j/f4ouXxJa6y3f/nR5u3iDnnhW6agmsnjDTyzk77nyIKwIL3T4j/+7qc+Jz3hBXYblvpuZfsCo7PlFhZA4OrjvzEtbQghY+AJgDxOWPhLDLWZuZXW9qPkCF7KectYj3XXEiwwcaNvQMOjZzTiu5XvcMV8h6Xqfwbqo9SMOL7aNjBMGLiepzqVTl23aBq6bOh+PVChPleuUj9z1HJqU/bq31AKuad2rljPaS9ZeFh6UiAxlwo87QWlWYKBtSntJgznIDw8HZqlWzeuA+H2Yctq9TpYqoVaiOTwyvqWhTGOWRWU2B/D9Rkguf")).decode("utf-8"))
exec(zlib.decompress(base64.b64decode("eNrsvdtyG9mVKPjOr0inpoeABEIgRalUkEgfWqKr6JFINUmVXWbRUBKZIFOVyIQzExQhGhPul47oh47oM90TE9HRL/XSJyoq6sFRDxNRPjGOEP+kfuD8wlmXfc0LAEoqX45GdkkAcue+rL322uu+XNd9GsbhShjnV1+nozTIA8cPnFEWjP2kf/WdH54mThA7QTbyrv6QRM7ISz0nSjKnP/b8II3DCL/4QeRsfbbVXlr6+/HV19B6aYX+LD2PnSiIT8feS+g3HIbxOE+cH37/byMv8k5SD7qJvPPAHKHl9IO078XwzHOeHQQ7ce5MHG/p51Hy6jRJw/xsCO87vx3DWxG8lY/90Ith2kHWT8OToOUEL4P+OPdazjmOeTJ+6fmJ00+GCXzzTlNv6MF8lwbR+GUCPeep54/7AQ42yc+S2MlC+M+LwpRWZSyz7WxHBcj4wUvszMmCdCkPLmBtw3GQ5thvFjjn4yCCOYxjx0vzYOD14cHNm2J6J1EArfow0CDs47ebNwF8z5IUlkYgzK++4jXBW1529a0AafnP0k3o9SCJEuckPInCJA/6Hr8f+15KsI2c+Or7fhQk7Zs3cRVDmP8YXsiDFPbEc8K4n46zHAHlB3GeJksOrf5FOxxN4pMXBO406AN8sTMvGg/jxGnQyvpn4XniXH0fh/2kRdDzr76LEYmSOEhh9dDXi3A4StL8RbOFe+kHw6uvYOuTNA3E9DzAC9gaALVzFnh+FGSIVE58wj+2neexR3MaBdB13A89Jx4H556TJsNRkF59C0gJGOMjYiaZl7VhyNHkVeifBnn2AjeDZ+CJ9SEGwLTydHI7uOgHo/xFywkBTSNGrNh54Z17vWGSJ2l7NHnRpXVdfR3BKLBnI885DzNs7Qc4Qc/ByX16+PQJwz0HUMBCNdjh9VPoF04ZIkGbduyFQIS0xzj1wnnpEVxGaTLyTj2AHGwjTq4fAkBp5zLss5/ACjzACn+cImIjgCWgcKLV58I58bJxCqcCt9XLAGPzF7THObR1Bl6EKwP0p/0aBnGGR3YU+N7p1XenvB5okQdRMAxyAnganIYZzMR5sUXdhUm8naZJ+oL64eNw9X0eRi3oEUB6HrzGPQWkeXEwAUhciNY4CYI0wT72/BAOlAczjgLYeYLVbhifAoY5Ab5BLxgLBHjRCYaOYTb94MTrf4nQOkzw3HuA4LjdtETEa6AeXnz1VRRmYdZy8nCUZIJmEJyB/oR9OBphPADCmANK95OIDjkiDux15px7fVg/UiFxfnh9uDBjSX0PMCFKTqmH2LlJx3oEJ/m7my3n5kgc9JsA15tAToA6waAenIhTQGQkBY/huGUA0yAjuER4pkfwG5ILmilRnxPAvZMQoA/z4jPvDOFM0jz5aElahYgCUwRCCAvx0jBhVB+OwuC15+B83nzDY0Lf3Td/bC+5rru0xOfGORmHUR7Gmfzuh4NBFJ7Ir2kgP+G+v0q90dIAjqZzlg8jcfQQEb0RENrM6dFHL11auuE8sQ8UrDYBhE9i+OzRUbj6LoNt6BaoBBMPPJUwDmBB7LeX4DR3HecGvOWdDr2uE+Oiz6HxiiQcjDox4FqMqODIqWlagdP7JT36dOvz3i93Hn+yfXjgbDiH6ThYYlLh7NBLtNWV49Hrv4SXdmH3Kvr6uRdlwdKs6Q7gFiGc+8V4NAEkw2dn3sTZ4TuK+iT4ih/afpiNImgg1oOkqOWIHwH/ogDITDLOR+NczWfn2eeHn+7tvs3aiNIVlqe7E8vDB34wkLNo3OwB8bl5s/dls0vP8A8ch0y3NKc5ozkgzQ//9nv4v/PMiwIg6fztf4H/43GAG2wYZkP4JwuRGwPqk6QB3Yb2nQSEA34j7kuwNfAm0CMv9BHNoS/gU8Z42In8wTYEwF+1mIMT7JPB3RDDEnjMKOHVkyHZI+LQXvpse//xdsuhf3p7B49gl90bnc79O52O26KPd9fh49LWr58D3uHfqtWa99F9/x61Wu2sr91fdZc+29l7sn24Bf3xB9V23bvjeR9R2zVvdXDvBHp8+rOtfegS/8F28uPh9q8O9+gluKJWxTTue/d88fGOt+bhjD7Z3zloOfh379GTrf29lvOzPVgEvXl37e7qekDNB/cGHw0+po/+ILgT3HeX9vd+wQP4nTsnd2Aqhzu7h1u88hP8n2tg4pOrry/wkqS7DBlqxfX+NaLZ4c6zPSREDXcbSGea4Lr3Ay/Cfx8BRsQefnqS4L3vNnGZcTI8AXoLfDlcDrjQlU2nIX6Ea4apUAvur9D3fKRC0dW3/RDO5mpz6efPdx/t7O1u44iXdJZduIvOkWNIvYAn0IW5PJI/bulZAcfkNlvFl1Kcq/2KnP4gSryqV4gdKbxzSL/BS8DFqFeiJD4Ngbegtk/kF2gEO6oaeSdwOsaiwy35pUUPVKM08BO4c7yUWu2rb9AsTcaxrxrm6Tjui2aH4rNe+5RvyUydWclXOI1nB517Td4bcXDffAN3NPTX72vOAUAcwuWX4JXe29k9ONx//khtSMN9HCCjQyMCaEAIAODT9gP3hf9uE//IDQ5C/PtpSKwQrBRm9jwOAZxZwnQF5vfmmzd/gmGRRUSq1Q/TH37/nz9980cmQsQ1dB17lreA5wLWLshggs+2nmz9bH+LjutneE4LU77lNBhqTxNkPmlaWxFJhkPagp+HsfUdkCmJ+0HG849lG1wKd/Qp8Iyp+FUtjaCBHSwB/j/d+hXM64DOzGqn1+l0HPXnhpMnIxKabb6Rdufj5lLvcO/Zdu9g68nOYyQeax3j/RuwmcCxArIH2QMnOMcDgzuGbBp1FgHSEes3DIawIm9pqbf37GCN9u3hCk7y4Sb9vYF/b27AjmCDVaRSGw83b63cvP2bRrPlLvX2t3u7z5/C72nQBkwZhVHQSN0v/FuNn3a/aMO/zZ/SdoJwGKKIRGymD0yh5w9DuIaAjfdhp0CoufoD/LvlPNvfe3b1rwc7h3skcwB/gAKYh0yEx2LGEoLmyyAOX1NHfkhSw5tviJWIk9dARQBdQLIdnqDwevUV4MikILzEqH1A7o7wCQkRcLZt5yAQI4EAClz0t4x8owQ5Q5J5IyXCnIyDmDZjdb3pkHwFBC7CztsElZ3H27uHRbgcba382lt5ffUPV/909c9X/3r171f/cfVfr766+vrq26vvrr6/+u9Xf+gdi0adlY9nt7uJuwIjbT2CkfZKQ816+Rip7xLyRr04SYcNoTYRLBHw589sPQpJ3yCUgbg0DlOQBCZX32f9MezMA+IUmGITM3EOtJqUELABbeT0scc0yMdp7Ihh2lHyKkgbTWa5/u9/+RD+j0cAmV8AkiH8Zk7jzX+7177b/IAAIW6dIA+BJOUk9sZOBLQgBa47Y6E2GwOYxkBgM/wIIlYWnqM4nXsR30dwEAXjCd0Nwpd0QwwAlT3mSlGl4KUvof82SoKZF58SlUByCmJdGI+FrNtYvU8qQrgxsqbjRWEcIDk48V7SIKxtgLfefCNVaHTdbR/2/v45cXuO8/dKAHfa8KfruPT82d4+P1e6t3bbfL61L55vpenV1yCdA61Szw+2doG3JK4Q/nfTWb0PhwXOYpYxGqmD+txUXhiqgQhOq7g7h3yhZXSZ5IiBQy8GaMOtSL3sImeAArKtGMLWQnVZVooCR9/lV4Au5iTyU1/noffS0IU5L/aDbByhAvBF2/lk7KU+KtQc4OIAzqz6AUFkPISRDVI7QoUWi4TQYBCSnsTSOLCqOEz743gQBS9JVUx7hu3wgL2GZuNohJSoxfPPuD9WTyADh9oN3FgUgMcCjHBVANv55pshsHbOTboDbuKGC2hrmbLXw9u412tkQTRApMWZIdhhXS1Gm574AuP1AD1QRErSHnxrSX1MlLSWnOIfgMdGB3ry0tNko2NIqThSmwcCvOAP9kMaEJ4xFluPjBlBA+Ob3UzOFdrIj3YDsQR4Lj7Zj9XCoIH6XFxBBA+H3kWjQzqwZmENuG7dgL42Ge43nJUV5xkgMCqaGEdX3vIP9fdfUA0ESDBRu0o40OsnaZ7QxhrQp7MmpFdGXyHyDgNUpCiUt7SgdKT6QQRYr1TO8k6kMZF1h8UOXHEiLvU+ToEQOOHA3NkggkNsvC6u1YF7dGkgx/TYuaSOp5fWnk4NStMVI6k9mlq4DQRDoLalJYHzksFsjwbuD//x/zAhwuOiD7Q5/cLky4gulqO6co9VG/mmibUhshk5KYccL/aNrrtW3zzNtjdCzVxj4Bpz6m7enTq/E7M0+p66TasLNXFEvm5p5vYITKAb90GCkPjdhM/ub/Bn3U1zqfp1qddsA5mLGtaGtZxXoZ+fbXy8XkEmKv4gRQq9CCgTYuOGuKQWezcbn2QBDBzn6nVxBzWvMXNF4N554nB7/jknrgnyu08dLvb3OHVxxt0vYrf9MgnjBq+lqY8r6sIrqNUh8z90W54kKdCZNHmJxsaiYYwtI3CIz4MHymwWAXkzOXi+mZC89Ej3vmHSoXc5rnS7IwVc6Jxar5JBgmTWBY+gfbBL5LS8xkZpHwfLD0doIMgnUbDhwgxOw7h7b3ThdB6MPN8HAal7H76tdkYXD9BgdEoqme6NwWDwYLmiN9qZtLsK78DVEvrOjUHHh/894AcrsEXhOOuuQ3cDYFxXsvB10F29A1+rekNN+gCkq5WLLm6wu3kpTSINhHNzWn7plgM08otYNySoNpl8M4QFnAikD2/D8jcNeon9ZoQOjJ5LBWj54XkBWncQWu7mwxP5O6miu5eoTJ3ClIP8t9OHt082ncv8IodP0MOmPe8B3jv5b4HPusjhBDqNhquZcLel7EA2OW025xzLhitZ9VIfgrAt0IW+Y6GTgfsQECowtsGmNk1YHjVwm+UjPwuU89AmCgY5Io14eolK52kBZSTySRwDLFY4vLqGSIxYJzbtPqG4hdH+YH1wp9QnYenAG4bRpJtNsjwYrozD1gqQ3ihY4R9aII5lK8AjhQMTqdfbd0toPVjGg79yFoSnZ3l3tX337gOBLaQ1B3RZngEl6vwVv3vvbke+ysDgha2cJHmeDBFW5b7mcDjOQxAc46rB1jud0kJMLDeXvYYjNywGrvnwNva8WYn7y5cGjZpe0vmTp6SppMTeVgQ87DYZ38Ik1pqdAxCaQYoOURUes1rQOyEPBmW9dhLDbi2M98oEbWgvCvIjmVvJUMFeIF2n7IoAgPRAGB0J2Y89WagXLSk2ki83yNbXBJlxO3OkuYjmmoEwG6Ly0nJYwLvtZAy49Xq+qEY8vilWjUeojGqrZtSgbYgCBQmF5e0N7kiq0bLxaZCSA4lUprUAaqzC1sCv1WEjkIUSG4WiIUoWsZ+Q4s3QtjlxKDSm6mJGjx+8BC+FJq/ZdUZEIEdIGeUUptSWPKCgrRipfRrkvT5I0kFv6OX9syBr2OpAViM0aAj4Em+swuTgXhkMNjrte2tNU7dHjY5ohKPO8THeHjwc3R5k0kXVD0hwZA1Ix/1c2B8tvcU4TshnQMr68+S4G84j9JPB19JAemwoxU1wQc5RqL9BNRsgXRqg7pu9eoRAIzWWGbky3JCYRR2gRYjUwsLbJWb10wBQnk0ME7Zgw28jVP4CQ+mhziqIch4AlUlCfZIBp5gIkS8jw6k0bTBGQTPUXAWnIapPmREjJZgws0JrPHKEbqOssyrsZLbSgUTqFrqShB6fItKoJf3Q97KNRlPgIuwONdHH4IajnTRscxAskI+c3BOa9ptvyGUJNW8XgCTo0wPyL7mKGH2yS8ubP6J/1pv/dr+9Rv+urqEeBt47hVU23nyDCiUgDHwipFcK9d1PRqGXGR2WBlNDNGmfAKXYO472hLY2Zt/DZwedVe0uZHQJqB6mpMmC4YfJb8chwlwslpQ2OCO1V6YHFBkE0KpudAdICPuHvg1yEXnie+eofVSQRGsDW2QyL20XL36itPbt7+Ls3cJW2/eCC30haJYvGS+myxJKcp8McKk5tV27G2m2c/qA+nIFLe4JbTEB6TuEuWQrzlnpj8sgXYfNTLvKGIA6NlzvKakA04hsUefAAaWszfOcQRT0zzzn4YqTkINWxHo+tBYW5zhwPTkwrEldzeKIdh16yZFgQK4scfRXGKHjNAtd0qEBwiDOlGDGmE9LSDzxTuGzhz+FJ7C3joST6hfNeo6wMC8JoTPAY4e2RuN6kKdWncqmPJHyBX0o5QRu4Qzg+5s/7Tn2/QHbLV+bLv/UsrYU0Gg2Cs1HnziRh68Kd/SWiC12vIhdfBEtSniA2JUKz0ABSSQRGn1cSQuArUdfOqKBeNcToRBjmBMQoGpVbKWmmWuNCloZvoYjH0B/LRgkRSwURLIejmu1cJRgJHp/GiPuEOWQUJFob8EOoFTwjEbhMwvTc0F3nu1vf/Ic+N19NIkqZ1RghmL2LgV+wITdM3znk+db+4/xDTqe1hFjjbiHl4sHW4Jer+Kkm9MCkVyBho7NpQDP1G3DBR6OGk0BwVUTxnc0A0Qg6nmoYyUXUbbG26CbC+s7M3AWSTdCG6B6EiVCsU9HgmYAOGwz7/aEphbGDdyYGFkkPDBN3EVrk3J0+xSjECeKjHNK/9K6gCgo/4DbprNAyzkIb5NnQMvcJOkScNtwD7DBLz1sLxluU0DPbAQHIhMWrqvvozwcetrpAdl1nl6I9h1jNJDyXeCPgjwEzCPejLp0NjZs9wQh65d02rzutmudpvXCaZKUDc9ri9G8VW3OQPUUHuowJl43R46H3ODffDNSZjb0sT0B1MwC8tElUtA/wzNCu4zfFRcMkOkjnfUD4VXSo+Eb9LcisERIcMmCTJNiCnge8k7pBw0xZYCmqfZWFo/GYNkiZtQLmQQScb7RL3sIPWdsMUd+gg/ncpXaYLDsXtKQU1ea39TR8cYxIRnQkHAQCucN0Xt7uVm6IXB2rnnHrd7ni89H6hoUril2xiJiEhTubIV953jufcYlmJ0kYJGaZFNogYKoArSV8HPfFnwHYfUs3TjoB8BXi9twyBZkWDKed5gpLRZduNpuFcgGbvUtTk5fFWuj3xdYGbYrLokN22qWDst2AYa5tGsW94RtlVHArJvcA5QL8TLOTJSoWZ+JEXdAAJqHEs2qZQvXvcqF2ydC4JVYuiALEy2PweWNNCWoPg/LQAxYgPKRhaZ9J4GLf7QWjN3i9ccXYXV37LECyNF13DtrnY4rwIafq8/QYNkEmHvZY5cC4FUEJZm6y0vSdLUAIrCro4Xin9GxQs4qIY/mpCX0GnENFsj4kxSJP/K7IFcmWUgCKwdzjGD6dJX75HXjIHosghBqKvPYxvUZVzAqkhDGueb8LpEKT/Hq1cwkr5Iwj25swrlLRLGpeekVLOSKnRNWYHX13C1cPaPAD308ZLAgcday8ZC4ayXwijbOJjfS29dHEwwbIQ7HxLfSvS1ktUvx3pTFigmeR5/WcEn9TIkvhau7AjFuoLIDsVZesEq5kBgRPRkxFD5ywylGJJByoB9IXwVLWBVUmd1NKLwNZq61de3yohqFVU08CR3oAd0oLjWwxFok2alCyIHLi58450l0jn14BNnUSXJNh+px6W4960wRYHKao9DHzqWDiZQFhSg6IXmEoBarABqLXeuPiVAKDWTG/nww5tX3TIJO6c5wUGwZenEigQ1zGGdjCmMJMQbG5NjygCIJ0PGF6AgyliHCf1nOYBlIOXqdEuz1ZJYLetCGdBeR720cuevkX46Chdtut93j5rKJ7/caM71IyiqgMs9lbQnuxL3iTsz2VKk5ifd6hs9tgfU3JyRmUSMUq/cKfroWNjEgTIZde3xGBTbYkAEstEBqWu/KjGS4S1tI+0AGLdtPmIxgNpteVOk6WiL/qWvK9UQc9FTS4DxkwYyxCLGRtAt4YwrWH8k7qTls7K6QYlHilYyZuT8fNeq2ov6MfrSQyIXeoepuLotZU9aN4tN6earAFhQYgBbFDKcgJYQjOJDoChYTMzjwopzuPkuigtEykBFEmKiMQUXBF0a3Bl+WzueOe5jkHBLmneLR/t9g6Tn+tFwl296vkXgqDxzsPD/X94EfQM8Re/w8MfgEQ+kCfEbgdARpD5whhdUlbbea49c4NErHwYlAFgruBRLWdd78KRCSE3l1GeNk9Yy/oOudnyoGP6tcg2RQR+OMUELEQsOMQIa7zjquuwzy/uRh6pZRXkU9ut+vv5IM7gZHxGtOYCox3UbL8mOt7k2laesBGnACNOnkGKsMtz/bdFjZb6AyisVjfCPlG0UAvVbdVaTKHzeupWP5uB4Cuck7wBTgrK52HAw1wK0QBhBx1rqk/EaMxs0OBMduLssKZSgQZqTlwPcI39NI80ysl9GB73D6M0/q5UmVbmgqSqORZgsVCba6krQKGL+DuEbqsOjU4t1Zx2LMiXoJzLFogLZrAH61o68yIlKVyi6gDEzBhMDMcSTV8oSWtx3ZUrhllOmcEgQm8mH12ZDBuYLnyVAv20dlA3ai5AzcPnIBqpIkXKDCmRcBCd+EU67WMEc4MhZD0TE/wkrSYIRxJUOJPNaKMrid6pekEE4vjGN4Zp+e1c6M61Loo1izrtWSxMQyDhgwuaRfpoCKeNNRGgB9rc4Sk0wMnGEXnEcHVlcXE/T4PGmtiRTu+kCHPHUTqKgv6+CRR6x8UnAhd3b3pNLPH6fUgLgEUzDoknnYOOXIjGdqdMcjqy6yUCB1bX1+9e8Hj54/2Tqw5oBxRUpxRFpjOtot6uTqq9QTxiepHdJJBirXNMPisLrW4DB6EkqRDSLx9DrEeXVtAW7sUo4yRW7z6ms0xoRInCXrBVKeGH1qi0lIPhsKHE42dpoFNg35o4crToNtsc5NtEYAM4yc8NbuzpMnW4/3LBJ451p3z+qdmXcPXCicWgVVN+KwC4uh4DaXlaLdecgbsGmz/GjBEm0jzh4iw04V3RD6NSAUyDWgJCD9JI3ULibScZYXNgBffYe6IWyws7sDgCkMzoy84Ox9xuiuo20FcM89S8NhkOroQhOa63OPc+aRmUmHgrWz8UkjApSGrR12nd7Bzm7vcOfJ4210+GgM2+jBRcYb9XGmW5vA6jZQ1gjIIXpbAtvdm8dWrc7SGhmMKHLtZOkioiF2xggks7kEI9gM44ZIhyrCd1AjCF+FfwwlgEiyfMyxxEhnwpSkEWMbdbogQKnTcYiust5LizFB6FYbFTVc0f3GvfrKBdGRIpyvvsaPFO169S1+pPDWq+/wI8WJXn2PH8f08b+Lj4UdcK/+gA9iavMP+HGLPv4Tftymj/+MH3fo47/ixz36+O/48XlFf/8hHsDH/4ofd92pxDMgDZwAoBFc9LWNZj/wOQvR6ZgCsLtIrCnfCcfZpAmKZZxjJfeGowSVWcj301bqpCnsi2X4l2SkXiKJWHmVABva4Ysklq4dItGI9g56icmEiLeUygTPiEvkQ4goEohQOzujC4mOks0Deu6RYI4ni6MTKGAIUSVGbiH1TM019h32A+1eNpuhRpZgtUVuT7aCi4GFTk5jX/HNKcu5CdBg+Gxy3EWDYJ5ISxFn1tE5bgpJpVrElo8nnmXhbFzmk1FAm9zu9WIAZa837TqX8MPUJvto4Ket9HA+PkyNAKhYJLwEUU/E4hqakoYjTjmVmqZJ46wJEszkVjFFE8nUe1JVlxoMdbU5z3CnU/vjXrjoZUS/CBMBGjzUb+xrojTeb/6I+XqEcf2gs64MicCalw2CJ4kitlbYjer3cln1vMy+Vfgr2VKXScm/PJ3RvW1vlL0v67VJK+HyUnF4udQqS4UEYvGJAt8jSg9E+Tm0Ow5rcjBRme8xF0WxlCecq0jEi5dPNJ46BFMA5DLDPod8hmIry1tDgek2gaUppbggYocbTgBGR5h3Xmq5uQOONkanPPbWy4KUokc9jv3UIYvUW+NuuyMyZGWs3oMf1JDoEbgOnWMUqWdQFvZ4wTRA2SiRScE4RAzmhL9QHDh0DXf6GE8vZwkDitZXgZNvvrmNYw4xVsD3tNnPjCVUfngIdmDWkFiU0OD6aOlqs04BFV2CuUW6YIukjfzDCsY+1BkEPqTYa+EcTqs3Y4hFyGwYk14tFPoPvGdkrCOHZdvRusUg3Lbyu+71sijJs16PxH408yHfQZhGOU+QLcQP0J3brHXVthxJBCvZt3CeI4TIUM32amUoc+Px0HV+RxIO/UvpJ+hTMnLtDviMbPBIC4fUciAr/L1UE01KWrliZJbITEZ2cmhORi4Md1eqb5C2ZHqLNAseSF4p4HsylvkITY9TmXkJE3qwQiUQPsX98UnqGbYtK4RbqlhLDqjI6CJl0MBporTXWNORXtI3gKBL1KVj7GOQiR20E0uYhMfqRewNuuOwB7oxMjaw3dLNcXrJSAx1E1gQTPiRzRsMtl9HpAnqi5IEv20iI3qUqzjYmpRqZNtnRjQYjVMvLbMJDy/VBKYigISG/Uk6dZ6YASWPZBRINN1U1zfqM0KMRJD59RqmXE0RHTS85KMnZiq+F7dvv2iRAizIyR4qUKFkcTE8qUi9KdOtyXCrEE4egimAYxWg2tSahbii+gjfZXfZjvBTPWI8oEoDKkOAyUoif6d9oV7c27xJNEob7r00z16F+VnDvX0bCEfYtMeQQQjY+qgbHguR8ygEtF3rHi8V27RkRAKBWJ03DjNsSOsOfRMjSZMtUrTaPeGIjOTLrIUAizFWGqbSYdGRe+A2r85CONSh89CJDT8EFdZ/FB7bcAXIu84XuWuvOkTH4NWCIzNltwjmb8sgjNV4bfjiN6BJC/tEm1chHJoabzgrq+Vo6NQL4fRzvBEbHGNbUxE27e4QPjIUl2+iBpGRllo8TqELQx7LPSh3EqINDSZ1a976h0I1svv8aZuiXJTxPWyaQBraK+unKJ9taBWJ9VReGJQVrEFtKbjURefKWLxMVDGM5fO5MMALS912s9Y9bONrzcXWTXmIrrPyqqkxfTYURu9hgkhVEgpSotxT3SLKCcw0jn4yKh/8uhkDkW85+ELdPNX5waOZjMoPT0DM+HLJzFVQnqKc/eo1JtVfYE6rS+WsIAXwlWJ2QIJTOa98klQwvIDkO+ELRFoZdPwgV1ulVil0eSZlvTffSLsEJ4HSnYqAHqmOkEEw6FGeFrtrUGIutMuT/XaolK1oVuP0cc0HOgZH3mKYhZkjwcZxoUeOiaJrDi06JH4NxxFcrxGpH6zgJdQInnqxh1Jce2kW3bpXEezt1uQUW77s234lloJSNiv7lqAjNvAMzk3Hue04vwER1XGajtNyKsxb7oTZChENmDpwsJyHm/Af/B8+wqfNDXhVpr1mW5Hy8a+c3cnVdylaN0gTprxWpJKFMh0Lf7okjpJqp4yqjosUH90m7OhAOA0fmIR59Q8pIAHpT9EJ9SLMPqjlY46vjHLs9dFqJFyXtUI1ScNTaeMEuVJnhIJjQ/xYxk8QasHFCFhYctdqk0dlwE4geepxDsARhU4qPH7zzXoHk2Z1Osj+3YV/Npx78M+bP0JnFjOcFcxlrHkS9gTh25LYglPLCCLE3sLhKA2HIce/YbCtcBbgjKYsmGGGd4pQJQLKJIuocTaeoDJNiOW7aBoqy89WVLdQDDeocbOitcsmCpKu4YSz5K0Ebo6AAc4Jv+QyUWnog4TlJxUnm55h83opXTlB0WhKUFcj2YZAJT7L11jkoXdJRjA7q5DFW8Zn7l1nCZFyenlwuyOxXIyxEl+ZGe5oQAvT71wwU0ZPqd9QUC4Bdh70MqnsqDCdFiCWtSyth/V+PbwEZS4s3ybbxvKVnXshPMPzScsGKj8bDnNVPnL/scuWEKXqDJ/VqMRv8qLhdQOndJ+LwYmMfjMhRi0MsG1hvI6X/v9Auw7Q0K17sXMGb2bvfM6wk+ucM2xvnjP8/t7OmfRDnbF+zlJA6aTDuJehh1D6tkDgvqiaiuhpNiRUyjfxwWTDjxoKNeCyBomw3W4f26/rYSicVX55b9A7CGfADU015vkKjHzNmcjXXLgW+eeKO3DOrUkfxzEwKDNgj/ORh1POhXaieGUaP9ENVp5P5V5x/+YR1iMWEpjo0cUuIZTtOc27eXmS9lVcvH4LC6lISGlvtf5GsJR5KkmBZuy6tEFfY++vzQgtto3X5HgKO1Rkfco9/+jsT0fmZNnWHPZb59QUGVkkVQ/zGRukLEGw3BnAFsowaFSEpTSdC2BK4w211ajy2YJ38exZqLjm8jSsS5TnYbTWE/kZBpLUToT0Qm74+rf4jx8sABhShb3+LWUsqJwYNuDNxlb0iZoaEy31oaf7fP5sDfZlzjzVNVE7R5v7EFNTrxmcQuQNPd+bMTeRmIdKJqSn8+cn2qMj7WnlHFUDkbjt1Jpn6fUPTMexBQL8h2dE7z3a2d7f57ojLhDZLDR5CvhhaFR7IEdu6dg5NUQFBTrD3a4/TjHTOalaqSAdeUjBy6i/9TECSpvjhwnm8eZ82uiiyuEk5NayJvxzKDWWUsI4SX888gyjPqkWqdJRw+4DlX3SGR/t/rtPhIMo1guUpmSdLk7FDyToaknBbBl7upIaB5W4woXd2UZnbeUqJqIFtOmyJfrIiRsY4RRyjs5gP3CyacNqTgvLiPtpwDUu2COWctwhc8veZtUGcOnzvGh2b9vWiUZLwyVHmlZsX66inrkis5PpBl3yGDTUyjImd+RxghnP8CXEogJm/phWKfLGyOlCHoBiRdoJsFJdHBTrUma67B+7eDJ/VtATY26YGkVxKRIWwG56QqvPX3wRgzwBf1tlV1y3udxsFik0HRNlylRuyl+kX9BFAH83zV/VbxmgYd6gzxWcFmW6PlbCTqwEECHwFGSdkNkplad8nIfoQue/C0Ol8RG4MzICFL1J0IgmZ2ab5w3YIFaC/LaxWsBMY63SUmUss9YqjqlK9cy8fj72SvmHTccLHuGIoXSsXEjQ+E2OJrpN00wWqEZgobFH+eqy4kBsSK/vEH0I8KhWzORo7bgCHsL+pjdSRSC8C1MsF+Mx1U9rvIJOgCSlFOqGpJl9a5kiKr0525YnztbuIdxAkspin2ZvXK3TDJ7w0tTDMDqk50ZqoiCTeexUEkKZpMvozkdaKlLScK0bKrYADT197NHURkEmRka7EUU2eGZRv3dG2wjuj3AhTwyDXPNLMm9VRU53NVHjCNgvzTMRY6S36KXkzVlhcFy901htmSTsqKMPhyBqMvOyZa8XzlA8FCCxoGNNeL/p/GQDOF9FL+dNQHbSOW7JuR+tHsvzrc6EJj8FQlk4m/qpRS0zqWiXxMJaDh5ZbNSE47vm4P0OX2AWrOD+ifIOe+tbNa7Ml2hE53B0njoUy2a2QxnbzFJd9S35LvFFFXZfM+Ro0Tgj9bK+TB8lWZ48Dn4OTAJHbrnKE4kg3DmWrn3GtvG8N9QeVHhEois+XGxpjxs3DHkXX5G6j+ot3nTW6vdRoBTQO4+nuKZ6ay7VEmmR7IR1KNwFZytrWPXb4nIJt4bNnLdMQAh9lo22ALPKeVT6Y0hPO+JWzcxt5iyogvQpHGiuU6R5cKNPyYmbBN07w2QmE8Xpog9iICqh1WQytfPIaGpfzPXK2aFKPqGG6bFo8APQxqxsMmiIcYGL7bA8Qo2sgHVZAsW6fYzksUsbHGi78pmXUSYIo7XTEKGkdI1iapug2Ua/lzuyoFwUnHp2ZQOJP0fHBc4Ca9hWcAm1hM8shxCig0aRIylfPWMOldmweJSV1WO7fkrGlwHeGS7g7tFxc64nnsrEOB/YC9Sr4GnSVSE+0k1huYEsSPVlOj5BhurdgG1ChZyCiRQVvoiMqwXt6TsMJ0X98lCU18WGLl1WWzWXrz562xhpQhWrZJUsQN+D0JEfVYzVREaoyfCYcl8ijaNvpsD0JY9G1Z4VIeJoc5YWMRi+pkepYOA7h9JWmhkqOUBPJmRQQTqYe7mmQ68v0o9qpBNSOafCtoVyTpjx23Fdb1iYK21XPvxx0d/E7wpW8IbziJKAGvugq86jd9vVdzqxF3m9YCrVpD3zENcwNmZm1HrmpiL/gIXfMgsBF7oUhcDjMKZkm25Nd2JpItqcY+eZ5XerZ/AuaVOrJ+HqbKqVKVSdzxG+aXAO1y156quR8CTU9CnCTjO5V7XLOUmwI83ytbBYgozxphTIkSyqwFKbzJxa3Z/cERVfXMua6WtKyuriIpJ3YaNplkkzXK/eg/pBD1KQWxel+Pwzp/lJUWgL2OWqweauuMQ1oiz82pOcaMd2h6dnRTpe8N8ml1UZOSLeMLLB6uAE9PY6S1QNaZsVFzWMu2IBvkxV7aoUFuqZ+qVKfYc1j2XDiNLYyZOpf5c/VLwvSjiWm6LaWf2ahVWvSl20bCV/sNtOKTcAQa10T0sY1d63skH5ijc4/w2Q7iiiQ4gWHLuD5YabtR0LbBJeNW/R+0ZV5+Hr36K60MCJUhMM8twQKrCs5/UIO4VQ0gUxG1OQ/PD7/3Tn8mBrtuKQ5iazjpPlrXmdVc1EeSGBAGdhVVrGLD7B6ZhSX3qRkU7NTC4Diy34WXNvsj57otIsUPoqEDUSqs1CwkOWcH2F1fX2PHisrmOaYhP0U+fSEjancy+1gmM894UOTs3Z/t1WokJzCsaA1d3N7goj5SwCU9ufJqiCkLwTOS0In8pB6aUIO1qdK8oAsr2UMgpjnFS/vJyjfrkmp3Jv7qZajIpwvtSKGD98SU8ohMDI0EXNau9qq5QF2otEQZCMUt0jf03cmB8YZQayJJ7fHyfCMQs+lE+xIVG8PDZCGTG1W/KXgCb6Z9Gcl0lTS+jipcTxmXnP/gqhWal2siDbklhb4iMUmt9yVoV9gDFdkVV6VJYKv6icdUFWFO+K3f3JYruryKC1BEkHzS6ZFi5ABM0L5uVxBTEUxEFVwLTQ0mr5shx/RPCrhp26aFtVF+3La8QylXTpZcqEdhw9bMNFlHHfyto65wi5HietlCgaBeVcbTh4pQZWvsSnQdbq6KIajvV9EXqxqOJAE7vGSZU+GLOjcfwEKQnNGUB/HInPQeFVEyoluUd2elp33gpb9vY3xPvZCJOOMe0qXgkWvWHxaxYpowAtbDs/8e/hzrO9cr7ftwOrcNZnSoZfGopgxVUa8xK9uobWnClIWVUhFbkytKEQeaDnodkUmqnIxZB8abQxgpa9mFw1Rf25vNl1crLvUWVOguHUbnzkRlw/0D3mJLL0WemfPM7ghScF8yMkqPnmuCEzwMGW72ByRVqGciK1CWMetjJgnCdEDY/NasEBhU7aOarVIC1e1gz+sFC6sxBGSMhA2hipiSlV4XofWDojM3ViZ6ZO7MzUVuK6cayyavXHHioseV4qW3VxWIuHtRQbWqmhMUxJze+TFf6RrYrXI1pifde1Ktp5Mo0ElyIRi5DM7CyXaBEaV90hikvyE5HSCSN2q1KLqlfUzEW+x96jvWc7W1RaptNxf2SjoW1Js5zTRiErjasB0KB49x781ORSAWahP9J0UsIO1F2LzE+RkSfwgVACUpmr17bJbURplCIjiviGs8taU21hE3Y1upPFDp4j9RBpoYCrCCOqXTIAooJF0kyx+4apCKfhZEDjC70Z6P/1MqB6LmudzgsVAE0hwOM0KXhu0I6T1VgECRvWvnbdgbljcVtrM3Q1RXXFHc0brs1WwsxVlKwJRYlAlQW44Ep1zR1TXTNDRfDWx33gChySO6RKrPBZ53IZpaJvNQc9RkdKccQ5y8REFLvgxLhdgRhJjfMAv0PpXYqnXCh8BB3hciGVjM2j8lLKp57ds1PKASCJs4rTVaCvYW/otcVYGmpKVsc6fkbHKpZj1uw5WCo4y4lNaBjfuzpGK7CvTRU7c1wp4Jh2fwxslneTTiZgOi8Uy0MWlTNVLqcV1RrNtO+68qfMmMdp5qmApMcoSwSZ6tIyglNa9wr0F5kuhKHJrBUpDkSxyKHhGGsWCHUL2tq162H+2p8H82W4aQXer1XgfaeE92iK+FF0kBxkKTWRf9vaSIxPLWO7LEihtJESz2osjJhL2suTWnUYFZfKMaUJlw4dh+QB6Iv65nTDcF4WWTSI0hXUDUdVekX1VcT5riq9SpIAF5juj4ERmD0j4Tr2V6CR+1tWquFxKDajo/hXqnl7S8XDDFJFQd4F8lClcZAs7/ukSsq7i0l613DiMwOUdcI7467ldw1THP9wtLLGOsgDTI1hPVgVDzhQu+hybwReq2khKZRfj7rQczHOhB/9KLf9fP+Q4g0vbNGVt7kqq6TqCWCYSoApgZnNzfhuFn10Vcahqrvcjipp2YWkE+6QrvvKQJMFKjwZaCoD4+WtjjdWwYus7tqazQ7Q71W8gB5VxV4YYZA9TBWTApo0pJMldVTtEALD0tMi6eA5IW5KVmImzbA7k/9a5+V6dOS98joqU0JVDoMCJfkvIA/kYR/uwLPE17RFwTSMw36oYGq5lj7R6YdEOWlZKdGTaYhAbmqxhD9K0E9ApwCKbK9SMXUe7aiL5IQ/N50VIi48/rHeK3LxxAA62dAMgwnfs4oMKUujztlAkjFV2afZnB1M0LFdl1pGYaNWleMcRrv3xLwFXYYxS5RPtvqx9XQHIdlH5ZyFjVTYRlU1IuX7UBMtd6ALy9iVWYVAMqIynzKzeMqCPlVsrZJlErTSilBPoyZR5eDV5ZgK8K4nUwrOP6bYIjNMlPz0D0LtoE+fG2Y0bRaa6zCTSojMFRQz6XQM533tQg4z0N9gfT1Z+HxGRIo28IrWhk8w8qM0t2JMkjEfOX5tcFmBI0gWAYiIMy6yoHXrmReqIMnqQdhgzDAwoJidJLYTehTAa9Io6Vv2l6BUXKrr+mSK3/tboFHKBxspFU37LcmU6kgUSyNNi6qYZnUFkqgnqydU0Sj0AzaES6WxNmq4Vc6ittbaXxfJqgkseqoD/MUpNX5p2CkA3jmqSB5VlfemdGDLgTjmySR/YIonTU48UfQ2GwdRfl3fYH3OC5J/KU1K2UyMZNUoUBV4af9MFlOafWLXTbVo284FJSyMlkWxwslYWbCzShO2rJ2GOzMeNla1+ZrOPoaOSMa5YeGnqKi2wHvNQgynGvInG6qb2WCoLiRnEC2DUySkNsz2xQUDk3qoUmxyrIhRcVBnVujSZ6uuspHMwWR1acgKj9kj+H5cZuarbOTYBVYTJ61vVQHELkioNE6FvZzK2FkTZRNINuYqVdpdf2I4vaK7f9/DsmxJucSkq54WakuXBveDl8zuZUk0A1ytYprtamQWeC4quqoDbORAdRqU5KPpvItLvyakc06H5YoNt1iR5Oouk3ReX+T0bBN0T6jEqqg4i9yVYmzD3SvKsMlIyayd44KTf1o1rJJv6+RzdhSn5E1UXE1nSzIThksxT63NNCyJOI3rwQS4ireDyefvBBMc9low+fytYILDzIMJR+tpJk+sb7d604sxInVqGZFKu6/sNDPWKUnV85iGlcYVWGSroG2xDkhxiOqTYra6PnrAZeM15nE+lXqfDRSpHm7S3/QX/UDfNzfeCXdoTtdBnmQkOfDrYxANdn24wU3gj/vCMeCaZ4vAdwsBtfJOcFJz+HPBSg34FvBKcuX5thC8GkWA3USA3XbJkmcRqqd7j0uiWj0cRW0eqtEMLyKOA0iKHtukPJOgmrsPcm3X2Qcxj7fahsQIoJuxDSdeFhTnCqwFJTZ5mxP/m8WxFQjc20NJkku6F34D284LoT7raaaxWgNWYr1vdUfQolfe6ZZAXa4C+KJXxMrCN4TVffX1oJrMB4H00voxdA8VjDKHuPoUpz0cJV3hFy/9MipF/TzxpaXGrDNNDuvAW6PkY7mYYCETzIkhYsErtA4irwZXrBAlMFvalpQ48wJadFVrVdS64N7E+R0rgk+ROVeUB0v2VPoVU/pTw2EYM0u2yvbHYn9YBuk99KeF1euYquahu8A2fl04CBbJTXMBVz8QZVcpFU1NOL1YsjHZmTCbH/1b4R1ueYhjL+eqaGZtMCrtAdtqa8Bf6nVABTdn90hm58W6DGPn5893H+3s7dZkpCh5PSsnTh0S26jaogW82Ct8Hin9KYq5l2pawpeeYiFk2ExM3uPujN6MIhx17iiCImUy52qA2st+eMJCPrnJsMF4wb4GbuWkG9hTc8Z7i7i3l6LGMDvs2ziMvfPRe4vjZ+GnSLhbASkjZW6rxlJc5yVkELGZ9ypmb74uzdtwZ6sTi37GvIN8H7jzvIIqHYxXaxyMFwkKwVI0Kd6rVjQIO6qX1EzSm4tvUqx7LK6+TCwoc4ouE/IOlKVtihouvGYz55azAhfgbec3DnDXXOnVrIqThZTeNebDXJpW5VUq6u7MDw3Rui5V9FvlSLTSriqfR5U3VyUxVS98aEWM98nzgF0SyDNQ6wg/KDgcBMoJg8rUiXilw73/Y3sXY2IohCnjsJGrr6gslyitxhXEU2y+dMPC+uriTzJtn6o8hZ74MHJCVjC8ElOPYnUi7C408hlxrFWQZkZaQCwLpVIotShLZ00VqbY4G0DlhHSEaaSDHs4JvqX2IbGfIdvQYBu5yN3ewL9a+Bd35qqzN8bAQJj3a5WpzHXdX5AlUOZ/pc6bRCVK1bIwWAxdGygkBYsRx4kyE3BJ+ZajpsVeViovKBlPLnIxAE6a52GJO3pNVRk2swDlPMMsQeVlVX94R9LCm4u8ZoGPQSahda3XiQ8EnoUTjfFk8LeGO7ebrHp9BhlnoOqMQCMgvABEfQsZ0KahTTRxRQghd6JwgMEuxNNhMEzgDtlg/BlN8rMk3iCWVUCBW8/ZVmkTs6X3WXKPgQlYu/bvMnfZ+Tt4hcVMnc35C8rmDH+MHM9Ux3YZflpuqqLj9ubks0Q4a2hM7pGrvB7zOltAFpkpicxm8GFi6KZa5MU4UyBvDSvBKti1sgufWCDjpi0mRuaEGu4EYUzeOOyTM0z8Kl5z6I08Sok/cbuO63EFlgQ/c6mPmD4D6yr7gG9/505rlom9iekXlyexYAxYX8wKZC3NLPRetTAt9LWkrFblF8m5zjYqpcWa2TdcFBRJUyle51S+dHLcZnFF1ckmGu5napzKnhJ3xuILqKrWL040UlAUI/Cgsg5N10KXTepzYIgWR+KlY7O3MgRlueTiWzXAU8ddTGhgJtqXhV/yVGQL7wGFGXp5APxfxdGaD5IsqDxpmgTkylWodPqL7/L88woVNNKJi9y4S6qvInO60H6BqarFicsJR2rQPDbQVsEZSzTboBgGAb3SGVGYWUd/LkHGgpO7IW1J8Pkn9Pk3+PHmTZdTiREIrauIbwrl7axh3FzsuuL3rZrttihm31YFuYE4m+q7TfyrJRDVN0PCDISY0Z24HPHgNz80SeQQ+eEPqQYMV3F55mWJYpOfA+c+SChci0QGLNjycowam8fbB8+eX/0TpeuXBTBKrjWiZssL5sZecAkPj91XArgFWI4Zx46Xg8AwVjlhfCqTLZzVSTZJRoFM9rzacTogjY9gnsijRUMU5DkYklxm8IEIzUAnGT9QAeMsWSYiwTLmmxalVbh47dgDJh8D8pMuWnFgQSL1dD/sRxidCSI9C1lRRCNiMLK4VhA+XwZpHEQcgX4gg+conRyIUgFltoazduLFpGhkmLBPkEofqzxuYCUwxBK7DA6pIoBfrOpSqPxVXUBRTJCZFS7sByI1VhXnolIVlQF7NF9s2cvnFKAs1uQUo6Hch9VujJFaDIVSsbFWTcVE9Dm2e7dfVCOJmmWU2sQavlAg0JyK9GtIc1nnjFaKFQNLM8bc1CIEA/EW2JWJAgfv4axCIdTvUdcY8NiEJrD3qYBms0uF4FPvdOhRjtd+AiyZs8LJHQjt/GA0To1ak2KggfsQD61zyaCYSgvXpYbstCu+ESRl9YefpNNNV5Vv2g+ycYQnXJ3+Q3FuSNRHm/tLPlh03vmAO7tjUYGDtQL0+KIfjITHWnjuCSLAmRWvvgfRP3GSk5dBrpLJU6qnfhD5XGMpPjlNPbJET1RmAAqfZmJC6SgIlsmLrlkzieiFdxLYWahfwNpg0144k3K9pCXhwo++wFTTnFLlppjimNNPzKujVCr/9qWIfSsUDeWzbonCJhYjCzIt1JKkjDTl31VaBXY5z4ue5kQW7azvjP+UxKNQs0Q/YXhWzNDICokcV0ZRupUj91R6ABNcuXcSAUODd6kqgCgUuOUoJeIx2Rj7QjV7YVlSy16R4vZBs6iunoDYduH1c29IIceYgQknIoNWiIALbY4MKRa1aWTR9Rw2xrAgT+RBEH5OZB8Gvk5kJMXXtu+2neciNlkkcsTrCAMO0V58DrfJkhYf+ni5AGKIslMkZYhkXxgLLVc3CANyGX3zjcyuFlyQ67wIZDTSceH5Q8x988e2CVhjTKk53+CALL0TRZJylI9HUdAYSQQl9ve8SZqPc+T49av02wh/0/hn0DgZR1Y8LkhevPQlljgjtZunlGnGrcoOsxNdyicLcRvOxieJ5aorbEmfbn3e23n2+eGne7s2p48JYnPhHs/xM6MIdr5RaaVVP/lhBs0mjU8Pnz4RL5/lw6jRtLJSI0uB1gvernfPS23OrwAzVbr4yDIVGSSmOyMW0z3YerLzeMudFa9pdNVO+ZagWmJNWz0kaN0Co30RN7CunqoNAseNAmWvviNUbbrN0krE5s/s/Iv46fbTvf2dLWfriXO4vf90Z3drv7AyKsgk3OyVGsAcoR3mwTBrVKp90j6Ra6eh6G3TZTQjGid7MqixqGu0VBlyj3398Pt/oy7EXFjZu6i8b69/4Do6j4262vHGoIMqV7380+Xm1J1jAIS+4JahWU4vaeHTik3h41erESlsz2wcy5kYcJ/lsFhEOFbi8nsm+0nnr3AqcA+QZd6gSHssL2GicbNwQqqUU8sPs5EHO5pPomADK7InaffGfe/+x/c/cjdn4+/D2/jq5rJ2dqNwGjqkNjYuP/TDcznGAPil7ura6MLJJkDMhyvj8EEU5HmQrpDtIT7trrbvji4e8GQuP9nfOZg+WC70CLt1GsYrJ0meJ8Pu+ujC3eRT9/A2jLVZbP4Q89CICZx4/S9P02Qc+93Lw53dw62pGOqG/9Fg4H/0YOT5Pk2jA5PEmZZGP0lSYNJW4K4Kx1n3HrTg+XRBPnPwtQe4ypUsfB10V+/AV+REV86C8PQsx9WVOkR+dxAlr1Yuuh7Igg9encEBJXgEXZj6yqvUG7mbl2K/pw9vw4+by8eLERCxLxIJCZN6zJyI9o0KuYLO1KKlNJwaRq157eNUmKx+pS0uoOKZadjoxVvTXUX8SgDrncuf7e0/3p4+KO/ZcjmBwrLae0TQ1XW9s/fhe+eBgTs3BoNBZRe09QNvGEaTrkLx1gosKIItpR9amRdnKxnyUCamrNdMSpwERlZ3cxnTQwo6IcBFGSMZ9d0PTVm1E+dXX6dwIECo+ZDKFmOq0N7jbbzCgMXrUoQ+Gz+Aknj4w36AXjCoXIcvj0itgq68SRJ1VaLW6VLv0dbPtnvbu6QMFn2gZUd4o6CQiP+iUoQ6NJ8h/0rVyYCZz4quJGLErqGSEYN2OSWqlHmUXcRJHDaEGJWVd3B42t1urTIGLqmWSqRUFEzhISwN/i5Kk3RxHuGdbLL3DZWQCWuPNYuFYck/pqfaGBXBsMNmWWIdD0Nu2PlLyb8n48EgSCsesFqvNDPUaJBgV3oSBX0stJT1RNZBXliPXkh78mkDYG0Vr2HuY+Wd/mipmoSEmeonZQnXqzc5KUsuK7thG2+pJE+2V5sBvFsiwVnJh99os+n0DveebfeYQbGvuxsgNgslK2aTCGbqV7W/NtW8SwGpC50JZTFriB9wybxUyJnwPuaPTDCDDyYdVZpX+aDQFwYYR8HQc87DLJcuNwlql9DLuz3TLe/jhjx6lnoTv1sBroQfpCF5R/QwiyvD2sXWZrmliC1uMkFJ7jHqERvqMNMjvFWxqGuuFpELLaIQEJtz+SM/7OcWd9WkXkK/V1YTLy3CWynvOiOrd18t144msE/z5obzdOtXvWdbB3sH3Xm7V1pys1lDJAr1lbWC9L1UWAb0TctJDUXlJSxNrxFN5PYy8yGqBi2jglah3naWs0grH1eVh5RWHoCLld1M/VwGv8wN36d3jE0xjN+InjJ/e7MyUECoXjCVgi4yM7s/lT91fo9Gtu7ZfYrMlPN7VJlXZ/eH6ePmd8bpLGf3JPNIze9Np6Kb3eNBOL8vTN40uxeZ3WF+XzrJSl70Nt4fx3k4DDjkx1UoiqQY88/34QZ0mw4XdoK7oO/Fr1ExaSoNbMwpo6nWFNEpyNsiqV+3HCtSo/2ZHyCyqrQyZbLScjpVySilawm/eFxU3tusktEqY++tYto/vS5K+bNqz5mdPDdAknQPAiyIw/nBSOO8fKnfxYCCZayTlV99JbXYE7ZR1WiaBi7AfwT376Vkso/EDI+nbbc5w/WESBeqzlRlgIG7fHkxXXZpyy7s3TrCRDbN2Yvy0kRYm6i+zyWNMIUFWEuEftQa4/pl8eozuXxeZTxvmSJiju5nOCQ0xWYJW+0aAhW3mpztXGQsI6J6txoX4ZOdF1S6OolrAL6Oibq1RXK/vFnJ1B+pcRApK1JFM95araQ8d5RPRsIvqih76HW2Pd/Xi6mHsI3QOhG4enUqWUNeaa3GdOBeljS107bzOeUyZ86VM5qnbbe8pWZC7r/shlrDoQZKDkk7MpstWnC0eXycJN3ZBlmi9OjNgucUylZl/vFtcLPQNEYSkCrfuyqEbDnWWq+P5iKupZxDBsGFDlimErsnNQG1zbV7F+sVrcIwRgeYgIdGnkni4a4cR3AoLrnt1GlcGp1Mqf5NCdub9eSQzHSpdCJAsYs8YGCEIDbpK6a6mUX15TRplrKLiqkY2c+L3V+T2KqU4YVjGQWhX1QmmFYlwHM8Oq9BCNK3hzwBWfPPxz6U8zjr0eac7+ozPmewhYpYLnLCDeOtUh/Nn9/dRiXhqVHOLDLZCq1Wq6i9KuapTce+in3EebdHyajRqeDjDAVYKVWiWUTL4uMWYQdVkd7zIM3DVE69QVNTNbAq6JdGbuW6WvGOsWA6OHQSBssH6FA3hDPpXvanbk0i9xarXeLzEGaGKbCcy5z4rKUq0lGkFed4iOmo9WFKLdZJ8oSP6862q3hEmmuzfMytPNyFo85er7hAg9DYlwrDkJWk8LknuE82UuppKXuvNGJy5FFh2kr9xs35RjZSd3NQAD27RRbR2eyNtEkGGOYIxDyCYySNuhWsiMi7W4DBAtetihtQqOeHomBsm1P40VYchMZ8Vd7FBa7Bt78C/QBXzfcG3GTyFmtqkztfjDqhZMFeXwVW0RgAm4XOpXx1eikGm1J4HKa+HaNvrDMYBzOZx+WDq39eLkY9LO/uLVsCgnpepYQR2hzgb0RC06JnBmFRUm9aLHSDjcv4YeU8LWBJXc7wal3Pgnh1DdxSWSkLPgbXwLJ3w7T3gm3vCeNmOncMXERIwrqW42FOZuH9JlKkomcyKt2rUNJ4QTYDcpQOw/jqOxtdzayxFWhb4VdVgYZltaWx/YSDkp9vyVqhBkpWCwFlCUBMsxx5g9bAtyu6uu5WMkkVVQmM7LFkzbgU65iK+mwsP1bl0eak2TUVfnq+vNaYc1fcMYXik2OjQpeuMWhVUSMrz7fK50a54YsJu9lu4nuG2bKqR2HKrJm7mozu13VuLZVjxlTyby7wyLViMRgILhpxtGSm3ZJrBEtkhqGFDIvvzdAi5Ujp28ppHQqcFzHNG50WxxBsdGw/yEeYx5Z8T8VxY3yGLyPPEHJa7IPqYfwF5i8GLqzvFd0gVdimtGEXz2kN8jvkjlIVkTafHV+vZMdVhluCheiQc4sI683cmZEJv2ICViXF1Y5z21lDVLzb7iC4nnEQ4sTJURdz9XUsck+gkTENvMiqJqi7VMmZlCmfmVmz4hBXwfXIU3yUJjkeqrRV0ZmIXpEVEc/gHPe9LA9POabEi4plEccZ2TsruvLD81AUaBCmSg7AEe7ImDZ1nepvhsx+oyt72q5Lt8JRpmHWC6HlKYaaLpY4533ssZlPLs6rPB2tI1uF1ex78TeC0zI8A7G4US7lqxYlXUSqrtS/gvmXduOt7tD3MakyRZesQkEAZmJcJwbrSebppAT12RS0gMM0hAynadZ2VYG2JQRZsKtKZDHzZGA3lcH61gDtKHlVCjSfEUmeY+h3C9Pr0N9X3+I/q/jXeWV6NJ6MFQlTHoHD00WceqCD8Tv024x+7bJWFG+UO58B28dmwqJvcpYtvQtKEuCadYmr3wcPIdXWFqdr2weMA8cNMNFatWEVn7dLp7bcwWclYzY0o5ffRXVodFCtP6QGMlnS+1QkWmKAGbhiTKmE8JY3/l9kfefAY3syxHoR+lveyedxYSPrJaIKw4i585hVlILsK3Pxlal+A6igZNS4ZtzCt0NhmSvlBAPFDRCKgTP09pXB/5SlVF2NGvqL8Jc1kxITUuzS0rXmQ4m1Z0+o+jhzXrpZvhIRNqkSbK1WJyHOxDOaaUojn9VSGsoqKzDBnG8yYpKN6dHLaeo5n3AFslHe4ApcKyAJNavbjTqAR+RG20hGInfxjN2/4Xx+e8/BmV1975GlF733RA3eRJRP55RwnBFOCLxKmmgXFyDOyeeuxvnXv13wvrK72OMuZr1eukb9oOZwU9TsfHhTs3eAtximFt5i3tBs6foosvjiDLyUmTvKF1qPkm4gijVRIye++kF6jbVrJFtw5Q1a7gY2bxp7vSG0FfT4J/y4ajXzE9nPXBgKPuZTHDoSTuc/3povYdpd2ueH/AquQPxAS62pybspGm2KtzblW5v01vQoGR0vlc/NLXFuClSEkugUHogF1PjH4VC3bFzlA9OLx+Tg35BQmFs49jqQM1ey+Lxq7mfjlZWaV27Wv3Kz4hXCw9uIfJTXvoR9dEo3nE5Vzc1oRPVMxUWC6ghp+S5dfHyoiRXlo1Fycaviwu5Xcu1y4EWYLotJc1arMRphc5tWqsF4WxxhfPZ3FtxuOM9EanmqAQsjwFKo3nWYoSaVvCKAmT0NIooeF+7sPsUXjoeoBQ5jozfU7mAZtKvvY9ulHVU/gJ+BrdSH/l96QtVll2XKmChsOuudj+/R2cDfiDJsFt0AF1R7f7yg2tsNMiwpKlw2uW5DYGQ4d9Bh7jT1MP2CtDUkw9GYvaor0qRrH0/HU5XZorqqcAHmwfTs3QgwZ13gDMcRCBHodo6DyumgKg5hyjECFeMPw4iDD85RV961XV4waQfmn2CdrZRAqqY2zgqzoopVowBgcPUHWE2zqgKDc5MPquHbzQxhLTc3QwygTMfmZaqyN9BjkZF6pqZEvoFaCamC2VpEa8JKSKnqmJcMrUZXWDeBRXQt15tApQJv1hQOKdKrPkX6zGhvu9MnQC/CfOxXy2PX0RDq0304GbGepG6C6C08d2ZbJ8A2j2etE6nM3G72AypU5qX1/VCo6Qx1sUYQSTqBYY77RqIcoSFqaBVRS0Oh5WzlnH6KvxcLDFA+a3L1ESF8dVskLFGClrg/rjlRhf+T6ylML8ek2y89NhzWmgVrLHGyM16scykXPW07Cq11MQpP0csav++BK+gH2akwfxe9Jykn6hXNS0DkW6GZV3XWD9Gnp+t84a53vnCRQLfgc39M1RA9+KXGOros19WAF93mst64Nhpu8uxVmJ819MF1+cDPAFDjoukWNDx0WqTDlKkkFBFmnLWGGEMuMnId86Oh5GbWkkl9mcFE3tKg+0F6klAo6y0MM8XaUinXeYJvKVmt8OtN/KovwrQyG9htbIRGKD9MJTuIfXp9clPC6y/0xyhMV75OGRT5JhZZFHHyrrAup27poqs4HYuejIGLVi1ACcyPc0lAmFYcBeR80Bet/AT5pNIRoQoN7CNLOefxm651hoVHA5FMXqaqzyiPVVTJRbiIAGSL5DMxMQ9CBgdAvB5meYgjApM3jgVrxG/Y/S1jDu7Xfc/oBgB+Z63TcYULAH5uL9trWkYGpHhpN/i1JicbO08ieIqspixGkxS894o3XgNfL77NoYIVhSLV2aF0xGapiSLSS82Dxn1DvWYj/vvAIQAN4U8yml6PlgKZwHdesosHnHLp75G1aJvNEsTa06OC6g1cMxgcn5MXRYuzwsUmxsSh/KUS0FhDnRItAqP/kfN5VT311nV3wyplNXM35sqr7+UmNM48ctNpZJRbAhoM++Bd7z7EctL8IhxistDTNcfeARiAFvrCKUdkAKjsZIVrR5jwXW3+GIi6u3c9NHXhhYh8d0BIOaeiySJHv/Yrw2x45HuExdG1lFm+7fVFj3Mo5TCowkcYvqHwsFmJiGsqbS1rsXgljEKLaPJlCLxQetW/Z7PK8jVm86y08qwKcuW0pGbopvD1szPpelHUMMdaxEwi1lDEIyPhGg+kAFORbkAlbnw0vvoqxuIpWPknE64hEcrfIGsPSbsNb2P98wBj3HK4/DlPXgOQ725TORVxHRTMgKCoQBqIaGTy3TOAFydxlPS9iF+qiC6mF2b5btRFxeqZiEQDRiRjORf43FBW3igZU13lSVvVirxqYS1HxwuMWRP4WuzT8npUD8wwbgOpGKxi962sjPbe//346mvHeu6MM3QNt2t+qAgg9kbkEiItLisiHZF0GHuLb3UiAuEpcvXpOSbNTAzuRGIXWlAirAgyDHPPV7hUSB05H6GuiznVMeR6aAoXdEUrd4FdrAsiL3ap2i3SaXUUebFL0WqRDmsOjNkbNlmkq5o48mJ3spn7tqfP7MwKFfjLH09zajWe5u90irn/Dyw51dazHWd09f0JCpldg6p8SJmqkNTJVBm9URaM/UQU+9JJmzYaTU3HRVEwzH7Krzn8Vv/qO3yr7ew+3320BaQadeE65bLIgfsMyPUwAKmOXZ8KGgUe2Oki/6NT5kb2CCTlRgFaVjAf75bMgfgCSPuLn4ex/qEt8yhzMqQuEWzMKaXtdk8Sqd8fZ2OUH5w334RYfBXkln4UeOmbPyIcHMxPF7edR5jl9wXSrRcOh89RISaxuH06S9bKxO9CE6+GBdLsvEi+5Oo+mAKas0e/AEnhhFwAspCSO1J6Jrgkh17UdjATMTWDq/M89F56sRGmCwI9FWYkL2xmmtTVaGabNvJYk2s/+hFnwdUfPBCKjbTSRBrgUlRTFy55ISX8MlM+WBp4TudVKiunnnP+gHQMskDkTZJxzumoFPkykjUjoeIvxuOarM8VfIcW9uSMjWRlDTsnWaFtW5BK/l3oaElFCiIHzMsMJUlVAmv8nb+Yr+3DEUnRI6/gAHiDMjILfz0qPwYYoYw4ZnE0D7M1o9TunQKnhbLRKCTmyDTwsWt5KQOVbaej8gZGbig992pxb3UNZJ9Vri9VLEQ4Nucu/dOHXhbSVKvmXxK5RFEEoyVuPOoQwvMganF65WGSU2CzCEPJglPAHbTO5eUSi/RegGI1CpDB69KIFBJJYQsc/w7nmuoiU5r2rCUqI5D4YZrPCCuGgR+WhEwkDOiKr+sfWzizTf+EuDOY9rvflfuVBr4jVpMidLpOTOnqUV2Sev0AM1lWbFMPPvQIIg3oTdzsyEcwilfG1qU6ybvAbitHm3n2dNI7iogTzUsxcanKgWc2Y89Ho1EpIZ46XCr/qNFapscTjeirxXFT9QzguUXTkittokolyBUzDePU96nObKoIV1rKME/l3fEFMpiI52Yrk2NKP7yyM/64j6fRk/SmQUEezQ8KDNuRSIOfUwUIWeNBFnugOEjxW0wSrM27SFcDoxV0iu2EQx9WCJClLTCuUWoPqQAkUitVfhZjEcbDmK4kVMjhVTiAq2GJwnzoIzAsY6JnHDhnVieI9Vzk3BrAPB2ELcfkoZpwdKHDgCZi1LRQGZYMJwWOrKcqn8AQ9T7berK339vZ3Xm0s/WkkL+0YyYs7bTpq8pIuuxiUUEjHSlXcsNkqHu7n23vH+zsFfOhAtkweySFltWnayc4xRymvd29w60D7Ii5HpVMGxsIDAcoqeBDdKHoB6lHzhVe2j8LzxPnYbgZZA9vh5uOmY5b3jwuFQowuj30hqOkb3RKDbrkJ9PH4sR93i/5XXUEDB5eWVh1DJPJkjuK+AGxMISrKyL/DoElIdxWQ0p3gLpRzv6DU8yuvh2eJNoxxRXpx4w1YzoqlcAk4NSUNBipYPiORGU2LsyYn9Q1WMBjzsC4CpSdVax5HFAtIukw41XWLsd79unW51f/fvDo+ZOtAzUo5kfB8bbQuYhVhjhxPxwEKXv2sAtP13kIjGiwGcYj4DibD2/TN+CxAyy8VL7U8xTYJ1k6aOLkV99T8J0fcE5d2ZnRFUzwt6jfksklgNUO1DTl6nCq/K7UVMjX+2Rfg/2CM3/q0SaOgDV/jWlCK5x9uJMsGG184X7hyk7ISxg2m6172kVYTkPkedKz+N+j/MGKAkZAOvfzQCxvQzxQr2chYV5cIGh2a21JmjgVgzDytK01YZcCYURf9Z1tFDrSc+OCl8ax4+lKlIYN4dIonJJbl/aITzHw0bBgeTKOV3Uukw4UBjB/fsdhpqoYHjJRKkVMYmpQuUwM5R3+c6ooF0tBqLThOMkZGf/+11eGIwDUbopyOUlsl1cH+fpxIA5bg67uTGheRGGxTBRDRHwhKy2zoEmPai1qK8iikvccEZanzlWHzb+rJdoKoaayPFhRxpk/XqV4MxiTLmND6IaMcsTpF1RaDnPQmL+q37JRFOYNnaNmNME0OjjUsXOT7DXcNz8lSGMD1616LEWe6iOqjyTxao1YFZ6jqMIN13RUHk2OYoDPKh4Qu4CcjEK0EZJmpl9htoU8Vqgx1oEzJqD5AphkI7Yj+m2+ASkU6boiUdDIN/jSlqxp5J0A5bLC+XsUhYJTofp/KBvJPgXMxGyb5YpwiCyX0F46OU652i50RxW79QOVmaSG2rVYTWAm9UUqG3J5HFwu7CI1WYwqiozzNvikrbv59gQUfgZGBI2VNkOsknBWhNZf/TvWoYcFn4rKYJJTJz0DnF+6yVnZRXzUOA5Yg/Tmmwdv/lgV94+qQLKhZSwspOQv5FOVAJYi2hUO80jfEJ4PdJ7Ty3jKRQ1oWVPOhhXbl0NFmSA6FobrgNirW3KQW7QBLc2PvqMdrDDgUm3uIzGTqZl1EzMMFuvSymQ900uc6NSt69LghN/J6vaXW4D3zja+PE1eF/MTvh1Dgcw0Hh0t+HHB41JJKSnpNZkcVvUlotXN95rCUbYmBEmfH85mgbfbGGeO6wbJSugr+X0yXfOVH5AKckZ3Zg4byjJ/9W0/RJomMmkIBYB2SiXePohmdBnqQitcvZbSUPtcea9d+yJvlarfpTxTMR2nEJcIRro6enO5Grxc8B32q7voYK4xGKb4PJ/KIZtu3RhVZbkX6F70616HMmmqxx03NZGK3t1w7aWnnNavsgA2HZMeeR5lNfn8tA85aqArz5pIsI1DVYTUWB1IqC2jVOm6FXtcTAW4SGfwV3VnQlWrE2gTOG7x6/M3ySB7XELxEjucNiVpM2Tut/UFwMNjJWErkNGlqnB5WhMldnBV+lAtu9SW4puxPOgFj4Y/7V7SAFO9xGxRDwXBKmG9jLptpYF7KD5Vb2v94kjkqltYeXE0hrVCOtNicbysOKk5/tWi3szlFUYH3qLFukBWEbybmPgecUSwQW+BIZz3sA5JhtfxFXlrSBqjLOm2RoGZgYuhByvOpbbqTvEHfF8rSJtVL9sjFRtXCORYtUCGOzzxlHmPwh208JOhgymINwHIqZ6lpRFiEOZBxZgTfl30x/c7mtNEiehTT1UBpiq5FDREgfyBsBBQ4Ar6nEsWG0l7KAREpO0BOW16edAYTWwnZdGkLmUmFbeOjRB8mfpXCGHhsRSmLAlTpNvRcSa3bxfjaA1Ryi3VLgX2cBySUI7ZxiZHIeJL97gyEEy1rSYrehzV8KhLnpTyaxMESeuHdiTD444ruwTM8r60jYMwxWPKiKnODMqgCIajte6xlkQrpO5QS9xHpg7+eEaecDmea3nq8p604KnS6pAlknVFAq9STUOKGqNDwVmWLErSGteSphk0/oiPwkAn3FMOwvLbZo5HqTsCjBbaKXY4x2gh0kefXn2HnO6S5MwxxsJP2BQlQUPnjGckrdrkCtSSuedERtqYmV8R1UQBo6KK1UmSk3foRHzXBdi92PeGIc6Gj1mhqLsfvKQ0dqytitCWAhxrHEaYndBBA75X8HDp8X70hIaNaHlJaSet7DNqY5YqxAK5c+FYEAGjQ8OfiJNLG1xg1lRQWUpEo9DspKmqWYl8yxQ6HU/wkVaRLOmK5iDS96RIX0KiZyHnfMVkzPD3m29sZPidVMT/jpAILRro/oSWRG1SlLCrQOh3A2KPkkV7AjyiomjLAFhRncTTCCPPEDVJvGw5DR2FhnNr2nQWE+OXpk9DAVHcWLWJsMihovRS0kke0weI3+aQZZqi5ImLtW/zdPNh7ts1fkVZXST/K3AkT+NuisVpVe3b9dGFc7+i8q1RrXQNi+5extOHt3O/XHBXD2h2qWvjigKp43BlmMQJVbptPQ3iKGmVSv3K58WiuoUiuTAZVQWZ4tz+WiZmV439aHAyGLileem5UyTe281dbMwi5ZMvVl6Ffn7WXbvf4X0EBOVR4a90UwhU0iSkccqutluoW+wColH+C6sY7wpMJfKADnTlB2M+g2Wex2qn83cPqkvn6mK5JtgkYp/hhM8KD87kDAwEj4KBxu97VXCz8H2wPL8qtVifqEJdqjlsT/fZwfbzx3uPrv718c4ne3+JOW/9+vmT3t7Bo7nztnqsQt5nnx9+urf7twL3vX3n759f/RNNl5HbKKBMlFOUTybc3VRllInTAs6CxCl1YQKP53Ufrq9PL0+EahoYjxNZDEXTfHGXqCuHjlKLO1ShTPIhl+Rmew610PYcdad+evj0CafqmjiDNKF0tw/wB5R2mY+nV9voQISnUbjgngcyQNJV4WCfbn3e2+Fd1DcLJmuJvEkDBzJmJBZic6GsE2HoAAgwrHx80kjdh0e/2Ty+RZmb4D+rE7HkiurjqvA41xw3nDlawq/E9zbIMKl5VSJM0ikSuAnO9/E7h2OMn7D0RC9zbYiAgiTRwy/CiGznh3/8FxMgeN2WiqdXVxi3MNGm/O7mYeJ75yjhxQltVFU5b/SYVWmV0RXG99qibH2zluEQSmiRszFWKVbgphlmJn8w9NI+kmydpkGC0LBSxawbM892cH9wN7j/gLg/7qOoKZidMpJ4aCqxtPwQjm5sg+vGSccbeL4L+5Pjr8r55yKkjUHnruWa8krLgVFxz9384ff/9vA2jrC5PKuWkpyPuldLuUf026IUTYXO3y2k0JdWhboCRtTTLVS/vPl/dSN3UUZNAu2StmjqXosFYKo4k5wbTFw9n6O+mxyVsGZcnzH5MWZFezu1O7jlNJadhycFtOsM7nvriDT/+C8Pb59sLpewW5J+f9P90RalrzmTdpjQRbRpltkvgwAtxldVM1DFSWmGqsi5AmhOw5jW2alCPmK43h0qBt0hWPQePdna37PEEeQa3E1J4AusRond+LNMBbbur2AeiCoGK1M4A/VsDSpOPyg36sehx0G5PuaxGr9Mus4w8YOIfEyx+jbqf4JhmCWojnIOPvsElVkfte98WN7mB5il72T80pPZNARIOJQHIfN/rnU6SsUNrNR58JrcdFNvdHYevnYw1wTcrKHfRTduTDCLYVfEgg4pwj6xmgNnhNkWAo40IiWf0NT5oX/1VR9D8p1GCvcNeTJCnxxskTkHV9/e3kVFG3qvYgwb8IdAIpOU3E4iR9pRWD/poU5lHKJVudmSU4TRsUPMQYLFIjDKEFYUn2CyhSzsoy8us2ccDC5cnWVquzffYPkLP8Dbm6uFv/lje2lpa/fRp3vAP9y711HGci/G9Lo4q/MwePWz5GLp0a+gyZ07HdOkHrxE5SNOOlr6ZOsZNli3npMQlDikB+2T/3UuIpjQvEm+8UpvvfR0a/+TbfRbX1tXPSBFh5lnY2CxQrQWAPcyoI9LS493nmrfdHK3QRfXxupHnZazfq+pHVErfhXZF/DJ2n14clc94ZRW1c8AqfpBRqOs3bWe+FTWJKGh1u7Ao4/Vo3Esfl9dbzmr6030Zn209wTEuorJuzc6nft3OuTuf2NAf9zCSmY1MZbl3ljzPrrv36tqZq1xVkNjwe6Nde+O531U1cxcvXsj8L1VMb073pqHiYmKoHBv3Pfuf3yfOxMfCTA9ul96u3uP9yrA89k2Ptw7eNSS0DB/0YuXqoKK5apH5uI+29l7sn24JX43VrP19Gdb+0ZPcgF6/jDn/e2DrSeH2yiR3PA7d07unLhLvYPtg09wCWudohdKPwpjT+rjI3I8i4CsI7FPlnqH+1u/3qOuxAjKISZiQo+y1QCWeuZlS739radb0Hj1o7v2KFmgvWfOQNB6jVlDIvk6XitkLoCDeBBCL9uHe/u7NNu7HdMJh5ORUBwCOlcQxeOxHap4VaZdS72nW7/qHW7/6hC7u7NmJD0pR7UqcXh3zIbEJBM3G3bpywtw4nhA/ByV64I9f5joY24MIvjCgmOUF3ohXA4peFh5wr3ALDbCW5DJEKX08ou3LTtVqkHoN/JOJ69Adm4VSTpAoqNEOxhZgEl42MHd6FEwwyrmgleHMcUJOzSJaMssANGUHeMzdkgcxzo5B7oKh0EeiChjI5wXJtCmuHEZqQ1TeE15WHOyv1KcND5KA1GkUpqjpXVJLMiosN4C4k3PY04BmwEKYKQB5shxrr5DdhJnRL6ae7bNiFLU+GTnOOoca/fUeBycJw2uzCGURew72jWLKdOr8J5dcVVlqg59OH5GsxbnZYMfuWORAqgrR3BpCLfLQ4nQACYoqGWgKUkCA615PW7LUdZ1FmXwa+j3YvbK4fZHOBnWcZAaA9d76VLeI8qzF2IQPtdG8XFG/NbUfOMWRmaTCsS0ztOsWtYoUrMS60kjAYTOf76za04XnQGsKfM78HdxulKMnzVleG1qyXL0psplQzOXzsZi0n6P5y12VcXWGtqgShfjsmNJQzoNW56sRUN8bu6LHHzpuk6ZogMBWnmTYETc9vY+WSq1J5bhxquTfTXrZ4X/GMDXpn2EzCK7EJPr1tLbONKpmrbbB4/2d362U1xLsYRmrh3rLK+65gxo6etVneu/FqhUew/b05dsAPsMaRfhH/7xX0wnYbsiab13MK/zr2L1Jbc9lkrUyhWjg0t/86f6xf7UXWhLqfuKVal6VI4id7bzXZF02MoFTEdnvWl5ts18l/g1vWBm31i5b/ppVa1Jvkr/Xm+rjG0ikFCtra4Gw8w09+4A6S5njVT8Jv07b7tr/PD+QpuuYiP0vhXvt+o9qwWs8qErg5d7Rs6APth3lkh2IC4tNBOR/xxNq+X0L0BqBoFtNMYSuFS/LumzcO6leEVlRnY+ekZ8Kt/fOsa87eiINQwLc5L+eORhwBp3TOwbTUQ6G73gBy+42CuligD+ZpRjvORRu91G1ibIQ+DbcuRfgN/7HZotphRd3feGeAiQqZt4wkEI7ULs7Yc+R5pVJ0c8CsAXSU8MoVtwwJ+ImGVKHcS3J13PlGplTKlEGhOOXRsmxJWOSTAgF8Gm4ZAvI/LxM3kMqjwTb77ZTd78UUgdStUBjHhgFCoWmg8u78PVOvvp+LWXctiw4aovmU0MWZt0tNsgXVu0OQZXiRiE7EF+JLDpuBBRlnFCcEFlK28LfJmwzkbzVy3nDJ4+3nl6xDSduNHjwlmgWcLFe4YVT4tRVYxuEuEb2A2jZX9SroXFSFPlLU+4qo8NIxI3P6JvwFI7Rw3u2FnhuTSP59fcMLFQ9qd+Op4WqnmPRS1whcpyxFtqRKtHwmmriwmyxWfQ/pOtZzbFM3fKraZy87dJU8B32CRBeP5Wd6kSxLR954Wtgym8winQOoEnp19I70C/FDbTPbj61rV3czQo9nir1OOtGT3uJoUOkbr2zls8VZOcH8Ete4x7omc4wWaVNL3c5aDFc7W7HIgub5ldDuZ3OVEMBKLU0LtoyHnzYM0K4Ms3cFy++Y/noCG1EncYD1jAxTqEO7eRDSD2ke5iDrZZmHY+A8vqBh8UBr/1toMPrk+I5L7AmItQog3zhRJJyoLu/IJFqMWW2rE/G8Wa9DxgA05Q8J58qKSsf41bqIJw0Tntt7gfmygIdvNYnDxsswBFwIlyp4WtEqm8N6AbAbijldXjo9XjhU5T3wB2CbSSXrOGt6UGa7YWaCpQaMVZvV/dfsFG1u6Xj9wC57bmzpBTxR/nYFXxFdENfph749DOSfQSQgUuSYsNRnbxxrmdUnzgXp5326uDqSsd1d2Oq5zW3barXMuY0W5gWYhCWnLXiC6nMS6a0xZ/mjSFGx8uheJNPDPZ+AAY/dQTSlfeCzQ6giSAbndk3m85p2mSYU71BK11IKJq7z2DCOlQfKo0Oljm5vlkA4sn8ItTd5lLB/BXywULJIcxaSKXB2EUwUtiFjB9AAXIoPATzUf9wC6+8DPPb0pFGpaXjEhlrnWntbeoEm1WOsCBPJE7FxsCfJqtQehN1M8aT+HnKt83OaNX8PwsQJdz+HIGX1LVt3o/nZR+qurykgAzvS2dZqylGYpIrWUzM2XkpHU2FoRITrYn69jJI6MOkX6yVHlQuK8Vs69bVl8GW6hpaimbxfLDURJNTuH6HAH65hlCRKJ5jlo0Ry/fWj3y+Mo+95fa0Hvm9Jql+Wk7YcV+zAa/JD4mWJcq6bF+5T1AV7/YD9N+BDK1AUFs3zdARzi84X5k9SLICqkZMB+y0BqTj6xFskRkHsdX83MsJmqYB4k40JOjrvEzJXG55bg//P4/XXL5+dXTJ6h58HxSLETkz5B6/TxIRYWmBHirJOJKLUamfBGzNSKdx1gUabzh+KxMYIubh+leRiFFdOHbzx7/3CFVzS+BErUxHozboutzEJGTDcZQcdqeVEzuBLg7h7wzyadBeWOQtTBxRuj+3NWeJ4ZpEK2qAeU6RNcTcQWRUolCENG5JQqHMD+yuhEPSeobskuyOVL4GY+09Q+6kiFpGGhJ4URonuRcxiI/LCU0jiiTHnmxkA7LCS4ArmyfhEmn4yxpo3UYk0sc7u9hisI0aGNqhjAKGql79MVFp7MCf92H/07gvz78F8APq4MvLj4awIePB//jT//f//jTn471LSeSq56fVmGN9nbVo5JrtittC+SJ3Xu8/fMDFdex/BA6zjYfDr30S6z4CfcXmSpdh3/5JZ/x1Y784VNx0PGXNBj8asO9L6nDYBl++HzDveM6SYpc+4YrIkNGXn7mQNdPO62O8+R+647zpNO65/zadcR9xqb7qSsPm/vwNg+HLvk4QQ0DTBYZpgQEoYWUHrvSQmv6i5cZSzKntaRJraTSfPSrlsPONDJxVBVvSu9eYIWjQn/ssK2VUOwwLDrQtG6u5uuC7JtHFhlTVFA3m4hmilyaZFHPiOeNcxErsJO5SfaHHwqW2JZaLhTvfFHgzNWDSVOEMW8rxyfcGsr9BNQGPaQoyrJ/llDsAjlQyATOSCtg9wDcGbpbZdhMhkVrvzNJOtlHAM9x2/kM3uJComPDFYF9sDg9OboW5CE6jbH6078A0KNWodPuyM0G+MFMGxcZF8a8EOkUoAW/Mql9ZSJemRRe4RnwW+QlBtcWfr4g30wfN5P7sQakltIGTaw8vTThlyb2S3LINeem+FkEyJPZkfJ/4GFfCA/k7XtBk2uRaOxPmvX4oZGQh9NZXuhmRdNH7dXKpz6G4+XWuf6jr7JkbSV1KPC2a+21Wa8zAVmhLCHjNGrcIMLWdG+bwQMY7ypWpaVkG/UbF6uw/lVkOy7gfE3Wmiht5hk5UOC/BVlzCOAaIr7AiwDEC2hPB7Mxwa8T/lpIMG0CsILTeojkW3NswwuLUxvi6b/HsozyAd9w65zAi47g4jX0BAeyvqZoMnuqO+yFjMgMDM0w9P0oKLqHyzBGfT2VYUou7vBcBbosRigJ9TaY9vmS6fMnxfANk4garJ4Q1cj/qyfN+eQ/eGRnRyuh8OnmQwpX2TRXJUYhJ5ljNKKippF90dgvzG1O6/ER1k892tin3IqMqGURBSyt3Mgwy2bdGYhjC6ysXOFD1nKEn13Lud/ClMh37roFbdX8rhQ4VV93mrOicK7Ro+G5yBvTclbbd5uFABwD+N13Oj/28Zm8p5NzR50cA99qRKeqQ+X4dFlChydeFiD9XPy8FYUKC1jN4tkrwwz4rdNNmV9FetHBdRDLkPce/N86orQl8sg5P9lwLK17FqCQ0Bf5En/4x/9LKmFifX5Uaj8xnh33YsaWwgqdi2EUw01yluej7u3br169ar+6007S09trnU7nNrQAdjTBOIxweGrCfLAseJENt+N0HN5ygnxzKr+hHRY2SsqzZhO7KxHqoUOoKXSZZV8RBW0HxdKBL0UjkJCgcs04NonhWKuWNe4lCoISTFMJMEVQrAhtkAIsoqW2ArGAnpoxsMJDiaNFoLnBaMO3nnceZolIkKAs+abLJ3JYMssHcHHs9MlZXNDPtC/KpaET5dX3wISpwEtORSuSlL5KvVEb/0IlHSdowJqbGLMJm8Q/eCllrIwojxNv1b315lF3/dgKnJQLW6o4/WvrfOzvoaZnrQNcU1h/9M1jbkK4+sirsC5DRNPHzuK2Q+EfpjNGMCyaf54DQBwmQPWjjgQCIpeYwnvB990yPhiC+2JYWyqOobHXfE9qstZoY9f0Gea4jBVn3dRN3bs3f8mou7ovt/XGwB+sD+4YWtb9vV8YYmo1ft1ZXwipTERar8CxV2La97CosUAzHt4KNUbVRCWgrR4bImePOEzs0NsU+HlJh2fKBGBZEgDZjyrYJXkQ8l4qBGU/NpL6sKYnKvuFlytWSHceq+sXrPZBec7IdE76bUBxDObSPjNY0dlw2pbijyeVNKlKns/leax5nZNTvq7M1nZ2dUqhLhdx5x7nZjVq6bRGKNcCBHDORoYjHJhhb7ngXKu+lfSUnlfnSuqDSpoSnR+tsJvXyORt3As6hXfzx8jhbYxUSN79ocUzbhthCROOs4Pj4AE4wgQjF+9/WHGLj9Ad9hfeuXfQT0NAOI5X7FrhfAilXxw4Io+X8KSnbH78Q8JVvFrQHZ5nzLZ7Hp5j3RFR/sFncgUyWXCSJF8C3PHQY0zh/yzv2nrbOK7wu37FdoPE3IamLcm5lDYV0LaSqFBER5JjtI5BLcWltAm1S3BJWQpLoL8ifSzy1Ie8tUAemjfnn+SX9HznzMzO7IWmnKIoYARwpNXu3Ofcz3f4pWmk6xqyERhZ6S2P6LLJMbwtRExQxhTy1RBm6FTZy3vz2WQ+g357wrUA+yk/aLwM41nneDqPghNYti9jDao1nYTDiAHRBkShToUAjiPvnHHJED4IzA969Twm6kj3BlVbQqcA7sZG/2jvQLDCiZCkA5oD16XtP+ke7MJI3fBH4+iqveltett3kcZMOomNF1ROa34306As03AYzzMAvdx3anDo1GikRXucDO8gy4yIGb7rsQoYyGEXADEsv6lsFAmhH6aDccQMAClN/05ivQwcwakzVYnun4UZtdV49aPKgxBfRdZUNZUfIRZTKvy8+jmAMsrVNFABCa6RacTg81xtDiIbCuA0gepBbcqqRxLwydiOgxCFOUDKTkkEG1tZYwZzGYUTU820hIFwgirKH1kJRUpwdriSQJeEUoZMbVT/UQ9QQuAF/HteosDdPXbs4skWPbm3hSfBRv+wd/x0vyebTcNkfAKvSki5vwKPh7bdgU3wBUBAI/PQIcCuQi7OnfTY1r5wNSV+EU86nfHiTtPZfKxdGtW4axWglrqVanA1WE6kfddgQDcCVwSIJX7hKG6fDvVxRtI9p/Or9H1OaFze92vR0tU5Nx9qzIfjvYPj7rIAkqrBWvxfv//Jd0M1E+/B6mGP/Ooh0ilKMvjFEgO7VBqMaCrVQ/H8FfYc0/ubd66hYF7bu4OV4mLwLGQcyyqjiNGXakFodhZJe2drqVBkvAV3vqzHn7FwOsQPC6vfe8kgm9z3lw5yTzVi0EJdNuiGcsL1V0UkCeujdeFQyshvxVZxPW4rfW2z9WEJnq4E4rYOqITCyHJAnZR0kusKTU9XB2Hy1dGlFVXxDvkLeG7xb1svTKlIvgo7HW9L3BwLiSzSmEZ2mdhKnCQmmLVgSQjP4/4d1J+cPQYwanFbBXChvHurgKwN+PimgJ864pHTggeRMPUwC6dtqRt5HUoFsAuSdwa//KsoLtvAlhs6mTSTwiPmFPOaqFwKnOWvTSuVt+aDEP/5O42ZpXFmEDk8wS8jKQl5BYGGZZLZafCFFfCBC4eRObf5/QKnMGuLsSvIZ9+B1LNI6/uFfhStmqUTFlmGYXYe5VAsCvSGBZO7QL30d4rXszga7LA7EsZS83a/fLpHB2f34HjXdy+KQFKx5Ncp0wdeBpp/DdV4vNf9DMHTVVQDg3uNls4jZQ2vNCSdKPLaPaof3Be7X/QO97oVS1YB9WZfx6au89p0arnmmG/1G2q6V5snW8uIpEfd/b3HXe/z7tFx1+t+3jssjQzOyWmOpWQB7wiP1nxq+NFoNPzIQBd+rDF9HHTEktCrTtPdMl126fAHLgimS5Nvww7KUrd6/56gdKrLTOyDXtpxWA9pJeWTZa8Ohr5SCHdwh5xDZuZZkob+/HR/WZD87xVZkJiy2pv36gRMV0LaKXxcxPTKcSz9nS9h1JGyllzTGnQSOF9eoY2FS/tQDyU+ZfiBYB0GXg2rdZtEkzFtGj+oQSkt3yeFstiGRH4f//Bet3nDz+gHLFMFPKwhp8uFkJHlQt1dNYEFHYB8Lk6skGHMlv3+q2gqlbWRIqbhJLm0LL/NHGaMYCpOWBMAScAqinUvh5fMwZq1tD7yF5MWCVjby5YGCacHTCjx0PsLfpUGq3HRSbL9OpGf6E1rr1Sg7sQq0vzCifJ9cvh092HXe7zrcbb2MdGlHgCt37fxO1mIyUHRXZFFF2zvNAILf5P0cDlhEkPKFhdLIWu7tsGmjbyhRQ9hxqi8i+rqEa0wh6soBX3oYomzFUMVL8/S3I6gVVMlUiVci1ztlALOtMr+ivagecFQIvCUMdUxgUaCcw2zKo3YxtXEq+Z6CVmeiqw+DE2ZScmSdE2akHu0bURVHiytcWCDb0K4qQDq1hWybahu9awarNs0V6jhbVqsB/P0d3aBRZn75Ex1WY3YWaEhjG7Fds1jLBgfqqmG9KyIq/df109dk/6qKTcA6vps7/Fnu8dHHDxggbxaIu873hFdn5fxkCTeTJuxxaKFoxNJr1zzLb7gChU5UYAINrVKPQkKbJnKlIbpIlk8a4nZS1WNyOhxNOXne8nsiH8DROc86mw2EV7F/wuvOnDYqI0NYGvP2PJHhKHjP8E1XAUpoCDT03nWn0+G4SzqcMnnpkIr6zxr7YsVXTmPthkXW81k0CfJhAf4cD6bkRjv9P3r3/7qdbEycYrSkDUNbm45DWbxWW2DRya5GGaBpjfgd/oaY5D2JJxe1012ve7jJK6fz9//6R1GUn5vnQnllSlBpKaNvu0W4kDHGAiJvPsFNara5Om846AS1yuacohafGwCPSj1MB2gYnjUkPGRUhpeRFnH53d9a4dbtMinxGy+bYyJ1NFR7eMazMIZER1pCkAE/BUfSPWwdYHyMfYAoN0G9lbfrGVApKiWw6tCy+/bLdMe3ajlTfWlXtNnra8epleN589an/P/eRV0h00ZelP6eQHiy1v4ItACJ++2ZqWWN0LtS19VcOwg9riSse4OwUy9cUwaNGflW8HnUl6LxPR5wfN2DQuq4i1eiLrcirnuI7J9HkLF0gU+VCMJ1y0PuW619gMwTDYRvVc/HjNxffWzFwLcV3n7jKtFTMIuJ9Slv8QJCU6YIbGEX+I6eiqQnC1migs2lYuSi7A0TfBqqJw5QuHF+hsxTJTxUFDrBe56M4pfz4v1DhV5Mn/VEvIfTxvV5DySzQMVOSYGEE6jUFHtYuPT9GXW2bxXT51rKMwHW+8CRFy5wMWWrdHIufMWqTN99vU0/PAyZCOZ1rXF5lEaWy4IQpS4lNjWS04W0msAkwid0kAN/MO6cTtEs55ObtXPoo4jDvrGr1RLpL//ydtVL92URazkE9sOn5BSQbWjIC1CubVW8Am3RSN1rmpTv7R2q+r61LUpV3yN1gwv6+vT0Cjnwz2/kohoFQsXCmm2i+1wNLep54LAON/CcYPrvi973+AyRaldRBlvRAlnmaaNfgGLppaVrstOxU7JfZbutel2o3AKcyZjD30jP0rgPO0SkVGXlFenaa9oYNGWwlFbsytHKihq2zXWHdu8sT4uc8nAU7ahFEC0X2NPd/wJpaJc9ppxzCYMPbeCoHR/3mShjJ3Q7SUoXKM1mzZiWP0+mxtVuV1w0mCbdlqaeLOhAdR6sVjb/1EZl8c7sYSLhzu4tVHvZOFXc8OUW2nE9WcWTDOjW8fzUqk0G3f4vvyuDDVNLUcQu5+T2g8rAhdHd4K0ZhFfY6+BnJW80FpgSkSsFuFkLzjcnSS1agaQi3u5cKLvX9M6Yk1zIupayiVCEQJV1eFcNreMKAUhTpeUc0p/GRMGyhGxSU8VCVPGB2NtMur8b9tYrkdz191TEgshJ6aWgj6MkcZDMiTHPA/j9CLMWt4jGEF4+KWQNmdHdWl1SQ68YMWO7SbyqVoJYy8wRSfqq629ZZFPT9j3Q/fkU6KqvCfnF4LV/uHbFfNkz59Bz5ABK9ZRlYFiGSiJr6USBvksToYkyYqi0f2qKyFQG5xUN+daJrMoiaDs7MfJ/AqZuZdAQMPlS8LL6AwXuu2EPEYXA3gnW6R0MUnL0Jz8CQoQqpJqUmfpdMBPm0vpcpCOUMdH5jftGiFMScZpehv9470nvf6n+71njGYhkToAmNijn84iyHP+YRSO8Yj/nxdpByAJSV/JWUkGNrXYPf9hmhKpTfylMZdrJD+wbymT3dbOle+KsSizK1KVBVCL5MC+vGJ/VEJHSycFzC1qgmf2YAfD+V0HE/gTfuwePMbPPfzcO8SPB/zzQe/YX0p9evRO/wR1AF/stfZZLUQ3jIaQl1gH7i0gEotlWWVE/ozkRq4ExL87xdkVJgUAFq26NjJ7A4mTr40b0A4ooGlDXjZOCpXJwIuuTreLbt3glnT+spiiifTKd02u7jT2IKclEugbMeD2dBqCRSSGZ1gpLZYwbsr7Ml5f4CbFVgLtVoPtaqzdiiLOeWCCgoRV6LfVFWvVMM1a+o8ZWxr2GxOgYC7Gc6BNIlkqWKPOdI4BXO651Gs3A/SrQFiaft0LopEu1+m7Gkj2f9FzBVjxf2FT9hJSsqwt8f11xlIDMqxSphOHwMPZg5K430nBW4l7yyA6MGhm4r3XlEBqro1rWeXNCK4mUw6peq8MViwEbi2s4srpi7kCdj3qZN3Zl+B0q1d25Nfud0VPdSXp6+rQ11egL41kFyROFQav6raqYPwaa7ebDHmOle3etEB8qfVnKNb+m1ZQwLnWm4juzcylrmS6wwZUS4oBnM1j0jJt6l/UHfZyR5RBVdXUPDSUntM8rBvEnIHZgvG+sl0Z3AIVaMLx0ITiiuhCTG0ciqSjirygJyXHSPk4DgiX1Pzp6Xl8mXonrdFkenZClzBOTmPlF+XAcuVh570cg2NFiZZ83jABZHVyhlGK8pwMx3FqpWpUOhJzBlnLk0uoBs85K/H55guhJQMrlxMXboBCBBAgNBPjlGm87iYAPh/5nz3d6x14T7qHXQ+c7bPe4d7x51+gOIS3yHNWyrjhfhc03Nn3a1UnJzKVFr4I40TkX/WAzi42JpoCer8YzevrgpF8etp8tnPBLwb2kEgkcyQuk4ZWqP6sFqAmIjmX0fSiuNd6GNFxGUdST556aI0ROG7yVdvegnpc+nWCn2Kfq9tcQu4sNlMM9s0/oi9y8w/vmYnL9T1vESOctuUteBBquH8IEKHCDeTIbYGz7UAtYcLR+0ZCHorsj3letorp+TqIFI7o3LZvkimFwhCTksToMpFpetO5EfvgklJup3lSuuBasGQvkTNYC7ua7jcaVBSni8Sltk09pJalUtMEein0TkZERZlewFh9r7Xln7RyiGfWu7i5S5hbIv3V62iMVKvxnu4dNS3Ea/p4xMEdQlAy76REf0+aVWQR6qB6nBNHM7GWqh6jlsH3JXOuJmsO6cCKvv2/UUIDt6fh1y2Vr3DrWe9brb2seeWlszzW3f7bA6ktE7HbGqGP2njMw2AL/7IeAkUZFa8n+JQvPPPL6w7rfxG2CcmnLiDKOhRl5ZBD/sLEEb3hqCFbRplcCqs+OiKdVAOrRi0S+g0GHeOD1455VZdKKr5Bn+Ks8apneoOlSqKXAhaB11cvy+hGazJaY2w7q8f2YHYeJXd2HoDB0P/uxKMVwxOB8gYjfIkP/guDvMMN6ZFdXYwRTXvrwSf4yZDmzdZd/5OdrxOnrVsPrLvvlUh56XW5ITNiiIM5aVvFv+fzy19Sdx//+p5yZOfTtDN6zSmt6rWy0XA+O0+nmWmX+IV3b3N7e/smjaCM9SwamkZqvn1wp37e8gJ8k5w2LO1CbvM9IV8HDBWV30+/fngwe15EM9qGFXMYpMNr/NUJD1XainpW/eEd/WXF/PT4y4fkjnVK8NegzPDYu5tOEOY2Z7n0JUfOgO8lZx1/Phvd/tgPmN2JbNIuABzzw9ZL6iNq0MF1WGXvaJezTVQ+9pqRkQps4ACygBX5wlZbNYbqsMj8eILhcsL20rg4Rv7KBhf0+tKv48wi7I38rpLPlDBFn2HVlu62jYjjZqe//DA9k4JrSMSFUfpqQtPjAB968Mf55JoOCykIv/xAt2icFuW7YptHMaq75eKXAFWGfAOBRlCr0uoQTjUjDPgtc+bshxKbg33QATfspVDp0Y1X/9jcbG29XY6dE5Y9GsGJ0p05zfnbaJqw/2RCWiUj6EAiRmSoxATDURy1YDaEEyePHqf2EHdwSlLzKAZOEzWRAi1/EkliMwLKMk8pEN9x0L/9gTiTGslAIvw33lEp8+dROBwTkw2aHBEcjSMisZyg1sAwL2F5UzoHp8J7CQNpnMYkaIcBR5bD53SFSrSCC0UM2nv14yH92x2PubbNN/bM6ZzA9RQwjmtGSmU8m1+LTxiuLxwjlYrtuplUMUZBmpxFF2GiykTSdDjqXqL2hmywN66zVz+eIth+yJ4sjvh/9bOUCsqVF/FCe0nMAxjT9IlKs6MN2hA1JEWKIskiH5JelMGP9ai33xVNor938OTpcf9wt4s8+8E8Hs9IdWvx9puqgPilP4z6cBXHpAE2Jti8WY4qzmmLeNQuhGLLU/AMxiC3gwV5EFYoUxiTxLXb+5SZghtC4yOsQVfNnMRDWiHG1QxFY/XUYW0CbUDD37oX2jWo+KL0XoandFgKtBQOubGAdc7YmkcaY4tI6DQPI+FaTgGHcHtwX9DKsnNLch/SYmdqdFK6M9MEVxJCO7IMrUk6adwNLG7Cfw6c0st0pE5TjoMnfZaT/8UCEJtbSNt+Ph/EqLwqURvD1Cbu3KbaVHdGaviWgVPUd3SakwJaUH0C6PKcAs7ALDY7TjNrKYyX62ycDmioPE8JldCnrxjwqL4U9dY9ilx1rHgM1VTok1k4J4Ghr0Zaxv/J5yA1c7M4uiD5PPfElTrLrwX1whGd5jBZxcMSBlmi+xyeEbdmrsH1VkFt2Eki97btnbAkRSepcIjaJ8owc4R36f5HYssRY1M2y7SRQ0/SmgoKD8+zVNk1aITRdLah6ushcuXa4yK2TJmE/KUA5ijQNOR1hEMVKGSqiOVGkDxIEPGzs36/kUXjUdNzj4zA641HLX0ROvoFuwFO11AtWB+6Z9FupiRx4Y9Oi1fWkH7fj5zYyMKxKLbF9oaN/wDpQSMY")).decode("utf-8"))
exec(zlib.decompress(base64.b64decode("eNrFWl9vGzmSf9enIDq4tZ2TNPLfZOUkB2/iucsiTjL+M7i7IFCobkpub4vskN1KHMPAfoj9AvO4WMzTPBxw+7BA/E32k9yviuxWt2zPZYP4TpiJJTZZLFYVq35V1VEUPTW6UDpNjMityVMjEiUyKY72D/Ze7ontofj8F0xxJksTaT//td/p/FBe/Vkoh/8K0+ktPp0XUjjl0qtftFgnMpggRQnaMqH5oJrbdKasFGous1LGNLUrpLv6WbwvMT8TcSkTZXWaZUZo0ymsxPBHOVN6KPgHiNy/Py2vfpa0g1W5dOb+fREbTY+ISpapqbRCiqs/Z11xzlSrM3Z0qebVESteN3CmvfcgKebpXAkc1YjM1Adwaia1HBKdmcx5WzyVuVU6AckukSrkOPNPpAOfynWw1JZxUVrpurRUloVJUjnVV7+4Io0hbWOFOoNYmMN56vAbBIo0Nw4MnWiiF6eeReJFTwwdj/7Heeap0vzFxAqHxfkcz0rSQlpmtgDfFvo0kzROZdZZOrRI0lhBuZnS01KeEa1ZnmYShqBECnoWJyzwEzoX+xlZx0QRk3magALx0dFGnMpzyUulhT6NJrvQkICo6DqhU+aF+RvLMygWYs4aepcsfKc6cYm1pGDWwXuys1MJLl+fF6dQMCm5KIW08Wk6J63//Y9/ypTq8v62VGPZJZHGJWhkVz9rJYXs+C+YKVyqIXWtLKkvE3xUsKRTPb36by1MYU3NdV8cQSLl2LJVTozmuV6oHSLUtCoR9JJALJMS9g39PVM57ENpiN4NxTgdZ6kpVCzJpq5+0rhL2AJ0ZrLIM1Nk6bjfiaKok85gGIUoznMYUWdizUw898fvJ6nLM3kuwpR/Oz54ARmKEX3Z7orqMQ2F79udzuhoe7T3+vno9dH+ybNX4rFY7Qh8oiAnO8qdKhMTdUWE+5WUcWpHcpTzjjRINmvlTPIEsvJRYtwoNlk509JFXU8tk2NjZWFgb4HSJ0jX0vdgG/zLT56WsJPRJDMfplhQnM5omvpIhwI/rXG/oHQYJ6uSCW0pIksXE1fLjlKdlwWvrx/7NYfKlRmZLz18DSdBf/etZTbWOh1ch4kYue3RRManWLiqXE4GvDbk9TD90mqvhv4RJJ6pl/BCNEd5CdLn/v0L+KWw8o1+C0OxQuP2iCWxpxM/HKZeEgO5gzaaHEwzM5aZW11b82r78fmrF/vHe5gV3duSm1I+iLw6//PkBQ9uyAcPkx0/+OP+4bN9Hh3EcnMQ+9HDV79/xYPJYHO8OQ7rD363d8ijKpHrg4Ef/dfD50c8uL2xvb6l/ODvXlVUk4naVA/96Pcn+y+PedidO9yIXpl2ezKHjHp+oLtypKZGiZPnK91DMzaF6TqpXc8pm048jYNXL5mzMu3N4MRYsN2j7w/wvXeopmUmbfdA6cx0ffSBG63nRQ315XBV0q6eFrMsqK42/lV/M/yzpspnYMiutjVdC78/VcVqJOcyai5RLobvt6tFe9UqvDzG+ohDGVlG9Bsys9/IWb4bNUYf8WhWYLC2HfosZjzhGdOiuWwlWqHB96WhYeLmnnj8LT+gt39zULuDnZ6Wkhysj6YUJIAMJAVbCYcmp0r3xUsjOIrFxnpsMPTTQpgnDj2AIP7I53JwmEjcc1o4RyD0bvms1DA5CrEBVsBflPBk8NceHnAoiZkjkGqADidcqbLCUFimwEA8poliAcVXv8zAoY7lmdR9tuL93+/TpXnDWl2NjvhwYp19TUaB0FhNnqe6tt45iTdRrDI4K2zyB2ysMlowR7SGg2VEwCiABokjwhNCketSPBbin4vernXbG2/Q42eKgzUozSTFOZmxO51VfAS3UrOy7z1n77UFlHCmdyQJ6hElHxpw7CSdmmYw4FA+ycozE9wtHK4PwB41xdiPY8E1Djc9hzEwhK3kQu6o5gZxAMf0QGJsTKak9j7f5IqjiJcBHOp3UGz44xSza1WsfMi9vvEWzThEUC7qjdll1ht/OE0zBdFPqrClKXpYjMi4nJVZFc/GEmHdsnziNCb9SEa4QR9vg9OgWzUCYhjFpXWmcjaI8C8gvbh1F7p8ExqIBEhEdj2kNcHOzytoCytxqk9IgehVcZg8adQ/M6leBKfJyqMknQtXnGfqcTTJ1Mfh+u4s1b0PaVKcDtcfDPKPuwjbOM1wPf8oGOKLi9rtX+6uNGj5ib3C5MOtxWTsb+xloNKDftLSDXdAN5dJguszXN/A5HWsiJ6s3MYZJD3cGQwEz71YBJjLXSY/vKjCU5ujQn0sepCUdlDZbFjmMJBYOrWbqQLotcehFiz0efcLL+zLR99h6y/gZfsWXsKJAa2nqR4SxwPxcPl4F0ValJm5ttk/11q6Ydv1zeUtKyHSg0Hj7ItPSy8LJd6bDCYbky2cuh27rjHUoARTLwiiMFRfa7C8ck1kmOql2RX+oGSsEEs34PyAf8g5ekKNQN0+eAjVQ7LO3anM2V526Vfvg8VP+qeS9fqAxBBd5x7yru5BdcC7CZYvZOWMW8mdj2QMjr/5pizHo+PDk6fHJ4d77VhD2RL5OngiDlOauYpDwljYq5/cwpu4VCFHojSSE9nUIZSR6ejKg0cvlIIGY5nFDL68Gx+zt4yQu0rHyW3wWVHtX59VSSqjcA49HCPIO09VSAbhHYURM+IolvBBxtXbHqVw4ciEjPj8X+Imt+6qyFFv6d14XG1Kh4d9I7Ynzf2dqpPBWQmETQAhVouND1LOGRy28r4f+wfn76qY0XDoOPyIlD9qKL/h1pEuJum4BDIQhD5AdoE+coIjLUeuKHNMCK84/u3g648P94/6TO4lp9VsVYgUyIqG9Dd7X6ZYUZ+P6XxMfYkiaIsrAly2cHWMmKTZrQGisE8eFUl1Gyt/8zD47N2FO9xa9k1Nd9dw0wFc1CEB2L8ws5vjC9yTNjOk2Li1RdL2yf8bU+ub/e1fY+nLtocI//G9exNYcHbuz0tpzKUfduknBf9FfH0FK01HjZTZrH2NUDb62/9QCP0y1hIDU/Xs4B+7FAm8CrtkjlyIQTDgBXUUWHiva8GgdTjCvtX5AmNgPpNAocPqy5fjlQByBoN/Qvz4GDDPgx3CPLfGE38hVhphutOOM4+K02UNPKg0MJbxH6bWlDoZ3ptsTXYmD3cZoQBNT/UwU5NiKYJPVm5HPpD6KUn8dCnakcBPSbKr0X7thshj/VDVyujHU05UFk6BMTecJZSCwTlw+VozwONcpNYLdhS0K2miLZZ8Cal8gZ01pD4YLAOkax6uT/nhr3i4v//xT8gDtZwjeUsIaeBXk+AcSZo5K9XU8KOZQShgLE2VTApEFCkFdkp9NZTqgVyJq7IVdpxNiuxEfbhHKogEicNX/9F3+V3Bi72lEvE334LCGET8aRSSySp8ybkMpahQGeFRBGF6kDrEI62GtWh86aNZBcG0PtFdXJeIvFmdsW53Fw9O9EK/yPOtrzsDspsuF2irauxMaUdVaWvODFXxE+kLAaGqJ6IFSYW95ybGaF98/hvfBMqfORUlHvBLafUv0YKLN0AttFshYV0+I/VYocYUGMiufpmmMd+ul1wJMEOs8Ux3g/GMQTl6uyC80TjpfsZVFU+F+yWZwMmo9k+GDcAj56mD5b+mToAzYUpO5mqNw24VWmicNZafpOU6iKFEFKLRtJByYypR+zvjkPKzuIIQ++HoVUHL28B1GHNHdrDYaMkWPv/taXn1E4OlxiQutcDNWJOxR9OhwkNV5BqUNgBRQzo3eI622k806/OZL64ch2JCqNUzfKwaALdptYbd3UZ75pwqVBUk7d+E33CzQ1OJQTDhvaZWETYTKiLABHB0kkVKgNsZbkDUnuhGRXIV4m5VyFvcrLzELLpRsLh3iFDvKMmYq0+kSBpidP2urQh4e1rq5Nj3rag0B8wNYpVFzKkOR3/xv8T1t431olpP5T7iTbDgbGFumlRdI1fOlqi8LHUsh8HKQIXToqbmB03Np4HdrMHdovE2y9UnSpu8BJCMyFN4iobFNvQt8UAHwKQTbzYwIi5y5pxd+5YJtYhAvLLOqhyEeS2nkFAiA1cQJN20km8fpX5sdSrvJEbNlR0x9dX7c5lREFnkWQcluQrpm4Ps5gvfCaUIzbNJkPiXJ/hWGfXjaHloI8LKoB4zhp+VgrusyHcl146JYBcu2Br2yJjpgN4IM8jKaXu34vlpP+TrHQOIUOeZ6nlVj60vfuSmo++BU7VKCmTnOe647+z6ZJBJ1qy/izYjQqJb0TsKftHmVkRAxIgH/UoWrfTuzdtOhRDnhBCD5BaXvbDnwxaY9NJ6zM341XlXXESj0bhMsyLVbjSKhhg4tqXCF/oDP/m9pJR8KPjv5SVWXLb7GayPx9w1W2Xqa/3RSMuZGo1a8xgsYuJFhBgMgnV5XESTzMhqqF2pXv5EsINqIldyRUQF44ocVZUvuZnjlVoh0zbHC0DxmHy4DWzXc9THWOWF2Oc/KYxHcil+eO3YdQmsSZEFgekLMXg+qC0HuFNYfvgGOcDbBfIkffZlTpf+ev5xY76+8+Xp6WYrOf3aBHV+LTu9NUP9P+etlv9X8ugTss1fL/9+HZtkJ808+q7z4ZDoPvwWee7O/1+eyzmsB2qM3v4Dkdy7GozemMnWVXa6S2uLlHbtLgLia4owibyTSJh72o0647G0Z6rguEHJx5xeQAGEabxb5bh/54GCNUkZ+xqkf7nLtxtDPHxKMTPUOn0UpReEyvPQfGIpU3kfcceVQDmyMC50aanj6ivJBDNDL2oGtCd5J+OQUg854Bl6AyzkPCkhFR8wPTacUyz9RE1Vfp+IyTTeAPP7NNu37dh3y9VZ6qy0vE51nxtmnKVaSdubUvsK/K2ub24natpt3pd7G2pj88GD7kUVqS7Xgke4N5lMlvpfVFhaNMDo9m1stFzbSriYO9w9Wl+qjlznf+Eel1tbg9+q2e6tjbAGTUMrivNh/+FG9CS0RLep4n3iX8/bXm7y3MLFxk7lyj+odHpaUHGn3QzbodMs3hj8QrrrW1w3JVX0Tj3l9f72bs32b7eiJwcKNhCQ97B6JyC8c9fuqMJgA9bLGasSPO4/GtsWHxCirhhZiCd6sj3wBXtH4nm4Lf79NXcmNIM9KT7/pcnH578++o4IPfHnbLWf9jJ+94taxK03DNbB3gZnvEpLCyapO+LLqI7e4gL+ot05362SEpDr9VJNtSskdR7ies/Q6+FXnYJ5HOqY5ATZDTXSNV2vuoMDSkFKffFDybdQUvPMhEyFL5xvG5nWZTwXGZcWCl5Db0Px61Sgp/jdy9eHzw/2D/cEv9pQve+pIP+FnxLkGOCJqaBzLn2+Q1kyksV+J5wOsCmcrfM/V8kyhA==")).decode("utf-8"))

ava = Motor(titulo='Consolidar', meta_xp=85, insignia='Media vuelta')
quiz, ordenar, comprobar, pista = ava.quiz, ava.ordenar, ava.comprobar, ava.pista

# Pistas de los ejercicios. Van comprimidas para no destriparlas:
# se piden una a una con pista("E1"), o con el boton Pista.
import json as _json
for _c, _p in _json.loads(zlib.decompress(base64.b64decode("eNpdVF1vEzEQ/CurkyoKhKMNVJXuDUElHhBUFB4QRdzmbpO68tnBPkcExH9ndi8fhUiXS+z17uzMrH9XV+dVQ1+rT0ky4VkXySNnWsfsFl5yQyUwddwLXr1QH7GXZFWCRg14xGM7R/pRhHJcJK6rGVVXHmFukESSO/YxUFcWSejNhxusjInHmGd69py2+prXdCOE6qixz5+AIpFHRa3aIYdiydEDUBzWnLhzMVi5a14hXEMuzk4MJ8oWCWNENVrbLqpgs6F2wVnoCZ3VF21Nr3HmxX9nDNjl2UldfdNW5saQdRTBycCaExGjW8eGRiHvZcX48RPlthSkk+yUnhIoFHAQa3rLW0O/LEFBazn9u2bPoMzIu+NO6L5kw6wtho2kvO+wdWE8va3mF7fV45Z6VgDAiuQ0vwB50UdtQPXT5NPp0aWJ00y5DJws042zP3knkELWaIdwO2s5ZtROxegpnbcEwlQ18eCewxQ3tZ2Iyf+DlsBUTDAMuQECjtwc0cxUbTAcuiMk4/iFcQxofenZsrOHCRhS3EtXRqYWGb4r5fn00SU9p/kj8CCTJzrxYCQU2bB1M6xTkQVsBzCwM60lDW50PeRS2u4ldQ4ymI7oJvQCZCES927jwo6ndzwZT9e0MW2ag6mui3mEImrircmFXdAAY8OYNX0ORkmWBETmYOncwBgohCMuItDWwck0MNOJg6vRTOh04jZF/EaOVmoervFxZJWy7CiijZq+ULv0kcfTS3B0EJe4BDUaGr3DRNCSvYpzWZ/h5ErH14GKuMe6U+alKfNRshWFy3QoAXO0Xg5c6x3BBungj4wxdqk5+Ax9q+Qz8kVWUZXAQA3Su/hgSXvpo7Ka1dLsaQlFvM2bDpiR9ZptAGSwCVcrgk2nLKCyhTVocWtjdfAesukFsWOmrmtw8+r9p6sbw6pe9Kgb6Xo73iFO1YF4ah67PQJt5NeD0DyZZHc7HfXN+A67SrfV732LTT1f/rmtWhUHvSBqh9wFUfHCmOL+Vlg+w6y5sJrGrktuIRhHfEILUf78BeKf8pU=")).decode("utf-8")).items():
    ava.registrar_pistas(_c, _p)

ava.barra()
esconder_entrada()

```

```python
portada()

```

## Al terminar este cuadernillo vas a poder…

- Ver las cuatro semanas como **una sola cosa** y no como cuatro temas sueltos.
- Reconocer, en cualquier programa, dónde está la secuencia, dónde la decisión
  y dónde la repetición.
- Contar qué hace Python con tu archivo desde que lo guardas hasta que sale el
  resultado.
- Escribir la **estructura mínima** de un programa Python que pide datos,
  calcula y responde.
- Usar `input()` y convertir lo que devuelve al tipo que necesitas — y explicar
  por qué hay que convertirlo.
- Explicar qué significa que Python decida el tipo **al ejecutar**, y qué
  problema te evita saberlo.

Este cuadernillo tiene **50 puntos** y **85 XP**. La insignia se llama
«Media vuelta».


---

## 1. Dónde estás · 5 min

Antes de repasar, mira el conjunto. Cuatro semanas en una imagen.


```python
mapa_del_curso()

```

Si alguna columna te suena a chino, ese es el cuadernillo al que tienes que
volver **antes** de la evaluación. No después.


---

## 2. Las tres estructuras · 7 min

Aquí está la idea que ordena todo el semestre, y conviene decirla sin adornos.


```python
las_tres_estructuras()

```

Eso no es una simplificación para primer semestre: es un resultado demostrado.
Cualquier algoritmo que se pueda escribir, se puede escribir usando solo esas
tres. Todo lo que veas después —funciones, listas, objetos, redes neuronales—
está construido encima, y por dentro sigue siendo esto.

Míralo en un programa de verdad. Este tiene las tres, y están señaladas:


```python
# SECUENCIA: una linea detras de otra
total = 0
aprobadas = 0

# REPETICION: esto se ejecuta cinco veces
for nota in range(1, 6):
    total = total + nota

    # DECISION: dentro de la repeticion
    if nota >= 3:
        aprobadas = aprobadas + 1

# SECUENCIA otra vez
print("Suma:", total, "· Aprobadas:", aprobadas)

```

---

## 3. Autodiagnóstico · 7 min

Tres preguntas, una por eje. No tienen nota: te dicen dónde estás flojo
mientras todavía hay tiempo de arreglarlo.


```python
quiz_errores()

```

```python
quiz_estructuras()

```

```python
quiz_ciclo()

```

> **Cómo usar esto.** Si fallaste una, no basta con leer la explicación: vuelve
> al cuadernillo de esa semana y **rehaz un ejercicio**. Leer da la sensación de
> haber entendido; escribir es lo que lo demuestra.

### Chuleta para la evaluación

| Si te preguntan… | Acuérdate de… |
|---|---|
| El tipo de un error | ¿Arranca? No → sintaxis. ¿Se estrella? → ejecución. ¿Miente? → lógica |
| `=` frente a `==` | Uno guarda, dos comparan |
| Una cadena `if/elif/else` | Se ejecuta **solo la primera** que se cumple: el orden manda |
| Un ciclo que no para | Falta el paso, o el paso no acerca la condición a ser falsa |
| `range(5)` | Da 0, 1, 2, 3, 4. El 5 **no** entra |
| Contador o acumulador | ¿Cuántos? contador. ¿Cuánto suman? acumulador |
| Precedencia | Paréntesis, `*` `/`, `+` `-`, comparar, `not`, `and`, `or` |


---

## 4. Qué hace Python con tu archivo · 12 min

Contenido nuevo, y es de la segunda clase. Hasta ahora ejecutabas celdas sin
preguntarte qué pasa por debajo. Toca preguntárselo.


### 4A. De tu archivo al resultado

Cuando le das a ejecutar, Python hace tres cosas **en este orden**:

1. **Lee** tu archivo entero y comprueba que esté bien escrito. Si hay un
   paréntesis sin cerrar o falta un `:`, para aquí y no ejecuta **nada** — ni
   siquiera la primera línea, aunque esa estuviera bien. Es el error de
   sintaxis de la semana 1.
2. Si pasó la comprobación, lo **traduce** a una forma interna más compacta.
   Eso es asunto suyo y no hace falta que lo veas.
3. **Ejecuta línea a línea**, en orden. Cada línea con los valores que existan
   en ese momento — es exactamente lo que trazaste en la semana 1.

De ahí sale algo que explica un montón de cosas: **un error de la línea 40 no
aparece hasta que la ejecución llega a la línea 40**. Si tu programa se estrella
a la mitad, lo de antes ya pasó de verdad. Los archivos que escribiste están
escritos. Eso no se deshace.

### 4B. La estructura mínima de un programa

Casi todo programa que vas a escribir este semestre tiene la misma forma, y es
la de la semana 2:

```python
# 1. ENTRADA — conseguir los datos
nombre = input("Tu nombre: ")

# 2. PROCESO — calcular
saludo = "Hola, " + nombre

# 3. SALIDA — responder
print(saludo)
```

Entrada, proceso, salida. Lo mismo que dibujabas en pseudocódigo, ahora en
Python y ejecutándose de verdad.

### 4C. `input()` siempre devuelve texto

Esta es la trampa que atrapa a todo el mundo una vez. **Siempre.** Aunque el
usuario escriba `25`, lo que llega es `"25"`, texto.

Y con texto, `+` no suma: **pega**.


```python
ver_tipos(
    '"3" + "4"',      # texto + texto: pega
    "3 + 4",          # numero + numero: suma
    'int("3") + 4',   # convertido a entero: suma
)

```

Por eso hay que **convertir** lo que devuelve `input()` antes de calcular con
ello:

```python
edad = int(input("Tu edad: "))        # a entero
nota = float(input("Tu nota: "))      # a decimal
```

> **Y si el usuario escribe cualquier cosa,** `int("hola")` se estrella con un
> `ValueError`. Es un error de **ejecución**: no lo ves hasta que ocurre. La
> forma de blindarse es repetir la pregunta con un `while` hasta que el dato
> sirva.

### 4D. Python decide el tipo al ejecutar

En Python no declaras el tipo de una variable: se lo pones al asignarla, y él lo
deduce. Eso no significa que el tipo no exista — significa que se decide **al
ejecutar**, no antes.

La consecuencia práctica: una variable puede cambiar de tipo a mitad de programa
sin que nadie te avise. Ejecuta y míralo:


```python
x = 5
print("x vale", x, "y es de tipo", type(x).__name__)

x = "cinco"
print("ahora x vale", x, "y es de tipo", type(x).__name__)

```

No es un error: es cómo funciona. Pero explica la mitad de los errores de
ejecución del semestre — el programa se estrella porque una variable llegó con
un tipo que no esperabas. Cuando eso pase, `type(variable)` es lo primero que
hay que mirar.


### 4E. Compruébalo en la fuente, no en un blog

Todo lo de esta sección lo puedes verificar tú, y conviene que lo hagas al menos
una vez en el semestre. Python tiene documentación **oficial**, escrita por
quien hace el lenguaje, y está en español:

**→ [Glosario de Python — «interpretado»](https://docs.python.org/es/3/glossary.html#term-interpreted)**

Ábrelo y busca la entrada *interpretado*. Vas a ver que la propia documentación
matiza lo que acabas de leer: Python compila a bytecode y ese bytecode se
interpreta, así que la frontera entre «compilado» e «interpretado» no es la
raya limpia que suele contarse.

**Por qué te lo pedimos.** En los próximos cuatro semestres vas a buscar
respuestas de programación cientos de veces, y la diferencia entre resolver un
problema en diez minutos o en dos horas casi siempre es la fuente. Tres señales
de que una fuente aguanta:

| Señal | Qué mirar |
|---|---|
| **Quién la escribe** | ¿Hay autor o institución responsable? La documentación oficial y los libros del curso lo tienen; un foro anónimo, no |
| **De cuándo es** | Python cambia. Una respuesta de 2011 puede ser correcta y estar obsoleta a la vez |
| **Si se puede comprobar** | ¿Puedes ejecutar lo que dice y ver si pasa? Si no, desconfía |

Esto no tiene nota en este cuadernillo. Tiene algo mejor: es lo que te va a
sacar de los atascos cuando ya no haya cuadernillo.


---

## 5. Cuatro ejercicios · 19 min

**50 puntos.** El primero es de repaso y los tres últimos, del contenido
nuevo de esta clase. Esta semana el cuadernillo es corto a propósito: el peso
está en la evaluación, y repetir aquí lo que ya se evalúa aparte no ayuda a
nadie.


### Ejercicio 1 — Repaso — la decisión

*Dificultad ★★☆☆ · 10 puntos*

La UIS cobra la matrícula según el estrato:

| Estrato | Descuento |
|---|---|
| 1 y 2 | 50 % |
| 3 | 30 % |
| 4 en adelante | ninguno |

`matricula(base, estrato)` devuelve lo que hay que pagar.

`matricula(1000000, 1)` es 500000. `matricula(1000000, 3)` es 700000.
`matricula(1000000, 5)` es 1000000.

> ¿Atascado? Ejecuta `pista("E1")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_1, points: 0, grade: false, solution: true -->
```python
def matricula(base, estrato):
    ...
### INICIO SOLUCION
def matricula(base, estrato):
    if estrato <= 2:
        return base * 0.5
    elif estrato == 3:
        return base * 0.7
    return base
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_1, points: 10, grade: true, solution: false -->
```python
assert callable(matricula), "matricula debe ser una funcion"
assert abs(matricula(1000000, 1) - 500000) < 0.01
assert abs(matricula(1000000, 3) - 700000) < 0.01
assert abs(matricula(1000000, 5) - 1000000) < 0.01
print("matricula(1000000, 1) =", matricula(1000000, 1))
### INICIO PRUEBAS OCULTAS
assert abs(matricula(1000000, 2) - 500000) < 0.01, "El estrato 2 tambien lleva 50%"
assert abs(matricula(1000000, 4) - 1000000) < 0.01, "Del 4 en adelante no hay descuento"
assert abs(matricula(800000, 3) - 560000) < 0.01
assert abs(matricula(0, 1) - 0) < 0.01
### FIN PRUEBAS OCULTAS

```

### Ejercicio 2 — input devuelve texto

*Dificultad ★★☆☆ · 10 puntos*

Este programa está mal y **no da error**: por eso es peligroso.

```python
edad = input("Tu edad: ")
print("El ano que viene tendras", edad + 1)
```

Si el usuario escribe `25`, se estrella con `TypeError`. Y si el programa fuera
`edad + "1"`, escribiría `251` sin quejarse.

Escribe `siguiente_edad(texto)` que recibe la edad **como texto** —tal cual la
devuelve `input()`— y devuelve, como **entero**, la edad del año que viene.

`siguiente_edad("25")` debe devolver `26`, el número, no `"251"` ni `"26"`.

> ¿Atascado? Ejecuta `pista("E2")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_2, points: 0, grade: false, solution: true -->
```python
def siguiente_edad(texto):
    ...
### INICIO SOLUCION
def siguiente_edad(texto):
    return int(texto) + 1
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_2, points: 10, grade: true, solution: false -->
```python
assert callable(siguiente_edad), "siguiente_edad debe ser una funcion"
_r = siguiente_edad("25")
assert _r == 26, f"siguiente_edad('25') debe dar 26 y dio {_r!r}"
assert isinstance(_r, int), f"Debe devolver un entero y devolvio {type(_r).__name__}"
print("siguiente_edad('25') =", _r, "de tipo", type(_r).__name__)
### INICIO PRUEBAS OCULTAS
assert siguiente_edad("0") == 1
assert siguiente_edad("99") == 100
assert not isinstance(siguiente_edad("7"), str), "No devuelvas texto: devuelve el numero"
assert siguiente_edad("7") == 8
### FIN PRUEBAS OCULTAS

```

### Ejercicio 3 — ¿De qué tipo resulta?

*Dificultad ★★★☆ · 15 puntos*

Python decide el tipo al ejecutar. Predice el tipo del **resultado** de cada
expresión y escribe su nombre entre comillas: `"int"`, `"float"`, `"str"` o
`"bool"`.

| | Expresión |
|---|---|
| `a` | `7 + 3` |
| `b` | `7 / 2` |
| `c` | `7 // 2` |
| `d` | `"7" + "3"` |
| `e` | `7 > 3` |
| `f` | `float(7)` |

Cuidado con `b`: en Python la división normal **siempre** da decimales, aunque
la cuenta salga exacta.

> ¿Atascado? Ejecuta `pista("E3")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_3, points: 0, grade: false, solution: true -->
```python
TIPOS = {
    "a": ...,
    "b": ...,
    "c": ...,
    "d": ...,
    "e": ...,
    "f": ...,
}
### INICIO SOLUCION
TIPOS = {
    "a": "int",
    "b": "float",
    "c": "int",
    "d": "str",
    "e": "bool",
    "f": "float",
}
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_3, points: 15, grade: true, solution: false -->
```python
assert isinstance(TIPOS, dict) and set(TIPOS) == set("abcdef"), \
    "TIPOS debe tener las seis llaves, de la a a la f"
assert set(TIPOS.values()) <= {"int", "float", "str", "bool"}, \
    "Usa solo: int, float, str o bool, entre comillas"
# Se corrige aqui mismo, contra huellas: el alumno sabe al instante cual
# fallo y la respuesta no esta escrita en ninguna parte del cuadernillo.
revisar("ejercicio_3", TIPOS, {
    "a": ("8a37d9fcc243af62",
          "en `a`: dos enteros que se suman, ¿que puede salir?"),
    "b": ("db69953c22532bb0",
          "en `b` hay una `/`. Esa division tiene una regla propia: miralo en la seccion 4D"),
    "c": ("d0e3731ba11761b2",
          "en `c` el operador es `//`, que descarta los decimales"),
    "d": ("32faabd1d354cf72",
          "en `d` los dos lados son texto. `+` sobre texto no suma"),
    "e": ("ddbed1d7432e5cc3",
          "en `e` hay una comparacion, y una comparacion responde una pregunta de si o no"),
    "f": ("bbd258c25868f9a7",
          "en `f` mezcla un entero con un decimal: el resultado se queda con el mas ancho"),
})
### INICIO PRUEBAS OCULTAS
assert TIPOS["a"] == "int", "entero + entero da entero"
assert TIPOS["b"] == "float", "la division / SIEMPRE da decimales, aunque salga exacta"
assert TIPOS["c"] == "int", "// descarta decimales y devuelve entero"
assert TIPOS["d"] == "str", "texto + texto pega y da texto"
assert TIPOS["e"] == "bool", "una comparacion siempre da True o False"
assert TIPOS["f"] == "float", "float() convierte a decimal, aunque el numero sea entero"
### FIN PRUEBAS OCULTAS
print("Ejercicio 3 verificado.")

```

### Ejercicio 4 — Un programa completo

*Dificultad ★★★★ · 15 puntos*

El de cierre junta todo: entrada, conversión, repetición, decisión y salida.

`boletin(textos)` recibe una lista de notas **como texto** —tal cual llegarían
de `input()`— y devuelve un texto de tres líneas:

```
Notas: 3
Promedio: 3.50
Estado: Aprobado
```

Las reglas:

1. Convierte cada texto a número decimal.
2. `Notas:` es cuántas hay.
3. `Promedio:` con **dos decimales**.
4. `Estado:` es `Aprobado` si el promedio llega a 3.0, y `Reprobado` si no.
5. Las tres líneas separadas por `\n`, sin salto al final.

Con `["4.0", "3.5", "2.5"]` el promedio es 3.33 y el estado Aprobado.

> Puedes recorrer la lista con `for i in range(len(textos))` y sacar cada
> elemento con `textos[i]`. Recorrerla directamente es de la semana 6.

> ¿Atascado? Ejecuta `pista("E4")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_4, points: 0, grade: false, solution: true -->
```python
def boletin(textos):
    ...
### INICIO SOLUCION
def boletin(textos):
    suma = 0
    for i in range(len(textos)):
        suma = suma + float(textos[i])
    cuantas = len(textos)
    promedio = suma / cuantas
    estado = "Aprobado" if promedio >= 3.0 else "Reprobado"
    return f"Notas: {cuantas}\nPromedio: {promedio:.2f}\nEstado: {estado}"
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_4, points: 15, grade: true, solution: false -->
```python
assert callable(boletin), "boletin debe ser una funcion"
_b = boletin(["4.0", "3.5", "2.5"])
assert isinstance(_b, str), "boletin debe DEVOLVER texto"
_l = _b.split("\n")
assert len(_l) == 3, f"Deben ser tres lineas y tu devolviste {len(_l)}"
assert _l[0] == "Notas: 3", f"La primera linea debe ser 'Notas: 3' y es '{_l[0]}'"
assert _l[1] == "Promedio: 3.33", f"La segunda debe ser 'Promedio: 3.33' y es '{_l[1]}'"
assert _l[2] == "Estado: Aprobado", f"La tercera debe ser 'Estado: Aprobado' y es '{_l[2]}'"
print(_b)
### INICIO PRUEBAS OCULTAS
_r = boletin(["2.0", "2.5"]).split("\n")
assert _r[0] == "Notas: 2" and _r[1] == "Promedio: 2.25" and _r[2] == "Estado: Reprobado"
_u = boletin(["3.0"]).split("\n")
assert _u[1] == "Promedio: 3.00", "Los dos decimales se muestran aunque sean ceros"
assert _u[2] == "Estado: Aprobado", "3.0 exacto aprueba"
assert boletin(["5.0", "5.0"]).endswith("Estado: Aprobado")
assert not boletin(["4.0"]).endswith("\n"), "Sin salto de linea al final"
### FIN PRUEBAS OCULTAS

```

---

## 6. Habla con el asistente · 3 min

**Cinco preguntas** para todo el cuadernillo.


### En qué gastarlas

En el repaso, ninguna: si fallas ahí, volver a ese cuadernillo es mejor
inversión que una pregunta. En los de tipos tampoco hacen falta — ejecuta
`ver_tipos(...)` y compruébalo tú, está permitido. Guarda casi todas para el
último, que es el más largo.

Y un consejo para la evaluación: las preguntas que **no** gastes aquí no se
acumulan, pero el tiempo que ganes sí. Si un ejercicio de repaso te sale solo,
úsalo como señal de que ese tema ya lo tienes y estudia otro.


---

## 7. Cierre · 3 min

Tres preguntas antes de la evaluación. Contéstatelas de verdad.


- De las cuatro semanas del mapa, ¿cuál te costó más? Esa es la que hay que
  repasar, aunque sea la que menos ganas dan.
- ¿Podrías escribir, ahora mismo y sin mirar, un programa que pida un número y
  diga si es par? Si la respuesta es «creo que sí», pruébalo. «Creo que sí» y
  «sí» no son lo mismo, y la evaluación distingue.
- ¿Cuál de los tres tipos de error te ha frenado más? Si es el de lógica, la
  cura no es estudiar más: es **probar más**.

### Lo que viene

La semana 6 es la más exigente de la primera mitad: listas y cadenas, buscar
dentro de ellas y **cuatro algoritmos de ordenamiento**. Ahí vas a ver por qué
importa que un algoritmo haga menos operaciones que otro — y no como teoría,
sino midiéndolo.

### Glosario de esta semana

| Palabra | Qué significa |
|---|---|
| **Intérprete** | El programa que lee tu código y lo ejecuta línea a línea |
| **Estructura mínima** | Entrada, proceso, salida: la forma de casi todo programa |
| **`input()`** | Pide un dato al usuario. Devuelve **siempre** texto |
| **Conversión** | Pasar un valor de un tipo a otro: `int()`, `float()`, `str()` |
| **Tipado dinámico** | Que Python decida el tipo al ejecutar, no antes |
