# Decidir
### Semana 3 · Unidad 3 · Operadores y estructuras condicionales

Hasta ahora todos tus algoritmos hacían siempre lo mismo: leer, calcular,
escribir. Siempre en el mismo orden, pasara lo que pasara.

Esta semana tus programas empiezan a **decidir**. Y con eso dejan de ser una
receta y empiezan a parecerse a algo que piensa.

**Empieza ejecutando la celda de abajo.**


```python
# Motor del cuadernillo: barra de progreso, pistas y verificadores.
# No necesitas leer este codigo. Ejecutalo una vez (Shift+Enter) y sigue.
# Codigo fuente legible en el repositorio del curso:
#   notebook/cuadernillos/**/ava_motor.py
#   notebook/cuadernillos/**/pseudo_uis.py
#   notebook/cuadernillos/**/contenido.py
import base64, zlib
exec(zlib.decompress(base64.b64decode("eNrlXN1y20h2vudTdOByRHpI6MeWx6ZMVflHO+tZ2XIs7zgbr4tqAk2qJRDg4oeWzGJV8g55gcndVGqu5iJVk7vVm+RJ8p3TDRAASWnkmUpla7U7Egk0Tp8+v9853bDjOK+iNIpFcPWzr71I+EoEUSK8TPoqDnVAX3wViKffPXUbjX/Krn4QKhGX4i/06fURvjQ6Sz+NgyRVYsyEJzpMpQik8OREEvlJfPW9l2pPMt3SRF0Q1Z5K2hgSDQI1ljR1I4p9FcoYV3WS4pJKPBlEofQlRg5kHFuq0ShWSSSU0GGiR6GWiSteSBGrNI5koMcKfHj66qewocOx8rUEW5ci0SE4S1I9ijCDBLG/ZEqAMZWkGQaFWEiihExjNZVCiqEMAhlDFPfuHUoRRin9EokMFHEhIZYf3Xv3RH7vklaeKiwGbFz9KHlkyFLGUjwV+GaR4WAUkyBE84Q/9LXfxXNJ2ldnKva0p6P+65NW23CXQEf8KVAgduJlSRqN3bPkRFw2TuRE99XFJIpTd3J54opjLWQwuvoZsxIDaqEaXyXQSaw81RZVVUAsI5BPtAr9qAF56yEUBp244oDtQQ8U2ITKcsUa+8GXKI7BYUgEz0ENf0jNMsjGkNOluPohaEwyBQHECpqM9UDHgdwT4FeQ+kIyQU/jPmmYF+llUJ2CyN9gDNvdVE95jiwU46uf/CwgyqSGTAVpZO5II90V1nnTT+NpwW8SgTRUGwRqJMXJpyg+3yyJydWTy3Bw0iZGZOyd6mkkIGj4UZfYOZWXkPlPIfQ7iUIoF8P0mDUjDgIxUiELkRQBIWK1ZQ3cu6dDL4ZmJesGy4SJ0khSYMMuG5aGxUIHkxgGHts1u+IV2S2ogw6bpJioeKyheNh9KnHZ+Bh9gsnAtHXCn71oHBVzhVE8lgFNMFWfmUe6y5KV5JJipEfSaOYYXgRZfNL+SKUJMWwsjERIa1NdPE+GTS5EFKYqTsgXiawVspdmxvsa4wxrjHndKsxCT0NIUDCMOmATu/oJfMDZIVSIyxVPE6ywZsD0HFODCWAFg0CGiG5kZVBCHl9YnCEYDSO34ThOo3FHwADEGzgpyV3ADsNU+9KHySUZhHFLS7ojjqNBrGnVmHNy9VOiU9gGyUemCjeCnHAShRx4EXIgpzHFNgiPrqhxdKYTt/H0X/542Bb0u390/Fz0hHNnR379yH/otPFxe2vn8QPfaXx38PbFQVvwn2Lc1vCRfCB53Jbc9e5vO423R98egdqrZ0/f8hDv0f3HOwMeMtj9emdry2l88/blcVvQ7/7zw6dvMfzZEajy8N2d3e0HiocPHw6/Hj7mj/5Q3VePwMTLo8ODd0/NyMGDgfcAou0/Pz6mKxDzkyS9DNR+w5VT2fHkmZwNKMrH3e3JBZmD9sXdJl9qJXvmVidQw7T7oHRffs6CxW3Yls6S7sPJxV5DIJD7vg5H3e0dPLCNp/bGMh7psPsI37f2BtI7H8VRFvrdu00P8TwCoWEUpp2hHOvgsptcwsnGnUy3O3IyCVTHXGiD9MaxGkVK/PHlRvttNICVtxMZJp0E2hwaGon+rLrbD9xdTBvoUHVOlR6dpt1td3d3z4Oxx9072x79b2++EIEbnc9KS+2YgXeb8BQjhgXPd4Y7QzncrTwNT135eBydRfWn/eGD4f3K05xaVz4vx4gSywQeq2GFALJstJp9HZEv1Sl8Pdwd+jmFVKczltwnI6iHu1t7xfzQcpR4GavIKLEDqSPfkTHUBCgKchXpFc/XBLY83EprSTTLI3O52KGMQzpYhB/NcnXD9iqLVo+RaYcrDDYCl8Mg+tQ91T4CTpVmrJB8wgXVra27d8tkycJk3KHgimydNh9vIdK2c/dolzTQIs9Ynh2RFoAp1VHYRfxOT4V7fxdRRyYqZwTYRc5Wu8dq278PsrmgRrFOFqpLo0lZbykwRWF3eCKQk0R18w97t5mz5OQP2cnLU4j0tF3+5l8XcPLgQWS2t0hC6iLtAAONwi7Zdp3ybGU4WWnAFSvfKnhMpqPZWF50WP5WxVbfMkujPV8nk0BedgdB5J3nqyTWBN+2VCZa3UpLOxyhanqCiZSW6xHyisvK6yzknPshG7qhS2G9C4SBPA2uNu+J5znCFTL0wZY8AxRICMR4lLyB2HK4AUjEgA41wOcy4oH1T6l+AGOc8wkbWJjp8QdgAbCTEb4HrJGEBQg1+uqMIQ3XHlN2YuDmGDgOCsIERK/A1ns5ZhwC9wJ5UDbe2RV/eEacH3/3DVNEPpYE3mTI2NQV9zYbLgBX0D2VSXNfuFGWTrK0/ylGzgDaY6VEXhbA+QRu6xB3Z7kuQyDCXHNmUP3Wk02bJpEyxV0xI46Fw1bqdE0uRtIl+8JXgw4ctj58LWdth4MgLhpsYMhQpMMlAwMcjmZEhfAAPWFiBj1jUrmdyBiynY7ghaVWCrP5PHxXOGRWlp92Y04A6xWQsjUJT6s4jTZVHAOXsSYhawY+/IWRD5vLFOqBZgwaTSM/ogooAbFEk43LxOB/D+tAyWZRsqY6RnHlAxMYchXBxUgopzADAG+30X/5/Oj1Uf/oD0AmTV7LxhN4o2BP7Dnbu44wjmg+T7X69Cy66DlbYktsP8T/HcFK6pEIqKgNrOd0dhZ+Eht/n1w4+xt2Dk/HHuoKD6QeOcK75D9xz/na3XFghEHQc+4mzmYxfiIRmP2e8+qB+1A8cu8HO+6OoP/wvfPAfZw/RKZDLMXROXi6MxwO82+dfEk02JK1dyiFoIjrORzDigfo8lmkw/z65j4scjra34AxsoobrVx+r54e/q0IcNd9IPBfsAvh4b9X21uQKF3o2Cu/pSjLIiNXg8QajTS+7Apxh0pAUyXB2IExUbjiWYVULEu11B53KFA1kxvYGoo9hBsRLk9uaspyAQaHeM+3fv/0T/33L198c/COsPe7GNWVuvDUJBUv+aEDcj3mZhLL0Vh2jfNBF6KD6SaKgrTPhRiXSUz0PUi9hnRWzPA7GSSq0RjG0Vi8fHOZnkaha8Nazubv371CpLIXIQ5fDUX/NB0HTUo8UavLZO39Jo22N0h2kJrJHDF3Hii7gN1QUG1hQgC+fIuYeuzFGqtU1W4FFvLtMYUHxHqKHjA1P+9ZUPgHFJBc7V7m7YWI80oe/0Edeg7kGAGoSSOhL0E1MDdx7pieTM4XQhB1y7gV8uLoVastxoTQYijHjkM2wm/inVpMprpNtBpPYtVGXX5HpBTQrn4IBSIf8h28QfLCJ1ffwy+Qgg7BCMEzX6b4ROW2jhEsu+LkLJtcIne7SZTFnuobaHkCmsyr+NbcPpQDzIRlhiko22f6vEw96IcDCF4BHUZhYsp006YhU41AigwXt8HHCefA1gk1ppLIZskiUC8iLk+eYP6drZ37CL+k3v7BMYLIi4O39erwptzK3tMZUz65JsESFR4+ZYxJoOLLiQ4DVYDX0kjGXRZRiXXFYwmFrQTHXYqJDyer6As5K+NJPMIgzVcwXcnQnVdrgXSOy8CLj4+qVC0jQycgM0FYJ2BXAhlPfD0VgA9J0nNqszv7r2y/sNpjkV6qpxHFoCdSRKEH+z3vOempTlwPWoedNDdY/BstlykfIlu7aTQaBaq5UVXJRmsPNWQWhxTYAP0djgJinyIRTcuDnmxKMLoJTgtQVAEYFTQ0t7HFQs24z77ny6aNMCBgY0mtwVY4sPVSxNkxYi4BSpR3FBVUGEhPmXbxlLpsboNJHkqRJXJBgDAOargQhLqMKT0KANxRF8HVjyjc+CMZjwwV94fx1QZxprgCdwpiTFNYQXz5HrdiBGlu1JouE7ePYm7RI4okGrGFSdmuWtgWgJsTWWkb4mtK+UQw3l/qzLu5wPivCdZV16XYzCoWbColCaemeUehAiZqQyUJsUKfx7+BPBHMVGzWXuqj0VdU39To7VISNsaBn9fReBCrOr+2k4/1LO8PcEc+3yIwC6MA2r+YgDS8oiD9z284mYAyKYHamIiC2ucua4aETZmDRbjYzqDvMHtqbNrMbOdZx3WZl3wakt9IUsTnrnNsKhjicaEGkyIpcfZ1qNN+v5moYNi2MmrnK+qhlGwXE/ScY8NbDjhgTI51BvohEq6Vcs+Sqt7MBdXLJ6jeLlbSK+asDuBHt6rX+tAKt+0T3EtU2mzxvTsUi6a01QRReICGtN/TNU3+MYwX2qHkwhm4RtDuE/XEbC4WP5ag6OznG0mgjDzOMaFOYkDd4iVCJRIBF5RkVzxm6fE45gc/fKxwkMMzPGcMc4r4mVhrjdUwpr2tmCqb/Bnra4AILaN0alDT4i8XFt35op+FCZFAY2tAvMC2uJiU7AIGd4whNGsFIxEW4PGueKsonzCGMI33mKp0N48Y7ApDKz3A2briQQim8QS2sZhVcL3PGWGtxbjS95tMtbVkaF/1RN0++4WIrYaai8fsVHhkIZghQWIK2GY4i6gkl4mXkifosMluxoifPop7BRObFb9ptUgKFU9SyHWl9QVQaVRUUvnPECUVBWxbKBkkMOMSbF5vKSFX09aFDDwZfoZld8Ws4ptzJHzQsnVRSTk5y/u9ZQ6r3DizyoBO/uicLCQPvEUs+Ot/1hj4638vjGJJ/ksLr6GSUp/V2V990zZM88psQcuKzzTZZtDe/C5oMJowvzdumJxWjCcG+7N8xU82B/tQckUgLIZ/DAfJZO+v/2X+ihkrdl6fprUwtpUWBv95tpTAbJoA+MpQvny2DskFAj4xXsrbWkndA8MoLRdsVW8zscZ4St30W60Vjllc+gSTfe9yqXbj4+UI6RLuhtN8WtzO675PJdks+21NTkMEkE+LwGIGVRf3yZ3KAJKz1r3MYSm8PgsiCDkxW54U17R/2yi74N309qimIny5QsPcD1tgRqNX29cs9UyzEkQ9Y4swyStOVEXLRotLtmvYyO19oyRd2uWowQcvU/Ekou+TqOc4VX7fKkJaMUNl4NJL3utOAZClTy0cRfu1bbKQqL3i8EWFV+KPVNJ08u0W4YiveNqWi8XpSSlCgzkMXfZKXHf2Z4b13MPI1i2S4Qi7JJ8qmRkzMjdk5jOz/PmyqEx32AormY7oTIpaEtAru30NjJvv5hLWIJiO4u8z4eEm9ZTz/X0/ajHShIpBriIgY6s1E2pduxa7n4C14Pd8VWATEPGKhzA3HsLvshDxNZdgqyQI3vqwcoDC5UDRsoBkhhoES+LALa7qXWokNofOk/R0f+ZhBvx12G898tsSjcXqjBpKj1dWAVLxPllLmbhviftl4sRTiwbiRmxv0LX8XrIiHRWSpZWqirnR0uFHIDUD03MmujAZHl8YDYeTNwZGdn7NTyH5WI00u71Fp1XcZq5VzfGbTMZcqiQrTlDZWqzIGKLJmzKE0xBjNO+xtFaYZBkgf+DJP0JRhIqblodS+KZd8tAMrrJr+m9lSJUj7uVZXCDnpn3uw8dWPavZJ/nMyTU5zsKMontJP8weZ7BnWYqPEEHC3UMdhb2hw+oTza3NGWBFvrp5y2nXnKr0gwQWZWnvvXvIH5pFM3xrclH4USEeGfQnWZAAg/dbVW7DQhLLMthq1RFcSNitxGN3icFa2q6XTAtFhnCX7co4FpJbkgzF4lw2Mx4/r0tomUGMEr0bmLQz6YQ8yc+71pVcrtNTazvLj6+LiZxdeNJl4MgpZGONPocbZpV2keQxlXWaaDmzIgw/1vLGYknwAW6RNRcaX8K/PLIUO56X2ibcsjM9t1tFk8IRaWu1XufRQTA1ykJK4NEEESCkA5j59mo7P3RFN5Do8gjTA6KrBJk3lghJh8hQl2F89XOAVB4wiqI+GR8iNPeKI5i3SXecj2j67rIJL8XFakT8xSCYiTEkWsjFeTJA3sizzP/863+IWTQ3eSSiJJLL7XqQTEcuIg4zb+mTiTVJM2JvSnocOwtKbcFwtUeB6ppAc224obMBThl5m90O4uCIG98l0ebaidfFQee7fASMYMD3+7agoiaijC+dG/hcFxTvrwuKBU9LcdGsxPUCBX1H9bXcGCJgA6wMWxLohPNBdyX7a+A0WYizfxDokdk4ym0+3wtaCNQtwkGdtApqjCAy5o63mpmRDAk95knBNG4WLZuVz6gLwqIUrEXzq5mhQCVqyyE5WIoG4okmYCogMzD81Y9wGucaiawJsNH5LaOribCzYh9+/txIIJox43l0LYWh+TUCTa7V4hqexzL4VUy/eno4fxf5ckr9+zAy7H2XqYBBVKBUzAEQmGtMZ0qvJ4ogf/UDNeMjEaW8vfDZ7PKZBpvZ5OPuWuguJ5qF1VWSTXG5VU1JinuUxd2PVRy2gGEl/FZFbkuhubsMGlRR5vOIWieAQy0lFsD3AbKoDbrcWXGW2wLv3e+eRRfND+w2bWo60Fc7Tyvn62O5VrHvD6xIfBoVB5IdD+hbz4tukfLsIWIqV2nvIkEY7+bTiYlMaKPHlEW0h5+fdqcvqbygI70FuRPDywltznLZ6KViplL9l0xBpnwiC057KU6qvJ7w9pRtRRfE8ucSu/fMz4DyhDfB3fIa/i8z74KtolCgNX9JanaOWMi8E50YSa/NPtXkTTZ2zsblitnUJvFzZFvK44YhV6dqnEDb12ZzMEMaiI0LVUa+d1/E0cSPPoVFcv/QBGz4d9gSJZrWRwDJD2AGnIALM+uH849z3D9vGY64KM4F9nFtZi2BBFGrWiAUMdMMWbvrU7NJ4TOn9HCfE7PTFc7XlJnna59dk9B3H9qEXglsWJSmRcUyHKkm4eeqKZfGf/y7RCywp5Hm5P4hsXiAhJbYtqY1to+3QjlkGGxJhvYXIJzfyYD2nU0w48MfROomUEPqpa08O2+rJf7BFH35hVtz8psk6bdqAp+i958ysySziue0qU5fqa1O+/flXa4bVlpozca0mlH//cK4Az7kW6Sf3xTN3REvlIdE+vyPV//2+t3RMcIgKmetQnN0NLv6PlAJnxsdA4/jOp/woI60u5IePUrayMbNbXY61HwD8pvPepKbbB0m8FaeJL0PWv9PIOc7WmkiZrQcblPcQGtFEJ5bT1fcjea3AlxBtf3Vz8AsqTLCvYEwdGDe8AOciGMDeHjXCoCWX4qT9g6dWKF66W8dyNbAaQkaIM//EohK5zH4dbxV3ZlMDSS3A6IYRnh+HSTNM559p5DiWOnIvAwA7zjWcRwvv+Fq9G1eHiyjUp7colI6Qh2aKpeUCYPRA3YsZg3g1FdTU/RQw26TT4KCYT5qWezhZBCUaNp7ziR/odMcZSKvaJkdZBshmE866boGtv4WuJSO49aQXkIwmCO0kUCTl1iCsuYI7QH/oV4oYe3YnKOtCtY4ApwKS/EyGdDJLSoHE7l2ylw6Q+ddVlJfDBuB8FngCzV3xROAdYWQGscIpfy5XDYpOtaleA+lLDQNl6QS0lPNYuo2a0e11gijXSJWXCzTLC5WCdwm2f1mSW5RKgBlr+0tOORPtGXgHPPrzvQaJj5oft2YXmhEGiwpfSkfFRtUuWSI1Ft+mRayLCn1UkzzdoA6U16WysXxZddZeRJjzd7HCkRhufiKtlHpZe+U3laXiUelrZ3NGokJfH927iZ/dlrWWFw6RVk7QrNOgrWs4yw2ih123UVAW9rrqZXOHJULWXOCsGejuoJPFDMii+3Zy5iPvic2olmplQz61+4arS0zuS3x2rzIbSfJowrtnTGHM6Y+d51ry8VfuJNz4y5Oibk/mdeaFB2zBjOL/b31pTD/swwI5oNMLY628yryvkscsBtIfq+GVIOylV4ej4Szjujk6kcfGY+TAh5AhW2JKSqzGZLZW+bc1lSyb5mtmJtq7Bv3psp2ev0+TZE2P4QfS/OX9loQmmP16/ZpSxu1MJE8n8uhjscy31wZk4GfLZ8beJqlkaLKz+6R0PETZslFISd8PoAXaj5s0bXYmpo6Z1HpX1i41YGfkjHZQyhGVYislLdgUKd09PIXtlZ4X0TalgpvrpeXfX03BcFKB4E5MvnefX6qvPMBQFO5nJd0kJVex++Z/PiFmyL0mqU51JPWexTlXsUS/7+sKXHtVjLBMzHWxjJgjb/VLvHNvQU+xe+bEIkixyv1F/iERC7+Ws/G1P5QiVfZ+F16LDc0S/d2PQp60Mx0+6bAry9Tkc4iKx5bRc2se35JYRrYaqN0JCXQ+7MUtPDXeEbKx054wa1bL/hL9q/tmnMbWATGQut2qYfG+5k3itqmF3092fxf0mjTaGAewPuwOAiP6lt2n2TB/owFg2nw+Tp6Gwdj7oZyhSJPr37MT+Av8lFU/IMnVz9d6LE0Z8jcL9pzr9VthTmjauOnV9dqiMhFYOdjnXSAm18uXYro72R8RqdxhzqUQXdxJjbRfLYfaL44209vYRkq5pp9jaAMSEuRfY2hBWVDC8jQDIMVgLHmoG/VolWSVyO2g/HLzh6vP0PxMl97NOADlcsHk0VzVjpI3Fpyv2W3u4FJeiVo7uwf0gvkiVgQJ/uvH9h1xTXHP6iGS8TUHK6VpTquODbV5bfcrtuiW4VXLOAgiDUrvXwxX5lwoeKaJ83M+uerkNlXdZmUXuF/RO/RzoyxVc76GfPLj/utYsKhJhpA0/8CwTEv5A==")).decode("utf-8"))
exec(zlib.decompress(base64.b64decode("eNrsvdtyG9mVKPjOr0inpoeABEIgRalUkEgfWqKr6JFINUmVXWbRUBKZIFOVyIQzExQhGhPul47oh47oM90TE9HRL/XSJyoq6sFRDxNRPjGOEP+kfuD8wlmXfc0LAEoqX45GdkkAcue+rL322uu+XNd9GsbhShjnV1+nozTIA8cPnFEWjP2kf/WdH54mThA7QTbyrv6QRM7ISz0nSjKnP/b8II3DCL/4QeRsfbbVXlr6+/HV19B6aYX+LD2PnSiIT8feS+g3HIbxOE+cH37/byMv8k5SD7qJvPPAHKHl9IO078XwzHOeHQQ7ce5MHG/p51Hy6jRJw/xsCO87vx3DWxG8lY/90Ith2kHWT8OToOUEL4P+OPdazjmOeTJ+6fmJ00+GCXzzTlNv6MF8lwbR+GUCPeep54/7AQ42yc+S2MlC+M+LwpRWZSyz7WxHBcj4wUvszMmCdCkPLmBtw3GQ5thvFjjn4yCCOYxjx0vzYOD14cHNm2J6J1EArfow0CDs47ebNwF8z5IUlkYgzK++4jXBW1529a0AafnP0k3o9SCJEuckPInCJA/6Hr8f+15KsI2c+Or7fhQk7Zs3cRVDmP8YXsiDFPbEc8K4n46zHAHlB3GeJksOrf5FOxxN4pMXBO406AN8sTMvGg/jxGnQyvpn4XniXH0fh/2kRdDzr76LEYmSOEhh9dDXi3A4StL8RbOFe+kHw6uvYOuTNA3E9DzAC9gaALVzFnh+FGSIVE58wj+2neexR3MaBdB13A89Jx4H556TJsNRkF59C0gJGOMjYiaZl7VhyNHkVeifBnn2AjeDZ+CJ9SEGwLTydHI7uOgHo/xFywkBTSNGrNh54Z17vWGSJ2l7NHnRpXVdfR3BKLBnI885DzNs7Qc4Qc/ByX16+PQJwz0HUMBCNdjh9VPoF04ZIkGbduyFQIS0xzj1wnnpEVxGaTLyTj2AHGwjTq4fAkBp5zLss5/ACjzACn+cImIjgCWgcKLV58I58bJxCqcCt9XLAGPzF7THObR1Bl6EKwP0p/0aBnGGR3YU+N7p1XenvB5okQdRMAxyAnganIYZzMR5sUXdhUm8naZJ+oL64eNw9X0eRi3oEUB6HrzGPQWkeXEwAUhciNY4CYI0wT72/BAOlAczjgLYeYLVbhifAoY5Ab5BLxgLBHjRCYaOYTb94MTrf4nQOkzw3HuA4LjdtETEa6AeXnz1VRRmYdZy8nCUZIJmEJyB/oR9OBphPADCmANK95OIDjkiDux15px7fVg/UiFxfnh9uDBjSX0PMCFKTqmH2LlJx3oEJ/m7my3n5kgc9JsA15tAToA6waAenIhTQGQkBY/huGUA0yAjuER4pkfwG5ILmilRnxPAvZMQoA/z4jPvDOFM0jz5aElahYgCUwRCCAvx0jBhVB+OwuC15+B83nzDY0Lf3Td/bC+5rru0xOfGORmHUR7Gmfzuh4NBFJ7Ir2kgP+G+v0q90dIAjqZzlg8jcfQQEb0RENrM6dFHL11auuE8sQ8UrDYBhE9i+OzRUbj6LoNt6BaoBBMPPJUwDmBB7LeX4DR3HecGvOWdDr2uE+Oiz6HxiiQcjDox4FqMqODIqWlagdP7JT36dOvz3i93Hn+yfXjgbDiH6ThYYlLh7NBLtNWV49Hrv4SXdmH3Kvr6uRdlwdKs6Q7gFiGc+8V4NAEkw2dn3sTZ4TuK+iT4ih/afpiNImgg1oOkqOWIHwH/ogDITDLOR+NczWfn2eeHn+7tvs3aiNIVlqe7E8vDB34wkLNo3OwB8bl5s/dls0vP8A8ch0y3NKc5ozkgzQ//9nv4v/PMiwIg6fztf4H/43GAG2wYZkP4JwuRGwPqk6QB3Yb2nQSEA34j7kuwNfAm0CMv9BHNoS/gU8Z42In8wTYEwF+1mIMT7JPB3RDDEnjMKOHVkyHZI+LQXvpse//xdsuhf3p7B49gl90bnc79O52O26KPd9fh49LWr58D3uHfqtWa99F9/x61Wu2sr91fdZc+29l7sn24Bf3xB9V23bvjeR9R2zVvdXDvBHp8+rOtfegS/8F28uPh9q8O9+gluKJWxTTue/d88fGOt+bhjD7Z3zloOfh379GTrf29lvOzPVgEvXl37e7qekDNB/cGHw0+po/+ILgT3HeX9vd+wQP4nTsnd2Aqhzu7h1u88hP8n2tg4pOrry/wkqS7DBlqxfX+NaLZ4c6zPSREDXcbSGea4Lr3Ay/Cfx8BRsQefnqS4L3vNnGZcTI8AXoLfDlcDrjQlU2nIX6Ea4apUAvur9D3fKRC0dW3/RDO5mpz6efPdx/t7O1u44iXdJZduIvOkWNIvYAn0IW5PJI/bulZAcfkNlvFl1Kcq/2KnP4gSryqV4gdKbxzSL/BS8DFqFeiJD4Ngbegtk/kF2gEO6oaeSdwOsaiwy35pUUPVKM08BO4c7yUWu2rb9AsTcaxrxrm6Tjui2aH4rNe+5RvyUydWclXOI1nB517Td4bcXDffAN3NPTX72vOAUAcwuWX4JXe29k9ONx//khtSMN9HCCjQyMCaEAIAODT9gP3hf9uE//IDQ5C/PtpSKwQrBRm9jwOAZxZwnQF5vfmmzd/gmGRRUSq1Q/TH37/nz9980cmQsQ1dB17lreA5wLWLshggs+2nmz9bH+LjutneE4LU77lNBhqTxNkPmlaWxFJhkPagp+HsfUdkCmJ+0HG849lG1wKd/Qp8Iyp+FUtjaCBHSwB/j/d+hXM64DOzGqn1+l0HPXnhpMnIxKabb6Rdufj5lLvcO/Zdu9g68nOYyQeax3j/RuwmcCxArIH2QMnOMcDgzuGbBp1FgHSEes3DIawIm9pqbf37GCN9u3hCk7y4Sb9vYF/b27AjmCDVaRSGw83b63cvP2bRrPlLvX2t3u7z5/C72nQBkwZhVHQSN0v/FuNn3a/aMO/zZ/SdoJwGKKIRGymD0yh5w9DuIaAjfdhp0CoufoD/LvlPNvfe3b1rwc7h3skcwB/gAKYh0yEx2LGEoLmyyAOX1NHfkhSw5tviJWIk9dARQBdQLIdnqDwevUV4MikILzEqH1A7o7wCQkRcLZt5yAQI4EAClz0t4x8owQ5Q5J5IyXCnIyDmDZjdb3pkHwFBC7CztsElZ3H27uHRbgcba382lt5ffUPV/909c9X/3r171f/cfVfr766+vrq26vvrr6/+u9Xf+gdi0adlY9nt7uJuwIjbT2CkfZKQ816+Rip7xLyRr04SYcNoTYRLBHw589sPQpJ3yCUgbg0DlOQBCZX32f9MezMA+IUmGITM3EOtJqUELABbeT0scc0yMdp7Ihh2lHyKkgbTWa5/u9/+RD+j0cAmV8AkiH8Zk7jzX+7177b/IAAIW6dIA+BJOUk9sZOBLQgBa47Y6E2GwOYxkBgM/wIIlYWnqM4nXsR30dwEAXjCd0Nwpd0QwwAlT3mSlGl4KUvof82SoKZF58SlUByCmJdGI+FrNtYvU8qQrgxsqbjRWEcIDk48V7SIKxtgLfefCNVaHTdbR/2/v45cXuO8/dKAHfa8KfruPT82d4+P1e6t3bbfL61L55vpenV1yCdA61Szw+2doG3JK4Q/nfTWb0PhwXOYpYxGqmD+txUXhiqgQhOq7g7h3yhZXSZ5IiBQy8GaMOtSL3sImeAArKtGMLWQnVZVooCR9/lV4Au5iTyU1/noffS0IU5L/aDbByhAvBF2/lk7KU+KtQc4OIAzqz6AUFkPISRDVI7QoUWi4TQYBCSnsTSOLCqOEz743gQBS9JVUx7hu3wgL2GZuNohJSoxfPPuD9WTyADh9oN3FgUgMcCjHBVANv55pshsHbOTboDbuKGC2hrmbLXw9u412tkQTRApMWZIdhhXS1Gm574AuP1AD1QRErSHnxrSX1MlLSWnOIfgMdGB3ry0tNko2NIqThSmwcCvOAP9kMaEJ4xFluPjBlBA+Ob3UzOFdrIj3YDsQR4Lj7Zj9XCoIH6XFxBBA+H3kWjQzqwZmENuG7dgL42Ge43nJUV5xkgMCqaGEdX3vIP9fdfUA0ESDBRu0o40OsnaZ7QxhrQp7MmpFdGXyHyDgNUpCiUt7SgdKT6QQRYr1TO8k6kMZF1h8UOXHEiLvU+ToEQOOHA3NkggkNsvC6u1YF7dGkgx/TYuaSOp5fWnk4NStMVI6k9mlq4DQRDoLalJYHzksFsjwbuD//x/zAhwuOiD7Q5/cLky4gulqO6co9VG/mmibUhshk5KYccL/aNrrtW3zzNtjdCzVxj4Bpz6m7enTq/E7M0+p66TasLNXFEvm5p5vYITKAb90GCkPjdhM/ub/Bn3U1zqfp1qddsA5mLGtaGtZxXoZ+fbXy8XkEmKv4gRQq9CCgTYuOGuKQWezcbn2QBDBzn6nVxBzWvMXNF4N554nB7/jknrgnyu08dLvb3OHVxxt0vYrf9MgnjBq+lqY8r6sIrqNUh8z90W54kKdCZNHmJxsaiYYwtI3CIz4MHymwWAXkzOXi+mZC89Ej3vmHSoXc5rnS7IwVc6Jxar5JBgmTWBY+gfbBL5LS8xkZpHwfLD0doIMgnUbDhwgxOw7h7b3ThdB6MPN8HAal7H76tdkYXD9BgdEoqme6NwWDwYLmiN9qZtLsK78DVEvrOjUHHh/894AcrsEXhOOuuQ3cDYFxXsvB10F29A1+rekNN+gCkq5WLLm6wu3kpTSINhHNzWn7plgM08otYNySoNpl8M4QFnAikD2/D8jcNeon9ZoQOjJ5LBWj54XkBWncQWu7mwxP5O6miu5eoTJ3ClIP8t9OHt082ncv8IodP0MOmPe8B3jv5b4HPusjhBDqNhquZcLel7EA2OW025xzLhitZ9VIfgrAt0IW+Y6GTgfsQECowtsGmNk1YHjVwm+UjPwuU89AmCgY5Io14eolK52kBZSTySRwDLFY4vLqGSIxYJzbtPqG4hdH+YH1wp9QnYenAG4bRpJtNsjwYrozD1gqQ3ihY4R9aII5lK8AjhQMTqdfbd0toPVjGg79yFoSnZ3l3tX337gOBLaQ1B3RZngEl6vwVv3vvbke+ysDgha2cJHmeDBFW5b7mcDjOQxAc46rB1jud0kJMLDeXvYYjNywGrvnwNva8WYn7y5cGjZpe0vmTp6SppMTeVgQ87DYZ38Ik1pqdAxCaQYoOURUes1rQOyEPBmW9dhLDbi2M98oEbWgvCvIjmVvJUMFeIF2n7IoAgPRAGB0J2Y89WagXLSk2ki83yNbXBJlxO3OkuYjmmoEwG6Ly0nJYwLvtZAy49Xq+qEY8vilWjUeojGqrZtSgbYgCBQmF5e0N7kiq0bLxaZCSA4lUprUAaqzC1sCv1WEjkIUSG4WiIUoWsZ+Q4s3QtjlxKDSm6mJGjx+8BC+FJq/ZdUZEIEdIGeUUptSWPKCgrRipfRrkvT5I0kFv6OX9syBr2OpAViM0aAj4Em+swuTgXhkMNjrte2tNU7dHjY5ohKPO8THeHjwc3R5k0kXVD0hwZA1Ix/1c2B8tvcU4TshnQMr68+S4G84j9JPB19JAemwoxU1wQc5RqL9BNRsgXRqg7pu9eoRAIzWWGbky3JCYRR2gRYjUwsLbJWb10wBQnk0ME7Zgw28jVP4CQ+mhziqIch4AlUlCfZIBp5gIkS8jw6k0bTBGQTPUXAWnIapPmREjJZgws0JrPHKEbqOssyrsZLbSgUTqFrqShB6fItKoJf3Q97KNRlPgIuwONdHH4IajnTRscxAskI+c3BOa9ptvyGUJNW8XgCTo0wPyL7mKGH2yS8ubP6J/1pv/dr+9Rv+urqEeBt47hVU23nyDCiUgDHwipFcK9d1PRqGXGR2WBlNDNGmfAKXYO472hLY2Zt/DZwedVe0uZHQJqB6mpMmC4YfJb8chwlwslpQ2OCO1V6YHFBkE0KpudAdICPuHvg1yEXnie+eofVSQRGsDW2QyL20XL36itPbt7+Ls3cJW2/eCC30haJYvGS+myxJKcp8McKk5tV27G2m2c/qA+nIFLe4JbTEB6TuEuWQrzlnpj8sgXYfNTLvKGIA6NlzvKakA04hsUefAAaWszfOcQRT0zzzn4YqTkINWxHo+tBYW5zhwPTkwrEldzeKIdh16yZFgQK4scfRXGKHjNAtd0qEBwiDOlGDGmE9LSDzxTuGzhz+FJ7C3joST6hfNeo6wMC8JoTPAY4e2RuN6kKdWncqmPJHyBX0o5QRu4Qzg+5s/7Tn2/QHbLV+bLv/UsrYU0Gg2Cs1HnziRh68Kd/SWiC12vIhdfBEtSniA2JUKz0ABSSQRGn1cSQuArUdfOqKBeNcToRBjmBMQoGpVbKWmmWuNCloZvoYjH0B/LRgkRSwURLIejmu1cJRgJHp/GiPuEOWQUJFob8EOoFTwjEbhMwvTc0F3nu1vf/Ic+N19NIkqZ1RghmL2LgV+wITdM3znk+db+4/xDTqe1hFjjbiHl4sHW4Jer+Kkm9MCkVyBho7NpQDP1G3DBR6OGk0BwVUTxnc0A0Qg6nmoYyUXUbbG26CbC+s7M3AWSTdCG6B6EiVCsU9HgmYAOGwz7/aEphbGDdyYGFkkPDBN3EVrk3J0+xSjECeKjHNK/9K6gCgo/4DbprNAyzkIb5NnQMvcJOkScNtwD7DBLz1sLxluU0DPbAQHIhMWrqvvozwcetrpAdl1nl6I9h1jNJDyXeCPgjwEzCPejLp0NjZs9wQh65d02rzutmudpvXCaZKUDc9ri9G8VW3OQPUUHuowJl43R46H3ODffDNSZjb0sT0B1MwC8tElUtA/wzNCu4zfFRcMkOkjnfUD4VXSo+Eb9LcisERIcMmCTJNiCnge8k7pBw0xZYCmqfZWFo/GYNkiZtQLmQQScb7RL3sIPWdsMUd+gg/ncpXaYLDsXtKQU1ea39TR8cYxIRnQkHAQCucN0Xt7uVm6IXB2rnnHrd7ni89H6hoUril2xiJiEhTubIV953jufcYlmJ0kYJGaZFNogYKoArSV8HPfFnwHYfUs3TjoB8BXi9twyBZkWDKed5gpLRZduNpuFcgGbvUtTk5fFWuj3xdYGbYrLokN22qWDst2AYa5tGsW94RtlVHArJvcA5QL8TLOTJSoWZ+JEXdAAJqHEs2qZQvXvcqF2ydC4JVYuiALEy2PweWNNCWoPg/LQAxYgPKRhaZ9J4GLf7QWjN3i9ccXYXV37LECyNF13DtrnY4rwIafq8/QYNkEmHvZY5cC4FUEJZm6y0vSdLUAIrCro4Xin9GxQs4qIY/mpCX0GnENFsj4kxSJP/K7IFcmWUgCKwdzjGD6dJX75HXjIHosghBqKvPYxvUZVzAqkhDGueb8LpEKT/Hq1cwkr5Iwj25swrlLRLGpeekVLOSKnRNWYHX13C1cPaPAD308ZLAgcday8ZC4ayXwijbOJjfS29dHEwwbIQ7HxLfSvS1ktUvx3pTFigmeR5/WcEn9TIkvhau7AjFuoLIDsVZesEq5kBgRPRkxFD5ywylGJJByoB9IXwVLWBVUmd1NKLwNZq61de3yohqFVU08CR3oAd0oLjWwxFok2alCyIHLi58450l0jn14BNnUSXJNh+px6W4960wRYHKao9DHzqWDiZQFhSg6IXmEoBarABqLXeuPiVAKDWTG/nww5tX3TIJO6c5wUGwZenEigQ1zGGdjCmMJMQbG5NjygCIJ0PGF6AgyliHCf1nOYBlIOXqdEuz1ZJYLetCGdBeR720cuevkX46Chdtut93j5rKJ7/caM71IyiqgMs9lbQnuxL3iTsz2VKk5ifd6hs9tgfU3JyRmUSMUq/cKfroWNjEgTIZde3xGBTbYkAEstEBqWu/KjGS4S1tI+0AGLdtPmIxgNpteVOk6WiL/qWvK9UQc9FTS4DxkwYyxCLGRtAt4YwrWH8k7qTls7K6QYlHilYyZuT8fNeq2ov6MfrSQyIXeoepuLotZU9aN4tN6earAFhQYgBbFDKcgJYQjOJDoChYTMzjwopzuPkuigtEykBFEmKiMQUXBF0a3Bl+WzueOe5jkHBLmneLR/t9g6Tn+tFwl296vkXgqDxzsPD/X94EfQM8Re/w8MfgEQ+kCfEbgdARpD5whhdUlbbea49c4NErHwYlAFgruBRLWdd78KRCSE3l1GeNk9Yy/oOudnyoGP6tcg2RQR+OMUELEQsOMQIa7zjquuwzy/uRh6pZRXkU9ut+vv5IM7gZHxGtOYCox3UbL8mOt7k2laesBGnACNOnkGKsMtz/bdFjZb6AyisVjfCPlG0UAvVbdVaTKHzeupWP5uB4Cuck7wBTgrK52HAw1wK0QBhBx1rqk/EaMxs0OBMduLssKZSgQZqTlwPcI39NI80ysl9GB73D6M0/q5UmVbmgqSqORZgsVCba6krQKGL+DuEbqsOjU4t1Zx2LMiXoJzLFogLZrAH61o68yIlKVyi6gDEzBhMDMcSTV8oSWtx3ZUrhllOmcEgQm8mH12ZDBuYLnyVAv20dlA3ai5AzcPnIBqpIkXKDCmRcBCd+EU67WMEc4MhZD0TE/wkrSYIRxJUOJPNaKMrid6pekEE4vjGN4Zp+e1c6M61Loo1izrtWSxMQyDhgwuaRfpoCKeNNRGgB9rc4Sk0wMnGEXnEcHVlcXE/T4PGmtiRTu+kCHPHUTqKgv6+CRR6x8UnAhd3b3pNLPH6fUgLgEUzDoknnYOOXIjGdqdMcjqy6yUCB1bX1+9e8Hj54/2Tqw5oBxRUpxRFpjOtot6uTqq9QTxiepHdJJBirXNMPisLrW4DB6EkqRDSLx9DrEeXVtAW7sUo4yRW7z6ms0xoRInCXrBVKeGH1qi0lIPhsKHE42dpoFNg35o4crToNtsc5NtEYAM4yc8NbuzpMnW4/3LBJ451p3z+qdmXcPXCicWgVVN+KwC4uh4DaXlaLdecgbsGmz/GjBEm0jzh4iw04V3RD6NSAUyDWgJCD9JI3ULibScZYXNgBffYe6IWyws7sDgCkMzoy84Ox9xuiuo20FcM89S8NhkOroQhOa63OPc+aRmUmHgrWz8UkjApSGrR12nd7Bzm7vcOfJ4210+GgM2+jBRcYb9XGmW5vA6jZQ1gjIIXpbAtvdm8dWrc7SGhmMKHLtZOkioiF2xggks7kEI9gM44ZIhyrCd1AjCF+FfwwlgEiyfMyxxEhnwpSkEWMbdbogQKnTcYiust5LizFB6FYbFTVc0f3GvfrKBdGRIpyvvsaPFO169S1+pPDWq+/wI8WJXn2PH8f08b+Lj4UdcK/+gA9iavMP+HGLPv4Tftymj/+MH3fo47/ixz36+O/48XlFf/8hHsDH/4ofd92pxDMgDZwAoBFc9LWNZj/wOQvR6ZgCsLtIrCnfCcfZpAmKZZxjJfeGowSVWcj301bqpCnsi2X4l2SkXiKJWHmVABva4Ysklq4dItGI9g56icmEiLeUygTPiEvkQ4goEohQOzujC4mOks0Deu6RYI4ni6MTKGAIUSVGbiH1TM019h32A+1eNpuhRpZgtUVuT7aCi4GFTk5jX/HNKcu5CdBg+Gxy3EWDYJ5ISxFn1tE5bgpJpVrElo8nnmXhbFzmk1FAm9zu9WIAZa837TqX8MPUJvto4Ket9HA+PkyNAKhYJLwEUU/E4hqakoYjTjmVmqZJ46wJEszkVjFFE8nUe1JVlxoMdbU5z3CnU/vjXrjoZUS/CBMBGjzUb+xrojTeb/6I+XqEcf2gs64MicCalw2CJ4kitlbYjer3cln1vMy+Vfgr2VKXScm/PJ3RvW1vlL0v67VJK+HyUnF4udQqS4UEYvGJAt8jSg9E+Tm0Ow5rcjBRme8xF0WxlCecq0jEi5dPNJ46BFMA5DLDPod8hmIry1tDgek2gaUppbggYocbTgBGR5h3Xmq5uQOONkanPPbWy4KUokc9jv3UIYvUW+NuuyMyZGWs3oMf1JDoEbgOnWMUqWdQFvZ4wTRA2SiRScE4RAzmhL9QHDh0DXf6GE8vZwkDitZXgZNvvrmNYw4xVsD3tNnPjCVUfngIdmDWkFiU0OD6aOlqs04BFV2CuUW6YIukjfzDCsY+1BkEPqTYa+EcTqs3Y4hFyGwYk14tFPoPvGdkrCOHZdvRusUg3Lbyu+71sijJs16PxH408yHfQZhGOU+QLcQP0J3brHXVthxJBCvZt3CeI4TIUM32amUoc+Px0HV+RxIO/UvpJ+hTMnLtDviMbPBIC4fUciAr/L1UE01KWrliZJbITEZ2cmhORi4Md1eqb5C2ZHqLNAseSF4p4HsylvkITY9TmXkJE3qwQiUQPsX98UnqGbYtK4RbqlhLDqjI6CJl0MBporTXWNORXtI3gKBL1KVj7GOQiR20E0uYhMfqRewNuuOwB7oxMjaw3dLNcXrJSAx1E1gQTPiRzRsMtl9HpAnqi5IEv20iI3qUqzjYmpRqZNtnRjQYjVMvLbMJDy/VBKYigISG/Uk6dZ6YASWPZBRINN1U1zfqM0KMRJD59RqmXE0RHTS85KMnZiq+F7dvv2iRAizIyR4qUKFkcTE8qUi9KdOtyXCrEE4egimAYxWg2tSahbii+gjfZXfZjvBTPWI8oEoDKkOAyUoif6d9oV7c27xJNEob7r00z16F+VnDvX0bCEfYtMeQQQjY+qgbHguR8ygEtF3rHi8V27RkRAKBWJ03DjNsSOsOfRMjSZMtUrTaPeGIjOTLrIUAizFWGqbSYdGRe+A2r85CONSh89CJDT8EFdZ/FB7bcAXIu84XuWuvOkTH4NWCIzNltwjmb8sgjNV4bfjiN6BJC/tEm1chHJoabzgrq+Vo6NQL4fRzvBEbHGNbUxE27e4QPjIUl2+iBpGRllo8TqELQx7LPSh3EqINDSZ1a976h0I1svv8aZuiXJTxPWyaQBraK+unKJ9taBWJ9VReGJQVrEFtKbjURefKWLxMVDGM5fO5MMALS912s9Y9bONrzcXWTXmIrrPyqqkxfTYURu9hgkhVEgpSotxT3SLKCcw0jn4yKh/8uhkDkW85+ELdPNX5waOZjMoPT0DM+HLJzFVQnqKc/eo1JtVfYE6rS+WsIAXwlWJ2QIJTOa98klQwvIDkO+ELRFoZdPwgV1ulVil0eSZlvTffSLsEJ4HSnYqAHqmOkEEw6FGeFrtrUGIutMuT/XaolK1oVuP0cc0HOgZH3mKYhZkjwcZxoUeOiaJrDi06JH4NxxFcrxGpH6zgJdQInnqxh1Jce2kW3bpXEezt1uQUW77s234lloJSNiv7lqAjNvAMzk3Hue04vwER1XGajtNyKsxb7oTZChENmDpwsJyHm/Af/B8+wqfNDXhVpr1mW5Hy8a+c3cnVdylaN0gTprxWpJKFMh0Lf7okjpJqp4yqjosUH90m7OhAOA0fmIR59Q8pIAHpT9EJ9SLMPqjlY46vjHLs9dFqJFyXtUI1ScNTaeMEuVJnhIJjQ/xYxk8QasHFCFhYctdqk0dlwE4geepxDsARhU4qPH7zzXoHk2Z1Osj+3YV/Npx78M+bP0JnFjOcFcxlrHkS9gTh25LYglPLCCLE3sLhKA2HIce/YbCtcBbgjKYsmGGGd4pQJQLKJIuocTaeoDJNiOW7aBoqy89WVLdQDDeocbOitcsmCpKu4YSz5K0Ebo6AAc4Jv+QyUWnog4TlJxUnm55h83opXTlB0WhKUFcj2YZAJT7L11jkoXdJRjA7q5DFW8Zn7l1nCZFyenlwuyOxXIyxEl+ZGe5oQAvT71wwU0ZPqd9QUC4Bdh70MqnsqDCdFiCWtSyth/V+PbwEZS4s3ybbxvKVnXshPMPzScsGKj8bDnNVPnL/scuWEKXqDJ/VqMRv8qLhdQOndJ+LwYmMfjMhRi0MsG1hvI6X/v9Auw7Q0K17sXMGb2bvfM6wk+ucM2xvnjP8/t7OmfRDnbF+zlJA6aTDuJehh1D6tkDgvqiaiuhpNiRUyjfxwWTDjxoKNeCyBomw3W4f26/rYSicVX55b9A7CGfADU015vkKjHzNmcjXXLgW+eeKO3DOrUkfxzEwKDNgj/ORh1POhXaieGUaP9ENVp5P5V5x/+YR1iMWEpjo0cUuIZTtOc27eXmS9lVcvH4LC6lISGlvtf5GsJR5KkmBZuy6tEFfY++vzQgtto3X5HgKO1Rkfco9/+jsT0fmZNnWHPZb59QUGVkkVQ/zGRukLEGw3BnAFsowaFSEpTSdC2BK4w211ajy2YJ38exZqLjm8jSsS5TnYbTWE/kZBpLUToT0Qm74+rf4jx8sABhShb3+LWUsqJwYNuDNxlb0iZoaEy31oaf7fP5sDfZlzjzVNVE7R5v7EFNTrxmcQuQNPd+bMTeRmIdKJqSn8+cn2qMj7WnlHFUDkbjt1Jpn6fUPTMexBQL8h2dE7z3a2d7f57ojLhDZLDR5CvhhaFR7IEdu6dg5NUQFBTrD3a4/TjHTOalaqSAdeUjBy6i/9TECSpvjhwnm8eZ82uiiyuEk5NayJvxzKDWWUsI4SX888gyjPqkWqdJRw+4DlX3SGR/t/rtPhIMo1guUpmSdLk7FDyToaknBbBl7upIaB5W4woXd2UZnbeUqJqIFtOmyJfrIiRsY4RRyjs5gP3CyacNqTgvLiPtpwDUu2COWctwhc8veZtUGcOnzvGh2b9vWiUZLwyVHmlZsX66inrkis5PpBl3yGDTUyjImd+RxghnP8CXEogJm/phWKfLGyOlCHoBiRdoJsFJdHBTrUma67B+7eDJ/VtATY26YGkVxKRIWwG56QqvPX3wRgzwBf1tlV1y3udxsFik0HRNlylRuyl+kX9BFAH83zV/VbxmgYd6gzxWcFmW6PlbCTqwEECHwFGSdkNkplad8nIfoQue/C0Ol8RG4MzICFL1J0IgmZ2ab5w3YIFaC/LaxWsBMY63SUmUss9YqjqlK9cy8fj72SvmHTccLHuGIoXSsXEjQ+E2OJrpN00wWqEZgobFH+eqy4kBsSK/vEH0I8KhWzORo7bgCHsL+pjdSRSC8C1MsF+Mx1U9rvIJOgCSlFOqGpJl9a5kiKr0525YnztbuIdxAkspin2ZvXK3TDJ7w0tTDMDqk50ZqoiCTeexUEkKZpMvozkdaKlLScK0bKrYADT197NHURkEmRka7EUU2eGZRv3dG2wjuj3AhTwyDXPNLMm9VRU53NVHjCNgvzTMRY6S36KXkzVlhcFy901htmSTsqKMPhyBqMvOyZa8XzlA8FCCxoGNNeL/p/GQDOF9FL+dNQHbSOW7JuR+tHsvzrc6EJj8FQlk4m/qpRS0zqWiXxMJaDh5ZbNSE47vm4P0OX2AWrOD+ifIOe+tbNa7Ml2hE53B0njoUy2a2QxnbzFJd9S35LvFFFXZfM+Ro0Tgj9bK+TB8lWZ48Dn4OTAJHbrnKE4kg3DmWrn3GtvG8N9QeVHhEois+XGxpjxs3DHkXX5G6j+ot3nTW6vdRoBTQO4+nuKZ6ay7VEmmR7IR1KNwFZytrWPXb4nIJt4bNnLdMQAh9lo22ALPKeVT6Y0hPO+JWzcxt5iyogvQpHGiuU6R5cKNPyYmbBN07w2QmE8Xpog9iICqh1WQytfPIaGpfzPXK2aFKPqGG6bFo8APQxqxsMmiIcYGL7bA8Qo2sgHVZAsW6fYzksUsbHGi78pmXUSYIo7XTEKGkdI1iapug2Ua/lzuyoFwUnHp2ZQOJP0fHBc4Ca9hWcAm1hM8shxCig0aRIylfPWMOldmweJSV1WO7fkrGlwHeGS7g7tFxc64nnsrEOB/YC9Sr4GnSVSE+0k1huYEsSPVlOj5BhurdgG1ChZyCiRQVvoiMqwXt6TsMJ0X98lCU18WGLl1WWzWXrz562xhpQhWrZJUsQN+D0JEfVYzVREaoyfCYcl8ijaNvpsD0JY9G1Z4VIeJoc5YWMRi+pkepYOA7h9JWmhkqOUBPJmRQQTqYe7mmQ68v0o9qpBNSOafCtoVyTpjx23Fdb1iYK21XPvxx0d/E7wpW8IbziJKAGvugq86jd9vVdzqxF3m9YCrVpD3zENcwNmZm1HrmpiL/gIXfMgsBF7oUhcDjMKZkm25Nd2JpItqcY+eZ5XerZ/AuaVOrJ+HqbKqVKVSdzxG+aXAO1y156quR8CTU9CnCTjO5V7XLOUmwI83ytbBYgozxphTIkSyqwFKbzJxa3Z/cERVfXMua6WtKyuriIpJ3YaNplkkzXK/eg/pBD1KQWxel+Pwzp/lJUWgL2OWqweauuMQ1oiz82pOcaMd2h6dnRTpe8N8ml1UZOSLeMLLB6uAE9PY6S1QNaZsVFzWMu2IBvkxV7aoUFuqZ+qVKfYc1j2XDiNLYyZOpf5c/VLwvSjiWm6LaWf2ahVWvSl20bCV/sNtOKTcAQa10T0sY1d63skH5ijc4/w2Q7iiiQ4gWHLuD5YabtR0LbBJeNW/R+0ZV5+Hr36K60MCJUhMM8twQKrCs5/UIO4VQ0gUxG1OQ/PD7/3Tn8mBrtuKQ5iazjpPlrXmdVc1EeSGBAGdhVVrGLD7B6ZhSX3qRkU7NTC4Diy34WXNvsj57otIsUPoqEDUSqs1CwkOWcH2F1fX2PHisrmOaYhP0U+fSEjancy+1gmM894UOTs3Z/t1WokJzCsaA1d3N7goj5SwCU9ufJqiCkLwTOS0In8pB6aUIO1qdK8oAsr2UMgpjnFS/vJyjfrkmp3Jv7qZajIpwvtSKGD98SU8ohMDI0EXNau9qq5QF2otEQZCMUt0jf03cmB8YZQayJJ7fHyfCMQs+lE+xIVG8PDZCGTG1W/KXgCb6Z9Gcl0lTS+jipcTxmXnP/gqhWal2siDbklhb4iMUmt9yVoV9gDFdkVV6VJYKv6icdUFWFO+K3f3JYruryKC1BEkHzS6ZFi5ABM0L5uVxBTEUxEFVwLTQ0mr5shx/RPCrhp26aFtVF+3La8QylXTpZcqEdhw9bMNFlHHfyto65wi5HietlCgaBeVcbTh4pQZWvsSnQdbq6KIajvV9EXqxqOJAE7vGSZU+GLOjcfwEKQnNGUB/HInPQeFVEyoluUd2elp33gpb9vY3xPvZCJOOMe0qXgkWvWHxaxYpowAtbDs/8e/hzrO9cr7ftwOrcNZnSoZfGopgxVUa8xK9uobWnClIWVUhFbkytKEQeaDnodkUmqnIxZB8abQxgpa9mFw1Rf25vNl1crLvUWVOguHUbnzkRlw/0D3mJLL0WemfPM7ghScF8yMkqPnmuCEzwMGW72ByRVqGciK1CWMetjJgnCdEDY/NasEBhU7aOarVIC1e1gz+sFC6sxBGSMhA2hipiSlV4XofWDojM3ViZ6ZO7MzUVuK6cayyavXHHioseV4qW3VxWIuHtRQbWqmhMUxJze+TFf6RrYrXI1pifde1Ktp5Mo0ElyIRi5DM7CyXaBEaV90hikvyE5HSCSN2q1KLqlfUzEW+x96jvWc7W1RaptNxf2SjoW1Js5zTRiErjasB0KB49x781ORSAWahP9J0UsIO1F2LzE+RkSfwgVACUpmr17bJbURplCIjiviGs8taU21hE3Y1upPFDp4j9RBpoYCrCCOqXTIAooJF0kyx+4apCKfhZEDjC70Z6P/1MqB6LmudzgsVAE0hwOM0KXhu0I6T1VgECRvWvnbdgbljcVtrM3Q1RXXFHc0brs1WwsxVlKwJRYlAlQW44Ep1zR1TXTNDRfDWx33gChySO6RKrPBZ53IZpaJvNQc9RkdKccQ5y8REFLvgxLhdgRhJjfMAv0PpXYqnXCh8BB3hciGVjM2j8lLKp57ds1PKASCJs4rTVaCvYW/otcVYGmpKVsc6fkbHKpZj1uw5WCo4y4lNaBjfuzpGK7CvTRU7c1wp4Jh2fwxslneTTiZgOi8Uy0MWlTNVLqcV1RrNtO+68qfMmMdp5qmApMcoSwSZ6tIyglNa9wr0F5kuhKHJrBUpDkSxyKHhGGsWCHUL2tq162H+2p8H82W4aQXer1XgfaeE92iK+FF0kBxkKTWRf9vaSIxPLWO7LEihtJESz2osjJhL2suTWnUYFZfKMaUJlw4dh+QB6Iv65nTDcF4WWTSI0hXUDUdVekX1VcT5riq9SpIAF5juj4ERmD0j4Tr2V6CR+1tWquFxKDajo/hXqnl7S8XDDFJFQd4F8lClcZAs7/ukSsq7i0l613DiMwOUdcI7467ldw1THP9wtLLGOsgDTI1hPVgVDzhQu+hybwReq2khKZRfj7rQczHOhB/9KLf9fP+Q4g0vbNGVt7kqq6TqCWCYSoApgZnNzfhuFn10Vcahqrvcjipp2YWkE+6QrvvKQJMFKjwZaCoD4+WtjjdWwYus7tqazQ7Q71W8gB5VxV4YYZA9TBWTApo0pJMldVTtEALD0tMi6eA5IW5KVmImzbA7k/9a5+V6dOS98joqU0JVDoMCJfkvIA/kYR/uwLPE17RFwTSMw36oYGq5lj7R6YdEOWlZKdGTaYhAbmqxhD9K0E9ApwCKbK9SMXUe7aiL5IQ/N50VIi48/rHeK3LxxAA62dAMgwnfs4oMKUujztlAkjFV2afZnB1M0LFdl1pGYaNWleMcRrv3xLwFXYYxS5RPtvqx9XQHIdlH5ZyFjVTYRlU1IuX7UBMtd6ALy9iVWYVAMqIynzKzeMqCPlVsrZJlErTSilBPoyZR5eDV5ZgK8K4nUwrOP6bYIjNMlPz0D0LtoE+fG2Y0bRaa6zCTSojMFRQz6XQM533tQg4z0N9gfT1Z+HxGRIo28IrWhk8w8qM0t2JMkjEfOX5tcFmBI0gWAYiIMy6yoHXrmReqIMnqQdhgzDAwoJidJLYTehTAa9Io6Vv2l6BUXKrr+mSK3/tboFHKBxspFU37LcmU6kgUSyNNi6qYZnUFkqgnqydU0Sj0AzaES6WxNmq4Vc6ittbaXxfJqgkseqoD/MUpNX5p2CkA3jmqSB5VlfemdGDLgTjmySR/YIonTU48UfQ2GwdRfl3fYH3OC5J/KU1K2UyMZNUoUBV4af9MFlOafWLXTbVo284FJSyMlkWxwslYWbCzShO2rJ2GOzMeNla1+ZrOPoaOSMa5YeGnqKi2wHvNQgynGvInG6qb2WCoLiRnEC2DUySkNsz2xQUDk3qoUmxyrIhRcVBnVujSZ6uuspHMwWR1acgKj9kj+H5cZuarbOTYBVYTJ61vVQHELkioNE6FvZzK2FkTZRNINuYqVdpdf2I4vaK7f9/DsmxJucSkq54WakuXBveDl8zuZUk0A1ytYprtamQWeC4quqoDbORAdRqU5KPpvItLvyakc06H5YoNt1iR5Oouk3ReX+T0bBN0T6jEqqg4i9yVYmzD3SvKsMlIyayd44KTf1o1rJJv6+RzdhSn5E1UXE1nSzIThksxT63NNCyJOI3rwQS4ireDyefvBBMc9low+fytYILDzIMJR+tpJk+sb7d604sxInVqGZFKu6/sNDPWKUnV85iGlcYVWGSroG2xDkhxiOqTYra6PnrAZeM15nE+lXqfDRSpHm7S3/QX/UDfNzfeCXdoTtdBnmQkOfDrYxANdn24wU3gj/vCMeCaZ4vAdwsBtfJOcFJz+HPBSg34FvBKcuX5thC8GkWA3USA3XbJkmcRqqd7j0uiWj0cRW0eqtEMLyKOA0iKHtukPJOgmrsPcm3X2Qcxj7fahsQIoJuxDSdeFhTnCqwFJTZ5mxP/m8WxFQjc20NJkku6F34D284LoT7raaaxWgNWYr1vdUfQolfe6ZZAXa4C+KJXxMrCN4TVffX1oJrMB4H00voxdA8VjDKHuPoUpz0cJV3hFy/9MipF/TzxpaXGrDNNDuvAW6PkY7mYYCETzIkhYsErtA4irwZXrBAlMFvalpQ48wJadFVrVdS64N7E+R0rgk+ROVeUB0v2VPoVU/pTw2EYM0u2yvbHYn9YBuk99KeF1euYquahu8A2fl04CBbJTXMBVz8QZVcpFU1NOL1YsjHZmTCbH/1b4R1ueYhjL+eqaGZtMCrtAdtqa8Bf6nVABTdn90hm58W6DGPn5893H+3s7dZkpCh5PSsnTh0S26jaogW82Ct8Hin9KYq5l2pawpeeYiFk2ExM3uPujN6MIhx17iiCImUy52qA2st+eMJCPrnJsMF4wb4GbuWkG9hTc8Z7i7i3l6LGMDvs2ziMvfPRe4vjZ+GnSLhbASkjZW6rxlJc5yVkELGZ9ypmb74uzdtwZ6sTi37GvIN8H7jzvIIqHYxXaxyMFwkKwVI0Kd6rVjQIO6qX1EzSm4tvUqx7LK6+TCwoc4ouE/IOlKVtihouvGYz55azAhfgbec3DnDXXOnVrIqThZTeNebDXJpW5VUq6u7MDw3Rui5V9FvlSLTSriqfR5U3VyUxVS98aEWM98nzgF0SyDNQ6wg/KDgcBMoJg8rUiXilw73/Y3sXY2IohCnjsJGrr6gslyitxhXEU2y+dMPC+uriTzJtn6o8hZ74MHJCVjC8ElOPYnUi7C408hlxrFWQZkZaQCwLpVIotShLZ00VqbY4G0DlhHSEaaSDHs4JvqX2IbGfIdvQYBu5yN3ewL9a+Bd35qqzN8bAQJj3a5WpzHXdX5AlUOZ/pc6bRCVK1bIwWAxdGygkBYsRx4kyE3BJ+ZajpsVeViovKBlPLnIxAE6a52GJO3pNVRk2swDlPMMsQeVlVX94R9LCm4u8ZoGPQSahda3XiQ8EnoUTjfFk8LeGO7ebrHp9BhlnoOqMQCMgvABEfQsZ0KahTTRxRQghd6JwgMEuxNNhMEzgDtlg/BlN8rMk3iCWVUCBW8/ZVmkTs6X3WXKPgQlYu/bvMnfZ+Tt4hcVMnc35C8rmDH+MHM9Ux3YZflpuqqLj9ubks0Q4a2hM7pGrvB7zOltAFpkpicxm8GFi6KZa5MU4UyBvDSvBKti1sgufWCDjpi0mRuaEGu4EYUzeOOyTM0z8Kl5z6I08Sok/cbuO63EFlgQ/c6mPmD4D6yr7gG9/505rlom9iekXlyexYAxYX8wKZC3NLPRetTAt9LWkrFblF8m5zjYqpcWa2TdcFBRJUyle51S+dHLcZnFF1ckmGu5napzKnhJ3xuILqKrWL040UlAUI/Cgsg5N10KXTepzYIgWR+KlY7O3MgRlueTiWzXAU8ddTGhgJtqXhV/yVGQL7wGFGXp5APxfxdGaD5IsqDxpmgTkylWodPqL7/L88woVNNKJi9y4S6qvInO60H6BqarFicsJR2rQPDbQVsEZSzTboBgGAb3SGVGYWUd/LkHGgpO7IW1J8Pkn9Pk3+PHmTZdTiREIrauIbwrl7axh3FzsuuL3rZrttihm31YFuYE4m+q7TfyrJRDVN0PCDISY0Z24HPHgNz80SeQQ+eEPqQYMV3F55mWJYpOfA+c+SChci0QGLNjycowam8fbB8+eX/0TpeuXBTBKrjWiZssL5sZecAkPj91XArgFWI4Zx46Xg8AwVjlhfCqTLZzVSTZJRoFM9rzacTogjY9gnsijRUMU5DkYklxm8IEIzUAnGT9QAeMsWSYiwTLmmxalVbh47dgDJh8D8pMuWnFgQSL1dD/sRxidCSI9C1lRRCNiMLK4VhA+XwZpHEQcgX4gg+conRyIUgFltoazduLFpGhkmLBPkEofqzxuYCUwxBK7DA6pIoBfrOpSqPxVXUBRTJCZFS7sByI1VhXnolIVlQF7NF9s2cvnFKAs1uQUo6Hch9VujJFaDIVSsbFWTcVE9Dm2e7dfVCOJmmWU2sQavlAg0JyK9GtIc1nnjFaKFQNLM8bc1CIEA/EW2JWJAgfv4axCIdTvUdcY8NiEJrD3qYBms0uF4FPvdOhRjtd+AiyZs8LJHQjt/GA0To1ak2KggfsQD61zyaCYSgvXpYbstCu+ESRl9YefpNNNV5Vv2g+ycYQnXJ3+Q3FuSNRHm/tLPlh03vmAO7tjUYGDtQL0+KIfjITHWnjuCSLAmRWvvgfRP3GSk5dBrpLJU6qnfhD5XGMpPjlNPbJET1RmAAqfZmJC6SgIlsmLrlkzieiFdxLYWahfwNpg0144k3K9pCXhwo++wFTTnFLlppjimNNPzKujVCr/9qWIfSsUDeWzbonCJhYjCzIt1JKkjDTl31VaBXY5z4ue5kQW7azvjP+UxKNQs0Q/YXhWzNDICokcV0ZRupUj91R6ABNcuXcSAUODd6kqgCgUuOUoJeIx2Rj7QjV7YVlSy16R4vZBs6iunoDYduH1c29IIceYgQknIoNWiIALbY4MKRa1aWTR9Rw2xrAgT+RBEH5OZB8Gvk5kJMXXtu+2neciNlkkcsTrCAMO0V58DrfJkhYf+ni5AGKIslMkZYhkXxgLLVc3CANyGX3zjcyuFlyQ67wIZDTSceH5Q8x988e2CVhjTKk53+CALL0TRZJylI9HUdAYSQQl9ve8SZqPc+T49av02wh/0/hn0DgZR1Y8LkhevPQlljgjtZunlGnGrcoOsxNdyicLcRvOxieJ5aorbEmfbn3e23n2+eGne7s2p48JYnPhHs/xM6MIdr5RaaVVP/lhBs0mjU8Pnz4RL5/lw6jRtLJSI0uB1gvernfPS23OrwAzVbr4yDIVGSSmOyMW0z3YerLzeMudFa9pdNVO+ZagWmJNWz0kaN0Co30RN7CunqoNAseNAmWvviNUbbrN0krE5s/s/Iv46fbTvf2dLWfriXO4vf90Z3drv7AyKsgk3OyVGsAcoR3mwTBrVKp90j6Ra6eh6G3TZTQjGid7MqixqGu0VBlyj3398Pt/oy7EXFjZu6i8b69/4Do6j4262vHGoIMqV7380+Xm1J1jAIS+4JahWU4vaeHTik3h41erESlsz2wcy5kYcJ/lsFhEOFbi8nsm+0nnr3AqcA+QZd6gSHssL2GicbNwQqqUU8sPs5EHO5pPomADK7InaffGfe/+x/c/cjdn4+/D2/jq5rJ2dqNwGjqkNjYuP/TDcznGAPil7ura6MLJJkDMhyvj8EEU5HmQrpDtIT7trrbvji4e8GQuP9nfOZg+WC70CLt1GsYrJ0meJ8Pu+ujC3eRT9/A2jLVZbP4Q89CICZx4/S9P02Qc+93Lw53dw62pGOqG/9Fg4H/0YOT5Pk2jA5PEmZZGP0lSYNJW4K4Kx1n3HrTg+XRBPnPwtQe4ypUsfB10V+/AV+REV86C8PQsx9WVOkR+dxAlr1Yuuh7Igg9encEBJXgEXZj6yqvUG7mbl2K/pw9vw4+by8eLERCxLxIJCZN6zJyI9o0KuYLO1KKlNJwaRq157eNUmKx+pS0uoOKZadjoxVvTXUX8SgDrncuf7e0/3p4+KO/ZcjmBwrLae0TQ1XW9s/fhe+eBgTs3BoNBZRe09QNvGEaTrkLx1gosKIItpR9amRdnKxnyUCamrNdMSpwERlZ3cxnTQwo6IcBFGSMZ9d0PTVm1E+dXX6dwIECo+ZDKFmOq0N7jbbzCgMXrUoQ+Gz+Aknj4w36AXjCoXIcvj0itgq68SRJ1VaLW6VLv0dbPtnvbu6QMFn2gZUd4o6CQiP+iUoQ6NJ8h/0rVyYCZz4quJGLErqGSEYN2OSWqlHmUXcRJHDaEGJWVd3B42t1urTIGLqmWSqRUFEzhISwN/i5Kk3RxHuGdbLL3DZWQCWuPNYuFYck/pqfaGBXBsMNmWWIdD0Nu2PlLyb8n48EgSCsesFqvNDPUaJBgV3oSBX0stJT1RNZBXliPXkh78mkDYG0Vr2HuY+Wd/mipmoSEmeonZQnXqzc5KUsuK7thG2+pJE+2V5sBvFsiwVnJh99os+n0DveebfeYQbGvuxsgNgslK2aTCGbqV7W/NtW8SwGpC50JZTFriB9wybxUyJnwPuaPTDCDDyYdVZpX+aDQFwYYR8HQc87DLJcuNwlql9DLuz3TLe/jhjx6lnoTv1sBroQfpCF5R/QwiyvD2sXWZrmliC1uMkFJ7jHqERvqMNMjvFWxqGuuFpELLaIQEJtz+SM/7OcWd9WkXkK/V1YTLy3CWynvOiOrd18t144msE/z5obzdOtXvWdbB3sH3Xm7V1pys1lDJAr1lbWC9L1UWAb0TctJDUXlJSxNrxFN5PYy8yGqBi2jglah3naWs0grH1eVh5RWHoCLld1M/VwGv8wN36d3jE0xjN+InjJ/e7MyUECoXjCVgi4yM7s/lT91fo9Gtu7ZfYrMlPN7VJlXZ/eH6ePmd8bpLGf3JPNIze9Np6Kb3eNBOL8vTN40uxeZ3WF+XzrJSl70Nt4fx3k4DDjkx1UoiqQY88/34QZ0mw4XdoK7oO/Fr1ExaSoNbMwpo6nWFNEpyNsiqV+3HCtSo/2ZHyCyqrQyZbLScjpVySilawm/eFxU3tusktEqY++tYto/vS5K+bNqz5mdPDdAknQPAiyIw/nBSOO8fKnfxYCCZayTlV99JbXYE7ZR1WiaBi7AfwT376Vkso/EDI+nbbc5w/WESBeqzlRlgIG7fHkxXXZpyy7s3TrCRDbN2Yvy0kRYm6i+zyWNMIUFWEuEftQa4/pl8eozuXxeZTxvmSJiju5nOCQ0xWYJW+0aAhW3mpztXGQsI6J6txoX4ZOdF1S6OolrAL6Oibq1RXK/vFnJ1B+pcRApK1JFM95araQ8d5RPRsIvqih76HW2Pd/Xi6mHsI3QOhG4enUqWUNeaa3GdOBeljS107bzOeUyZ86VM5qnbbe8pWZC7r/shlrDoQZKDkk7MpstWnC0eXycJN3ZBlmi9OjNgucUylZl/vFtcLPQNEYSkCrfuyqEbDnWWq+P5iKupZxDBsGFDlimErsnNQG1zbV7F+sVrcIwRgeYgIdGnkni4a4cR3AoLrnt1GlcGp1Mqf5NCdub9eSQzHSpdCJAsYs8YGCEIDbpK6a6mUX15TRplrKLiqkY2c+L3V+T2KqU4YVjGQWhX1QmmFYlwHM8Oq9BCNK3hzwBWfPPxz6U8zjr0eac7+ozPmewhYpYLnLCDeOtUh/Nn9/dRiXhqVHOLDLZCq1Wq6i9KuapTce+in3EebdHyajRqeDjDAVYKVWiWUTL4uMWYQdVkd7zIM3DVE69QVNTNbAq6JdGbuW6WvGOsWA6OHQSBssH6FA3hDPpXvanbk0i9xarXeLzEGaGKbCcy5z4rKUq0lGkFed4iOmo9WFKLdZJ8oSP6862q3hEmmuzfMytPNyFo85er7hAg9DYlwrDkJWk8LknuE82UuppKXuvNGJy5FFh2kr9xs35RjZSd3NQAD27RRbR2eyNtEkGGOYIxDyCYySNuhWsiMi7W4DBAtetihtQqOeHomBsm1P40VYchMZ8Vd7FBa7Bt78C/QBXzfcG3GTyFmtqkztfjDqhZMFeXwVW0RgAm4XOpXx1eikGm1J4HKa+HaNvrDMYBzOZx+WDq39eLkY9LO/uLVsCgnpepYQR2hzgb0RC06JnBmFRUm9aLHSDjcv4YeU8LWBJXc7wal3Pgnh1DdxSWSkLPgbXwLJ3w7T3gm3vCeNmOncMXERIwrqW42FOZuH9JlKkomcyKt2rUNJ4QTYDcpQOw/jqOxtdzayxFWhb4VdVgYZltaWx/YSDkp9vyVqhBkpWCwFlCUBMsxx5g9bAtyu6uu5WMkkVVQmM7LFkzbgU65iK+mwsP1bl0eak2TUVfnq+vNaYc1fcMYXik2OjQpeuMWhVUSMrz7fK50a54YsJu9lu4nuG2bKqR2HKrJm7mozu13VuLZVjxlTyby7wyLViMRgILhpxtGSm3ZJrBEtkhqGFDIvvzdAi5Ujp28ppHQqcFzHNG50WxxBsdGw/yEeYx5Z8T8VxY3yGLyPPEHJa7IPqYfwF5i8GLqzvFd0gVdimtGEXz2kN8jvkjlIVkTafHV+vZMdVhluCheiQc4sI683cmZEJv2ICViXF1Y5z21lDVLzb7iC4nnEQ4sTJURdz9XUsck+gkTENvMiqJqi7VMmZlCmfmVmz4hBXwfXIU3yUJjkeqrRV0ZmIXpEVEc/gHPe9LA9POabEi4plEccZ2TsruvLD81AUaBCmSg7AEe7ImDZ1nepvhsx+oyt72q5Lt8JRpmHWC6HlKYaaLpY4533ssZlPLs6rPB2tI1uF1ex78TeC0zI8A7G4US7lqxYlXUSqrtS/gvmXduOt7tD3MakyRZesQkEAZmJcJwbrSebppAT12RS0gMM0hAynadZ2VYG2JQRZsKtKZDHzZGA3lcH61gDtKHlVCjSfEUmeY+h3C9Pr0N9X3+I/q/jXeWV6NJ6MFQlTHoHD00WceqCD8Tv024x+7bJWFG+UO58B28dmwqJvcpYtvQtKEuCadYmr3wcPIdXWFqdr2weMA8cNMNFatWEVn7dLp7bcwWclYzY0o5ffRXVodFCtP6QGMlnS+1QkWmKAGbhiTKmE8JY3/l9kfefAY3syxHoR+lveyedxYSPrJaIKw4i585hVlILsK3Pxlal+A6igZNS4ZtzCt0NhmSvlBAPFDRCKgTP09pXB/5SlVF2NGvqL8Jc1kxITUuzS0rXmQ4m1Z0+o+jhzXrpZvhIRNqkSbK1WJyHOxDOaaUojn9VSGsoqKzDBnG8yYpKN6dHLaeo5n3AFslHe4ApcKyAJNavbjTqAR+RG20hGInfxjN2/4Xx+e8/BmV1975GlF733RA3eRJRP55RwnBFOCLxKmmgXFyDOyeeuxvnXv13wvrK72OMuZr1eukb9oOZwU9TsfHhTs3eAtximFt5i3tBs6foosvjiDLyUmTvKF1qPkm4gijVRIye++kF6jbVrJFtw5Q1a7gY2bxp7vSG0FfT4J/y4ajXzE9nPXBgKPuZTHDoSTuc/3povYdpd2ueH/AquQPxAS62pybspGm2KtzblW5v01vQoGR0vlc/NLXFuClSEkugUHogF1PjH4VC3bFzlA9OLx+Tg35BQmFs49jqQM1ey+Lxq7mfjlZWaV27Wv3Kz4hXCw9uIfJTXvoR9dEo3nE5Vzc1oRPVMxUWC6ghp+S5dfHyoiRXlo1Fycaviwu5Xcu1y4EWYLotJc1arMRphc5tWqsF4WxxhfPZ3FtxuOM9EanmqAQsjwFKo3nWYoSaVvCKAmT0NIooeF+7sPsUXjoeoBQ5jozfU7mAZtKvvY9ulHVU/gJ+BrdSH/l96QtVll2XKmChsOuudj+/R2cDfiDJsFt0AF1R7f7yg2tsNMiwpKlw2uW5DYGQ4d9Bh7jT1MP2CtDUkw9GYvaor0qRrH0/HU5XZorqqcAHmwfTs3QgwZ13gDMcRCBHodo6DyumgKg5hyjECFeMPw4iDD85RV961XV4waQfmn2CdrZRAqqY2zgqzoopVowBgcPUHWE2zqgKDc5MPquHbzQxhLTc3QwygTMfmZaqyN9BjkZF6pqZEvoFaCamC2VpEa8JKSKnqmJcMrUZXWDeBRXQt15tApQJv1hQOKdKrPkX6zGhvu9MnQC/CfOxXy2PX0RDq0304GbGepG6C6C08d2ZbJ8A2j2etE6nM3G72AypU5qX1/VCo6Qx1sUYQSTqBYY77RqIcoSFqaBVRS0Oh5WzlnH6KvxcLDFA+a3L1ESF8dVskLFGClrg/rjlRhf+T6ylML8ek2y89NhzWmgVrLHGyM16scykXPW07Cq11MQpP0csav++BK+gH2akwfxe9Jykn6hXNS0DkW6GZV3XWD9Gnp+t84a53vnCRQLfgc39M1RA9+KXGOros19WAF93mst64Nhpu8uxVmJ819MF1+cDPAFDjoukWNDx0WqTDlKkkFBFmnLWGGEMuMnId86Oh5GbWkkl9mcFE3tKg+0F6klAo6y0MM8XaUinXeYJvKVmt8OtN/KovwrQyG9htbIRGKD9MJTuIfXp9clPC6y/0xyhMV75OGRT5JhZZFHHyrrAup27poqs4HYuejIGLVi1ACcyPc0lAmFYcBeR80Bet/AT5pNIRoQoN7CNLOefxm651hoVHA5FMXqaqzyiPVVTJRbiIAGSL5DMxMQ9CBgdAvB5meYgjApM3jgVrxG/Y/S1jDu7Xfc/oBgB+Z63TcYULAH5uL9trWkYGpHhpN/i1JicbO08ieIqspixGkxS894o3XgNfL77NoYIVhSLV2aF0xGapiSLSS82Dxn1DvWYj/vvAIQAN4U8yml6PlgKZwHdesosHnHLp75G1aJvNEsTa06OC6g1cMxgcn5MXRYuzwsUmxsSh/KUS0FhDnRItAqP/kfN5VT311nV3wyplNXM35sqr7+UmNM48ctNpZJRbAhoM++Bd7z7EctL8IhxistDTNcfeARiAFvrCKUdkAKjsZIVrR5jwXW3+GIi6u3c9NHXhhYh8d0BIOaeiySJHv/Yrw2x45HuExdG1lFm+7fVFj3Mo5TCowkcYvqHwsFmJiGsqbS1rsXgljEKLaPJlCLxQetW/Z7PK8jVm86y08qwKcuW0pGbopvD1szPpelHUMMdaxEwi1lDEIyPhGg+kAFORbkAlbnw0vvoqxuIpWPknE64hEcrfIGsPSbsNb2P98wBj3HK4/DlPXgOQ725TORVxHRTMgKCoQBqIaGTy3TOAFydxlPS9iF+qiC6mF2b5btRFxeqZiEQDRiRjORf43FBW3igZU13lSVvVirxqYS1HxwuMWRP4WuzT8npUD8wwbgOpGKxi962sjPbe//346mvHeu6MM3QNt2t+qAgg9kbkEiItLisiHZF0GHuLb3UiAuEpcvXpOSbNTAzuRGIXWlAirAgyDHPPV7hUSB05H6GuiznVMeR6aAoXdEUrd4FdrAsiL3ap2i3SaXUUebFL0WqRDmsOjNkbNlmkq5o48mJ3spn7tqfP7MwKFfjLH09zajWe5u90irn/Dyw51dazHWd09f0JCpldg6p8SJmqkNTJVBm9URaM/UQU+9JJmzYaTU3HRVEwzH7Krzn8Vv/qO3yr7ew+3320BaQadeE65bLIgfsMyPUwAKmOXZ8KGgUe2Oki/6NT5kb2CCTlRgFaVjAf75bMgfgCSPuLn4ex/qEt8yhzMqQuEWzMKaXtdk8Sqd8fZ2OUH5w334RYfBXkln4UeOmbPyIcHMxPF7edR5jl9wXSrRcOh89RISaxuH06S9bKxO9CE6+GBdLsvEi+5Oo+mAKas0e/AEnhhFwAspCSO1J6Jrgkh17UdjATMTWDq/M89F56sRGmCwI9FWYkL2xmmtTVaGabNvJYk2s/+hFnwdUfPBCKjbTSRBrgUlRTFy55ISX8MlM+WBp4TudVKiunnnP+gHQMskDkTZJxzumoFPkykjUjoeIvxuOarM8VfIcW9uSMjWRlDTsnWaFtW5BK/l3oaElFCiIHzMsMJUlVAmv8nb+Yr+3DEUnRI6/gAHiDMjILfz0qPwYYoYw4ZnE0D7M1o9TunQKnhbLRKCTmyDTwsWt5KQOVbaej8gZGbig992pxb3UNZJ9Vri9VLEQ4Nucu/dOHXhbSVKvmXxK5RFEEoyVuPOoQwvMganF65WGSU2CzCEPJglPAHbTO5eUSi/RegGI1CpDB69KIFBJJYQsc/w7nmuoiU5r2rCUqI5D4YZrPCCuGgR+WhEwkDOiKr+sfWzizTf+EuDOY9rvflfuVBr4jVpMidLpOTOnqUV2Sev0AM1lWbFMPPvQIIg3oTdzsyEcwilfG1qU6ybvAbitHm3n2dNI7iogTzUsxcanKgWc2Y89Ho1EpIZ46XCr/qNFapscTjeirxXFT9QzguUXTkittokolyBUzDePU96nObKoIV1rKME/l3fEFMpiI52Yrk2NKP7yyM/64j6fRk/SmQUEezQ8KDNuRSIOfUwUIWeNBFnugOEjxW0wSrM27SFcDoxV0iu2EQx9WCJClLTCuUWoPqQAkUitVfhZjEcbDmK4kVMjhVTiAq2GJwnzoIzAsY6JnHDhnVieI9Vzk3BrAPB2ELcfkoZpwdKHDgCZi1LRQGZYMJwWOrKcqn8AQ9T7berK339vZ3Xm0s/WkkL+0YyYs7bTpq8pIuuxiUUEjHSlXcsNkqHu7n23vH+zsFfOhAtkweySFltWnayc4xRymvd29w60D7Ii5HpVMGxsIDAcoqeBDdKHoB6lHzhVe2j8LzxPnYbgZZA9vh5uOmY5b3jwuFQowuj30hqOkb3RKDbrkJ9PH4sR93i/5XXUEDB5eWVh1DJPJkjuK+AGxMISrKyL/DoElIdxWQ0p3gLpRzv6DU8yuvh2eJNoxxRXpx4w1YzoqlcAk4NSUNBipYPiORGU2LsyYn9Q1WMBjzsC4CpSdVax5HFAtIukw41XWLsd79unW51f/fvDo+ZOtAzUo5kfB8bbQuYhVhjhxPxwEKXv2sAtP13kIjGiwGcYj4DibD2/TN+CxAyy8VL7U8xTYJ1k6aOLkV99T8J0fcE5d2ZnRFUzwt6jfksklgNUO1DTl6nCq/K7UVMjX+2Rfg/2CM3/q0SaOgDV/jWlCK5x9uJMsGG184X7hyk7ISxg2m6172kVYTkPkedKz+N+j/MGKAkZAOvfzQCxvQzxQr2chYV5cIGh2a21JmjgVgzDytK01YZcCYURf9Z1tFDrSc+OCl8ax4+lKlIYN4dIonJJbl/aITzHw0bBgeTKOV3Uukw4UBjB/fsdhpqoYHjJRKkVMYmpQuUwM5R3+c6ooF0tBqLThOMkZGf/+11eGIwDUbopyOUlsl1cH+fpxIA5bg67uTGheRGGxTBRDRHwhKy2zoEmPai1qK8iikvccEZanzlWHzb+rJdoKoaayPFhRxpk/XqV4MxiTLmND6IaMcsTpF1RaDnPQmL+q37JRFOYNnaNmNME0OjjUsXOT7DXcNz8lSGMD1616LEWe6iOqjyTxao1YFZ6jqMIN13RUHk2OYoDPKh4Qu4CcjEK0EZJmpl9htoU8Vqgx1oEzJqD5AphkI7Yj+m2+ASkU6boiUdDIN/jSlqxp5J0A5bLC+XsUhYJTofp/KBvJPgXMxGyb5YpwiCyX0F46OU652i50RxW79QOVmaSG2rVYTWAm9UUqG3J5HFwu7CI1WYwqiozzNvikrbv59gQUfgZGBI2VNkOsknBWhNZf/TvWoYcFn4rKYJJTJz0DnF+6yVnZRXzUOA5Yg/Tmmwdv/lgV94+qQLKhZSwspOQv5FOVAJYi2hUO80jfEJ4PdJ7Ty3jKRQ1oWVPOhhXbl0NFmSA6FobrgNirW3KQW7QBLc2PvqMdrDDgUm3uIzGTqZl1EzMMFuvSymQ900uc6NSt69LghN/J6vaXW4D3zja+PE1eF/MTvh1Dgcw0Hh0t+HHB41JJKSnpNZkcVvUlotXN95rCUbYmBEmfH85mgbfbGGeO6wbJSugr+X0yXfOVH5AKckZ3Zg4byjJ/9W0/RJomMmkIBYB2SiXePohmdBnqQitcvZbSUPtcea9d+yJvlarfpTxTMR2nEJcIRro6enO5Grxc8B32q7voYK4xGKb4PJ/KIZtu3RhVZbkX6F70616HMmmqxx03NZGK3t1w7aWnnNavsgA2HZMeeR5lNfn8tA85aqArz5pIsI1DVYTUWB1IqC2jVOm6FXtcTAW4SGfwV3VnQlWrE2gTOG7x6/M3ySB7XELxEjucNiVpM2Tut/UFwMNjJWErkNGlqnB5WhMldnBV+lAtu9SW4puxPOgFj4Y/7V7SAFO9xGxRDwXBKmG9jLptpYF7KD5Vb2v94kjkqltYeXE0hrVCOtNicbysOKk5/tWi3szlFUYH3qLFukBWEbybmPgecUSwQW+BIZz3sA5JhtfxFXlrSBqjLOm2RoGZgYuhByvOpbbqTvEHfF8rSJtVL9sjFRtXCORYtUCGOzzxlHmPwh208JOhgymINwHIqZ6lpRFiEOZBxZgTfl30x/c7mtNEiehTT1UBpiq5FDREgfyBsBBQ4Ar6nEsWG0l7KAREpO0BOW16edAYTWwnZdGkLmUmFbeOjRB8mfpXCGHhsRSmLAlTpNvRcSa3bxfjaA1Ryi3VLgX2cBySUI7ZxiZHIeJL97gyEEy1rSYrehzV8KhLnpTyaxMESeuHdiTD444ruwTM8r60jYMwxWPKiKnODMqgCIajte6xlkQrpO5QS9xHpg7+eEaecDmea3nq8p604KnS6pAlknVFAq9STUOKGqNDwVmWLErSGteSphk0/oiPwkAn3FMOwvLbZo5HqTsCjBbaKXY4x2gh0kefXn2HnO6S5MwxxsJP2BQlQUPnjGckrdrkCtSSuedERtqYmV8R1UQBo6KK1UmSk3foRHzXBdi92PeGIc6Gj1mhqLsfvKQ0dqytitCWAhxrHEaYndBBA75X8HDp8X70hIaNaHlJaSet7DNqY5YqxAK5c+FYEAGjQ8OfiJNLG1xg1lRQWUpEo9DspKmqWYl8yxQ6HU/wkVaRLOmK5iDS96RIX0KiZyHnfMVkzPD3m29sZPidVMT/jpAILRro/oSWRG1SlLCrQOh3A2KPkkV7AjyiomjLAFhRncTTCCPPEDVJvGw5DR2FhnNr2nQWE+OXpk9DAVHcWLWJsMihovRS0kke0weI3+aQZZqi5ImLtW/zdPNh7ts1fkVZXST/K3AkT+NuisVpVe3b9dGFc7+i8q1RrXQNi+5extOHt3O/XHBXD2h2qWvjigKp43BlmMQJVbptPQ3iKGmVSv3K58WiuoUiuTAZVQWZ4tz+WiZmV439aHAyGLileem5UyTe281dbMwi5ZMvVl6Ffn7WXbvf4X0EBOVR4a90UwhU0iSkccqutluoW+wColH+C6sY7wpMJfKADnTlB2M+g2Wex2qn83cPqkvn6mK5JtgkYp/hhM8KD87kDAwEj4KBxu97VXCz8H2wPL8qtVifqEJdqjlsT/fZwfbzx3uPrv718c4ne3+JOW/9+vmT3t7Bo7nztnqsQt5nnx9+urf7twL3vX3n759f/RNNl5HbKKBMlFOUTybc3VRllInTAs6CxCl1YQKP53Ufrq9PL0+EahoYjxNZDEXTfHGXqCuHjlKLO1ShTPIhl+Rmew610PYcdad+evj0CafqmjiDNKF0tw/wB5R2mY+nV9voQISnUbjgngcyQNJV4WCfbn3e2+Fd1DcLJmuJvEkDBzJmJBZic6GsE2HoAAgwrHx80kjdh0e/2Ty+RZmb4D+rE7HkiurjqvA41xw3nDlawq/E9zbIMKl5VSJM0ikSuAnO9/E7h2OMn7D0RC9zbYiAgiTRwy/CiGznh3/8FxMgeN2WiqdXVxi3MNGm/O7mYeJ75yjhxQltVFU5b/SYVWmV0RXG99qibH2zluEQSmiRszFWKVbgphlmJn8w9NI+kmydpkGC0LBSxawbM892cH9wN7j/gLg/7qOoKZidMpJ4aCqxtPwQjm5sg+vGSccbeL4L+5Pjr8r55yKkjUHnruWa8krLgVFxz9384ff/9vA2jrC5PKuWkpyPuldLuUf026IUTYXO3y2k0JdWhboCRtTTLVS/vPl/dSN3UUZNAu2StmjqXosFYKo4k5wbTFw9n6O+mxyVsGZcnzH5MWZFezu1O7jlNJadhycFtOsM7nvriDT/+C8Pb59sLpewW5J+f9P90RalrzmTdpjQRbRpltkvgwAtxldVM1DFSWmGqsi5AmhOw5jW2alCPmK43h0qBt0hWPQePdna37PEEeQa3E1J4AusRond+LNMBbbur2AeiCoGK1M4A/VsDSpOPyg36sehx0G5PuaxGr9Mus4w8YOIfEyx+jbqf4JhmCWojnIOPvsElVkfte98WN7mB5il72T80pPZNARIOJQHIfN/rnU6SsUNrNR58JrcdFNvdHYevnYw1wTcrKHfRTduTDCLYVfEgg4pwj6xmgNnhNkWAo40IiWf0NT5oX/1VR9D8p1GCvcNeTJCnxxskTkHV9/e3kVFG3qvYgwb8IdAIpOU3E4iR9pRWD/poU5lHKJVudmSU4TRsUPMQYLFIjDKEFYUn2CyhSzsoy8us2ccDC5cnWVquzffYPkLP8Dbm6uFv/lje2lpa/fRp3vAP9y711HGci/G9Lo4q/MwePWz5GLp0a+gyZ07HdOkHrxE5SNOOlr6ZOsZNli3npMQlDikB+2T/3UuIpjQvEm+8UpvvfR0a/+TbfRbX1tXPSBFh5lnY2CxQrQWAPcyoI9LS493nmrfdHK3QRfXxupHnZazfq+pHVErfhXZF/DJ2n14clc94ZRW1c8AqfpBRqOs3bWe+FTWJKGh1u7Ao4/Vo3Esfl9dbzmr6030Zn209wTEuorJuzc6nft3OuTuf2NAf9zCSmY1MZbl3ljzPrrv36tqZq1xVkNjwe6Nde+O531U1cxcvXsj8L1VMb073pqHiYmKoHBv3Pfuf3yfOxMfCTA9ul96u3uP9yrA89k2Ptw7eNSS0DB/0YuXqoKK5apH5uI+29l7sn24JX43VrP19Gdb+0ZPcgF6/jDn/e2DrSeH2yiR3PA7d07unLhLvYPtg09wCWudohdKPwpjT+rjI3I8i4CsI7FPlnqH+1u/3qOuxAjKISZiQo+y1QCWeuZlS739radb0Hj1o7v2KFmgvWfOQNB6jVlDIvk6XitkLoCDeBBCL9uHe/u7NNu7HdMJh5ORUBwCOlcQxeOxHap4VaZdS72nW7/qHW7/6hC7u7NmJD0pR7UqcXh3zIbEJBM3G3bpywtw4nhA/ByV64I9f5joY24MIvjCgmOUF3ohXA4peFh5wr3ALDbCW5DJEKX08ou3LTtVqkHoN/JOJ69Adm4VSTpAoqNEOxhZgEl42MHd6FEwwyrmgleHMcUJOzSJaMssANGUHeMzdkgcxzo5B7oKh0EeiChjI5wXJtCmuHEZqQ1TeE15WHOyv1KcND5KA1GkUpqjpXVJLMiosN4C4k3PY04BmwEKYKQB5shxrr5DdhJnRL6ae7bNiFLU+GTnOOoca/fUeBycJw2uzCGURew72jWLKdOr8J5dcVVlqg59OH5GsxbnZYMfuWORAqgrR3BpCLfLQ4nQACYoqGWgKUkCA615PW7LUdZ1FmXwa+j3YvbK4fZHOBnWcZAaA9d76VLeI8qzF2IQPtdG8XFG/NbUfOMWRmaTCsS0ztOsWtYoUrMS60kjAYTOf76za04XnQGsKfM78HdxulKMnzVleG1qyXL0psplQzOXzsZi0n6P5y12VcXWGtqgShfjsmNJQzoNW56sRUN8bu6LHHzpuk6ZogMBWnmTYETc9vY+WSq1J5bhxquTfTXrZ4X/GMDXpn2EzCK7EJPr1tLbONKpmrbbB4/2d362U1xLsYRmrh3rLK+65gxo6etVneu/FqhUew/b05dsAPsMaRfhH/7xX0wnYbsiab13MK/zr2L1Jbc9lkrUyhWjg0t/86f6xf7UXWhLqfuKVal6VI4id7bzXZF02MoFTEdnvWl5ts18l/g1vWBm31i5b/ppVa1Jvkr/Xm+rjG0ikFCtra4Gw8w09+4A6S5njVT8Jv07b7tr/PD+QpuuYiP0vhXvt+o9qwWs8qErg5d7Rs6APth3lkh2IC4tNBOR/xxNq+X0L0BqBoFtNMYSuFS/LumzcO6leEVlRnY+ekZ8Kt/fOsa87eiINQwLc5L+eORhwBp3TOwbTUQ6G73gBy+42CuligD+ZpRjvORRu91G1ibIQ+DbcuRfgN/7HZotphRd3feGeAiQqZt4wkEI7ULs7Yc+R5pVJ0c8CsAXSU8MoVtwwJ+ImGVKHcS3J13PlGplTKlEGhOOXRsmxJWOSTAgF8Gm4ZAvI/LxM3kMqjwTb77ZTd78UUgdStUBjHhgFCoWmg8u78PVOvvp+LWXctiw4aovmU0MWZt0tNsgXVu0OQZXiRiE7EF+JLDpuBBRlnFCcEFlK28LfJmwzkbzVy3nDJ4+3nl6xDSduNHjwlmgWcLFe4YVT4tRVYxuEuEb2A2jZX9SroXFSFPlLU+4qo8NIxI3P6JvwFI7Rw3u2FnhuTSP59fcMLFQ9qd+Op4WqnmPRS1whcpyxFtqRKtHwmmriwmyxWfQ/pOtZzbFM3fKraZy87dJU8B32CRBeP5Wd6kSxLR954Wtgym8winQOoEnp19I70C/FDbTPbj61rV3czQo9nir1OOtGT3uJoUOkbr2zls8VZOcH8Ete4x7omc4wWaVNL3c5aDFc7W7HIgub5ldDuZ3OVEMBKLU0LtoyHnzYM0K4Ms3cFy++Y/noCG1EncYD1jAxTqEO7eRDSD2ke5iDrZZmHY+A8vqBh8UBr/1toMPrk+I5L7AmItQog3zhRJJyoLu/IJFqMWW2rE/G8Wa9DxgA05Q8J58qKSsf41bqIJw0Tntt7gfmygIdvNYnDxsswBFwIlyp4WtEqm8N6AbAbijldXjo9XjhU5T3wB2CbSSXrOGt6UGa7YWaCpQaMVZvV/dfsFG1u6Xj9wC57bmzpBTxR/nYFXxFdENfph749DOSfQSQgUuSYsNRnbxxrmdUnzgXp5326uDqSsd1d2Oq5zW3barXMuY0W5gWYhCWnLXiC6nMS6a0xZ/mjSFGx8uheJNPDPZ+AAY/dQTSlfeCzQ6giSAbndk3m85p2mSYU71BK11IKJq7z2DCOlQfKo0Oljm5vlkA4sn8ItTd5lLB/BXywULJIcxaSKXB2EUwUtiFjB9AAXIoPATzUf9wC6+8DPPb0pFGpaXjEhlrnWntbeoEm1WOsCBPJE7FxsCfJqtQehN1M8aT+HnKt83OaNX8PwsQJdz+HIGX1LVt3o/nZR+qurykgAzvS2dZqylGYpIrWUzM2XkpHU2FoRITrYn69jJI6MOkX6yVHlQuK8Vs69bVl8GW6hpaimbxfLDURJNTuH6HAH65hlCRKJ5jlo0Ry/fWj3y+Mo+95fa0Hvm9Jql+Wk7YcV+zAa/JD4mWJcq6bF+5T1AV7/YD9N+BDK1AUFs3zdARzi84X5k9SLICqkZMB+y0BqTj6xFskRkHsdX83MsJmqYB4k40JOjrvEzJXG55bg//P4/XXL5+dXTJ6h58HxSLETkz5B6/TxIRYWmBHirJOJKLUamfBGzNSKdx1gUabzh+KxMYIubh+leRiFFdOHbzx7/3CFVzS+BErUxHozboutzEJGTDcZQcdqeVEzuBLg7h7wzyadBeWOQtTBxRuj+3NWeJ4ZpEK2qAeU6RNcTcQWRUolCENG5JQqHMD+yuhEPSeobskuyOVL4GY+09Q+6kiFpGGhJ4URonuRcxiI/LCU0jiiTHnmxkA7LCS4ArmyfhEmn4yxpo3UYk0sc7u9hisI0aGNqhjAKGql79MVFp7MCf92H/07gvz78F8APq4MvLj4awIePB//jT//f//jTn471LSeSq56fVmGN9nbVo5JrtittC+SJ3Xu8/fMDFdex/BA6zjYfDr30S6z4CfcXmSpdh3/5JZ/x1Y784VNx0PGXNBj8asO9L6nDYBl++HzDveM6SYpc+4YrIkNGXn7mQNdPO62O8+R+647zpNO65/zadcR9xqb7qSsPm/vwNg+HLvk4QQ0DTBYZpgQEoYWUHrvSQmv6i5cZSzKntaRJraTSfPSrlsPONDJxVBVvSu9eYIWjQn/ssK2VUOwwLDrQtG6u5uuC7JtHFhlTVFA3m4hmilyaZFHPiOeNcxErsJO5SfaHHwqW2JZaLhTvfFHgzNWDSVOEMW8rxyfcGsr9BNQGPaQoyrJ/llDsAjlQyATOSCtg9wDcGbpbZdhMhkVrvzNJOtlHAM9x2/kM3uJComPDFYF9sDg9OboW5CE6jbH6078A0KNWodPuyM0G+MFMGxcZF8a8EOkUoAW/Mql9ZSJemRRe4RnwW+QlBtcWfr4g30wfN5P7sQakltIGTaw8vTThlyb2S3LINeem+FkEyJPZkfJ/4GFfCA/k7XtBk2uRaOxPmvX4oZGQh9NZXuhmRdNH7dXKpz6G4+XWuf6jr7JkbSV1KPC2a+21Wa8zAVmhLCHjNGrcIMLWdG+bwQMY7ypWpaVkG/UbF6uw/lVkOy7gfE3Wmiht5hk5UOC/BVlzCOAaIr7AiwDEC2hPB7Mxwa8T/lpIMG0CsILTeojkW3NswwuLUxvi6b/HsozyAd9w65zAi47g4jX0BAeyvqZoMnuqO+yFjMgMDM0w9P0oKLqHyzBGfT2VYUou7vBcBbosRigJ9TaY9vmS6fMnxfANk4garJ4Q1cj/qyfN+eQ/eGRnRyuh8OnmQwpX2TRXJUYhJ5ljNKKippF90dgvzG1O6/ER1k892tin3IqMqGURBSyt3Mgwy2bdGYhjC6ysXOFD1nKEn13Lud/ClMh37roFbdX8rhQ4VV93mrOicK7Ro+G5yBvTclbbd5uFABwD+N13Oj/28Zm8p5NzR50cA99qRKeqQ+X4dFlChydeFiD9XPy8FYUKC1jN4tkrwwz4rdNNmV9FetHBdRDLkPce/N86orQl8sg5P9lwLK17FqCQ0Bf5En/4x/9LKmFifX5Uaj8xnh33YsaWwgqdi2EUw01yluej7u3br169ar+6007S09trnU7nNrQAdjTBOIxweGrCfLAseJENt+N0HN5ygnxzKr+hHRY2SsqzZhO7KxHqoUOoKXSZZV8RBW0HxdKBL0UjkJCgcs04NonhWKuWNe4lCoISTFMJMEVQrAhtkAIsoqW2ArGAnpoxsMJDiaNFoLnBaMO3nnceZolIkKAs+abLJ3JYMssHcHHs9MlZXNDPtC/KpaET5dX3wISpwEtORSuSlL5KvVEb/0IlHSdowJqbGLMJm8Q/eCllrIwojxNv1b315lF3/dgKnJQLW6o4/WvrfOzvoaZnrQNcU1h/9M1jbkK4+sirsC5DRNPHzuK2Q+EfpjNGMCyaf54DQBwmQPWjjgQCIpeYwnvB990yPhiC+2JYWyqOobHXfE9qstZoY9f0Gea4jBVn3dRN3bs3f8mou7ovt/XGwB+sD+4YWtb9vV8YYmo1ft1ZXwipTERar8CxV2La97CosUAzHt4KNUbVRCWgrR4bImePOEzs0NsU+HlJh2fKBGBZEgDZjyrYJXkQ8l4qBGU/NpL6sKYnKvuFlytWSHceq+sXrPZBec7IdE76bUBxDObSPjNY0dlw2pbijyeVNKlKns/leax5nZNTvq7M1nZ2dUqhLhdx5x7nZjVq6bRGKNcCBHDORoYjHJhhb7ngXKu+lfSUnlfnSuqDSpoSnR+tsJvXyORt3As6hXfzx8jhbYxUSN79ocUzbhthCROOs4Pj4AE4wgQjF+9/WHGLj9Ad9hfeuXfQT0NAOI5X7FrhfAilXxw4Io+X8KSnbH78Q8JVvFrQHZ5nzLZ7Hp5j3RFR/sFncgUyWXCSJF8C3PHQY0zh/yzv2nrbOK7wu37FdoPE3IamLcm5lDYV0LaSqFBER5JjtI5BLcWltAm1S3BJWQpLoL8ifSzy1Ie8tUAemjfnn+SX9HznzMzO7IWmnKIoYARwpNXu3Ofcz3f4pWmk6xqyERhZ6S2P6LLJMbwtRExQxhTy1RBm6FTZy3vz2WQ+g357wrUA+yk/aLwM41nneDqPghNYti9jDao1nYTDiAHRBkShToUAjiPvnHHJED4IzA969Twm6kj3BlVbQqcA7sZG/2jvQLDCiZCkA5oD16XtP+ke7MJI3fBH4+iqveltett3kcZMOomNF1ROa34306As03AYzzMAvdx3anDo1GikRXucDO8gy4yIGb7rsQoYyGEXADEsv6lsFAmhH6aDccQMAClN/05ivQwcwakzVYnun4UZtdV49aPKgxBfRdZUNZUfIRZTKvy8+jmAMsrVNFABCa6RacTg81xtDiIbCuA0gepBbcqqRxLwydiOgxCFOUDKTkkEG1tZYwZzGYUTU820hIFwgirKH1kJRUpwdriSQJeEUoZMbVT/UQ9QQuAF/HteosDdPXbs4skWPbm3hSfBRv+wd/x0vyebTcNkfAKvSki5vwKPh7bdgU3wBUBAI/PQIcCuQi7OnfTY1r5wNSV+EU86nfHiTtPZfKxdGtW4axWglrqVanA1WE6kfddgQDcCVwSIJX7hKG6fDvVxRtI9p/Or9H1OaFze92vR0tU5Nx9qzIfjvYPj7rIAkqrBWvxfv//Jd0M1E+/B6mGP/Ooh0ilKMvjFEgO7VBqMaCrVQ/H8FfYc0/ubd66hYF7bu4OV4mLwLGQcyyqjiNGXakFodhZJe2drqVBkvAV3vqzHn7FwOsQPC6vfe8kgm9z3lw5yTzVi0EJdNuiGcsL1V0UkCeujdeFQyshvxVZxPW4rfW2z9WEJnq4E4rYOqITCyHJAnZR0kusKTU9XB2Hy1dGlFVXxDvkLeG7xb1svTKlIvgo7HW9L3BwLiSzSmEZ2mdhKnCQmmLVgSQjP4/4d1J+cPQYwanFbBXChvHurgKwN+PimgJ864pHTggeRMPUwC6dtqRt5HUoFsAuSdwa//KsoLtvAlhs6mTSTwiPmFPOaqFwKnOWvTSuVt+aDEP/5O42ZpXFmEDk8wS8jKQl5BYGGZZLZafCFFfCBC4eRObf5/QKnMGuLsSvIZ9+B1LNI6/uFfhStmqUTFlmGYXYe5VAsCvSGBZO7QL30d4rXszga7LA7EsZS83a/fLpHB2f34HjXdy+KQFKx5Ncp0wdeBpp/DdV4vNf9DMHTVVQDg3uNls4jZQ2vNCSdKPLaPaof3Be7X/QO97oVS1YB9WZfx6au89p0arnmmG/1G2q6V5snW8uIpEfd/b3HXe/z7tFx1+t+3jssjQzOyWmOpWQB7wiP1nxq+NFoNPzIQBd+rDF9HHTEktCrTtPdMl126fAHLgimS5Nvww7KUrd6/56gdKrLTOyDXtpxWA9pJeWTZa8Ohr5SCHdwh5xDZuZZkob+/HR/WZD87xVZkJiy2pv36gRMV0LaKXxcxPTKcSz9nS9h1JGyllzTGnQSOF9eoY2FS/tQDyU+ZfiBYB0GXg2rdZtEkzFtGj+oQSkt3yeFstiGRH4f//Bet3nDz+gHLFMFPKwhp8uFkJHlQt1dNYEFHYB8Lk6skGHMlv3+q2gqlbWRIqbhJLm0LL/NHGaMYCpOWBMAScAqinUvh5fMwZq1tD7yF5MWCVjby5YGCacHTCjx0PsLfpUGq3HRSbL9OpGf6E1rr1Sg7sQq0vzCifJ9cvh092HXe7zrcbb2MdGlHgCt37fxO1mIyUHRXZFFF2zvNAILf5P0cDlhEkPKFhdLIWu7tsGmjbyhRQ9hxqi8i+rqEa0wh6soBX3oYomzFUMVL8/S3I6gVVMlUiVci1ztlALOtMr+ivagecFQIvCUMdUxgUaCcw2zKo3YxtXEq+Z6CVmeiqw+DE2ZScmSdE2akHu0bURVHiytcWCDb0K4qQDq1hWybahu9awarNs0V6jhbVqsB/P0d3aBRZn75Ex1WY3YWaEhjG7Fds1jLBgfqqmG9KyIq/df109dk/6qKTcA6vps7/Fnu8dHHDxggbxaIu873hFdn5fxkCTeTJuxxaKFoxNJr1zzLb7gChU5UYAINrVKPQkKbJnKlIbpIlk8a4nZS1WNyOhxNOXne8nsiH8DROc86mw2EV7F/wuvOnDYqI0NYGvP2PJHhKHjP8E1XAUpoCDT03nWn0+G4SzqcMnnpkIr6zxr7YsVXTmPthkXW81k0CfJhAf4cD6bkRjv9P3r3/7qdbEycYrSkDUNbm45DWbxWW2DRya5GGaBpjfgd/oaY5D2JJxe1012ve7jJK6fz9//6R1GUn5vnQnllSlBpKaNvu0W4kDHGAiJvPsFNara5Om846AS1yuacohafGwCPSj1MB2gYnjUkPGRUhpeRFnH53d9a4dbtMinxGy+bYyJ1NFR7eMazMIZER1pCkAE/BUfSPWwdYHyMfYAoN0G9lbfrGVApKiWw6tCy+/bLdMe3ajlTfWlXtNnra8epleN589an/P/eRV0h00ZelP6eQHiy1v4ItACJ++2ZqWWN0LtS19VcOwg9riSse4OwUy9cUwaNGflW8HnUl6LxPR5wfN2DQuq4i1eiLrcirnuI7J9HkLF0gU+VCMJ1y0PuW619gMwTDYRvVc/HjNxffWzFwLcV3n7jKtFTMIuJ9Slv8QJCU6YIbGEX+I6eiqQnC1migs2lYuSi7A0TfBqqJw5QuHF+hsxTJTxUFDrBe56M4pfz4v1DhV5Mn/VEvIfTxvV5DySzQMVOSYGEE6jUFHtYuPT9GXW2bxXT51rKMwHW+8CRFy5wMWWrdHIufMWqTN99vU0/PAyZCOZ1rXF5lEaWy4IQpS4lNjWS04W0msAkwid0kAN/MO6cTtEs55ObtXPoo4jDvrGr1RLpL//ydtVL92URazkE9sOn5BSQbWjIC1CubVW8Am3RSN1rmpTv7R2q+r61LUpV3yN1gwv6+vT0Cjnwz2/kohoFQsXCmm2i+1wNLep54LAON/CcYPrvi973+AyRaldRBlvRAlnmaaNfgGLppaVrstOxU7JfZbutel2o3AKcyZjD30jP0rgPO0SkVGXlFenaa9oYNGWwlFbsytHKihq2zXWHdu8sT4uc8nAU7ahFEC0X2NPd/wJpaJc9ppxzCYMPbeCoHR/3mShjJ3Q7SUoXKM1mzZiWP0+mxtVuV1w0mCbdlqaeLOhAdR6sVjb/1EZl8c7sYSLhzu4tVHvZOFXc8OUW2nE9WcWTDOjW8fzUqk0G3f4vvyuDDVNLUcQu5+T2g8rAhdHd4K0ZhFfY6+BnJW80FpgSkSsFuFkLzjcnSS1agaQi3u5cKLvX9M6Yk1zIupayiVCEQJV1eFcNreMKAUhTpeUc0p/GRMGyhGxSU8VCVPGB2NtMur8b9tYrkdz191TEgshJ6aWgj6MkcZDMiTHPA/j9CLMWt4jGEF4+KWQNmdHdWl1SQ68YMWO7SbyqVoJYy8wRSfqq629ZZFPT9j3Q/fkU6KqvCfnF4LV/uHbFfNkz59Bz5ABK9ZRlYFiGSiJr6USBvksToYkyYqi0f2qKyFQG5xUN+daJrMoiaDs7MfJ/AqZuZdAQMPlS8LL6AwXuu2EPEYXA3gnW6R0MUnL0Jz8CQoQqpJqUmfpdMBPm0vpcpCOUMdH5jftGiFMScZpehv9470nvf6n+71njGYhkToAmNijn84iyHP+YRSO8Yj/nxdpByAJSV/JWUkGNrXYPf9hmhKpTfylMZdrJD+wbymT3dbOle+KsSizK1KVBVCL5MC+vGJ/VEJHSycFzC1qgmf2YAfD+V0HE/gTfuwePMbPPfzcO8SPB/zzQe/YX0p9evRO/wR1AF/stfZZLUQ3jIaQl1gH7i0gEotlWWVE/ozkRq4ExL87xdkVJgUAFq26NjJ7A4mTr40b0A4ooGlDXjZOCpXJwIuuTreLbt3glnT+spiiifTKd02u7jT2IKclEugbMeD2dBqCRSSGZ1gpLZYwbsr7Ml5f4CbFVgLtVoPtaqzdiiLOeWCCgoRV6LfVFWvVMM1a+o8ZWxr2GxOgYC7Gc6BNIlkqWKPOdI4BXO651Gs3A/SrQFiaft0LopEu1+m7Gkj2f9FzBVjxf2FT9hJSsqwt8f11xlIDMqxSphOHwMPZg5K430nBW4l7yyA6MGhm4r3XlEBqro1rWeXNCK4mUw6peq8MViwEbi2s4srpi7kCdj3qZN3Zl+B0q1d25Nfud0VPdSXp6+rQ11egL41kFyROFQav6raqYPwaa7ebDHmOle3etEB8qfVnKNb+m1ZQwLnWm4juzcylrmS6wwZUS4oBnM1j0jJt6l/UHfZyR5RBVdXUPDSUntM8rBvEnIHZgvG+sl0Z3AIVaMLx0ITiiuhCTG0ciqSjirygJyXHSPk4DgiX1Pzp6Xl8mXonrdFkenZClzBOTmPlF+XAcuVh570cg2NFiZZ83jABZHVyhlGK8pwMx3FqpWpUOhJzBlnLk0uoBs85K/H55guhJQMrlxMXboBCBBAgNBPjlGm87iYAPh/5nz3d6x14T7qHXQ+c7bPe4d7x51+gOIS3yHNWyrjhfhc03Nn3a1UnJzKVFr4I40TkX/WAzi42JpoCer8YzevrgpF8etp8tnPBLwb2kEgkcyQuk4ZWqP6sFqAmIjmX0fSiuNd6GNFxGUdST556aI0ROG7yVdvegnpc+nWCn2Kfq9tcQu4sNlMM9s0/oi9y8w/vmYnL9T1vESOctuUteBBquH8IEKHCDeTIbYGz7UAtYcLR+0ZCHorsj3letorp+TqIFI7o3LZvkimFwhCTksToMpFpetO5EfvgklJup3lSuuBasGQvkTNYC7ua7jcaVBSni8Sltk09pJalUtMEein0TkZERZlewFh9r7Xln7RyiGfWu7i5S5hbIv3V62iMVKvxnu4dNS3Ea/p4xMEdQlAy76REf0+aVWQR6qB6nBNHM7GWqh6jlsH3JXOuJmsO6cCKvv2/UUIDt6fh1y2Vr3DrWe9brb2seeWlszzW3f7bA6ktE7HbGqGP2njMw2AL/7IeAkUZFa8n+JQvPPPL6w7rfxG2CcmnLiDKOhRl5ZBD/sLEEb3hqCFbRplcCqs+OiKdVAOrRi0S+g0GHeOD1455VZdKKr5Bn+Ks8apneoOlSqKXAhaB11cvy+hGazJaY2w7q8f2YHYeJXd2HoDB0P/uxKMVwxOB8gYjfIkP/guDvMMN6ZFdXYwRTXvrwSf4yZDmzdZd/5OdrxOnrVsPrLvvlUh56XW5ITNiiIM5aVvFv+fzy19Sdx//+p5yZOfTtDN6zSmt6rWy0XA+O0+nmWmX+IV3b3N7e/smjaCM9SwamkZqvn1wp37e8gJ8k5w2LO1CbvM9IV8HDBWV30+/fngwe15EM9qGFXMYpMNr/NUJD1XainpW/eEd/WXF/PT4y4fkjnVK8NegzPDYu5tOEOY2Z7n0JUfOgO8lZx1/Phvd/tgPmN2JbNIuABzzw9ZL6iNq0MF1WGXvaJezTVQ+9pqRkQps4ACygBX5wlZbNYbqsMj8eILhcsL20rg4Rv7KBhf0+tKv48wi7I38rpLPlDBFn2HVlu62jYjjZqe//DA9k4JrSMSFUfpqQtPjAB968Mf55JoOCykIv/xAt2icFuW7YptHMaq75eKXAFWGfAOBRlCr0uoQTjUjDPgtc+bshxKbg33QATfspVDp0Y1X/9jcbG29XY6dE5Y9GsGJ0p05zfnbaJqw/2RCWiUj6EAiRmSoxATDURy1YDaEEyePHqf2EHdwSlLzKAZOEzWRAi1/EkliMwLKMk8pEN9x0L/9gTiTGslAIvw33lEp8+dROBwTkw2aHBEcjSMisZyg1sAwL2F5UzoHp8J7CQNpnMYkaIcBR5bD53SFSrSCC0UM2nv14yH92x2PubbNN/bM6ZzA9RQwjmtGSmU8m1+LTxiuLxwjlYrtuplUMUZBmpxFF2GiykTSdDjqXqL2hmywN66zVz+eIth+yJ4sjvh/9bOUCsqVF/FCe0nMAxjT9IlKs6MN2hA1JEWKIskiH5JelMGP9ai33xVNor938OTpcf9wt4s8+8E8Hs9IdWvx9puqgPilP4z6cBXHpAE2Jti8WY4qzmmLeNQuhGLLU/AMxiC3gwV5EFYoUxiTxLXb+5SZghtC4yOsQVfNnMRDWiHG1QxFY/XUYW0CbUDD37oX2jWo+KL0XoandFgKtBQOubGAdc7YmkcaY4tI6DQPI+FaTgGHcHtwX9DKsnNLch/SYmdqdFK6M9MEVxJCO7IMrUk6adwNLG7Cfw6c0st0pE5TjoMnfZaT/8UCEJtbSNt+Ph/EqLwqURvD1Cbu3KbaVHdGaviWgVPUd3SakwJaUH0C6PKcAs7ALDY7TjNrKYyX62ycDmioPE8JldCnrxjwqL4U9dY9ilx1rHgM1VTok1k4J4Ghr0Zaxv/J5yA1c7M4uiD5PPfElTrLrwX1whGd5jBZxcMSBlmi+xyeEbdmrsH1VkFt2Eki97btnbAkRSepcIjaJ8owc4R36f5HYssRY1M2y7SRQ0/SmgoKD8+zVNk1aITRdLah6ushcuXa4yK2TJmE/KUA5ijQNOR1hEMVKGSqiOVGkDxIEPGzs36/kUXjUdNzj4zA641HLX0ROvoFuwFO11AtWB+6Z9FupiRx4Y9Oi1fWkH7fj5zYyMKxKLbF9oaN/wDpQSMY")).decode("utf-8"))
exec(zlib.decompress(base64.b64decode("eNrFW91uG0eWvudT1LSRMaWQ1J9lO9TPQLbpWQ8sS5HsILtBQBXZRarsZle7upsRLQiYq32Axb5ALoMgV75YILkYwHqTeZL9zqlqspukPN4Za2MYElVdferU+f3OqWIQBI9NnKlYh0Yk1iTaiFCJSIrTzuHBiwOx1RYffn6i+jrU9sNvrVrt6/z6J6FS/M9MrTn7VzsbmczYNTmWXf7USiZnIrMS1Izom9H1r7GQIjOhTEE/FakayRi/6z1praRVvz1uiLe57qu0UTM2VLG0DZHoNJNpQ2TSvlb4tCImwi+VpCoPTTfXKa+lIjHSsW7qOLv+ySZWZQpUa25W//p9qIemJVZXOykejPB3DsZiI6xKNEZiGTIXKjWt1VVx8Da//kWM9ViJ1GBiZNzDTNYc520MpZ5fPIpEX0YqzuRI46dpkAwz2YuY5ljZUIYCnCkr+5keS3pe65/nEfZEM0yCJ6GxoDWhnaixjHIa4FUvsJtUm1il0MDj6/cjA/GBXN/mxFBz7l+tA25MnGY270NQQr1W/RzrKNq5tP1zPTZCxcRiX0Vu26SE+G2uGvSARZmOTE2liew7m4jNqOfYW10NMQ47SCEn2jmrAw8w62ymlLOWeAK65xBjCElZNYTa27Vak8ibHrRpxBnM5YyNYnz9iyRtqAtNXPZzGcMkKxxjy34nbfrcy9M+9hTXhMi0GiXOciM5IkVi++Ksm245W6yvnDVEnMc8vzQ7VANYTF9fv493wBcptNimcwaYagRVxPgTU18bR/NMrOHDKX1IpGW2YabSMoGpQMBXrP1QRS5k5YpcAi9nJG0vZOJZlYwTun6iEhXDE/oaohM93Yu0yVSfVXn9YxxKSxZ9qmMxklkSmSzSvdVViAek8hhLP1wXh4+IgZODQ5GAqTfKxuCQjSxVdqydjcFgcpuyjOEPxBcklVnFumFDgfKgCvo4Vu8aQqbQK/ZRwwT2jrFOcxmJMcke9nn6zZ+FEf/28vA5eO1bDWVLMBljV0EQ1DQ0YDORTRKV1moDa0bi2fEkOzdxK9RpEsmJ8FOYBOJElz5sNUTxmIb8561arXZH7H3Of6D3VPbPyZSgurNbIH82F75i59O8aWfghduGpAnDSoqh/Z4xb2byF2megtoAxs0BgsZCGGyaqwjh0ntzYWLDyPRk1BIdGsryUEuEJHb5vhz1NIctIsduJFJyFKucQsFvS0ZDU2+1Witn4u9//W9KA1A8MREpCiewt5GKU/naR0TnQyDnEoJA5JVi4MVqlQ9REwqr7CFn7AKe5LmcyCr3rRr53MHxs+7xaefVkyOxJ+pwMiECLynbdTINGiJA5gnzvrZd2U3YrmgQ+x1a7IwnUHDuIjJ1+ybKR0hFQcNRi2TPWAmOdUHpHUVi+uyc1v3lJg9ziL07iMwPQ7yQnY9oGgI2adFWx90LOQJFlzQKIaQ02VJSyS1GdZzkGb8/fezeOVFpDm26nR3LlH93rGU2VmD9CGSCIpMXbt1LbKXN7yMV5jZ2ztY6hYlF6oUc8RzlJEj/Vlcv43Yh6+/i78UAkSFGlhFzYtcDN+ynXjEF4iJJoZIyG87c0vrKyq146DPyDE2Z1QWfz74EbfybZ0fPOy8PsLPgzj25JeWDwNnhf7x6zoOb8sHD8L4b/KZz8qTDo+t9ubXed6MnR3854sFwfau31fPvHz46OOFR+OvG+rob/fPJs1Me3N7c3rin3OCjo4JqOFBb6qEbffqq8+IlD6cTZI1RM9eNpkyg3KYbaNw9VUOjxKtndxsnpgcfbKQyTpsI+3rgaBwevWDOct0cGeQLsojG6dNDfG6eqGEeAYIdqjgyDcBE4CAgsem8oGR3CaCNtPXzbBR5m5vG5roL3O5Z2VZ9ZnbTkRM6c6igkql9Xm9QmKwmbmRJIvD4o4ChIWYwhRKUtXhVAHMViyWImorxB5zRMsUpKlkeQh8R/LAIVpKRohprBLc8FuRZ7JocxUL1WloQYYqcZWWIDKwjvJjChzx+nPAf3x63CmGU/XbqRa2hyuoB+A3KckR6lQhK9UxdZHMeX0d89eMtpPWI/D34IwWPP8pRshOURnd5NMowOI0I9G82Y59nDLPya3eDuzT4Njc0fDte/riMqznvZA6gwT5jSnGuYHGoXGx+dg5IzlDTu65K0sJcoQMf6rwV8ygCIz3QqXiBVNyeytFpo6wYTGsRzVn0DUiTCSWE7cZssMMgXTkkjxwhOi4zNI+tgd2Y5qmMtAPwMD1Kz8g1I5dAUY/0c3o7EsGMYmYyoDSez0gtsXCOD397W5R1lLUV2Xkew6CRAv8UzPj5LngVS+GzE5nDK0LGzIn7C2pghoq/KLdZ3WdHDb6fEVov7zEl3kMkXGYbHMz2ESvQJscinnraOrydGPiQQGhQ7wgAdyob9PwwOAHiRN7USaT70qF8thQnM7dblkfL77FwK1Y35/3bVTgvsahyUk1ic9XzBSnDZ0vozCsKMNpjuz5BqYF2ZWZRclZV1iFQHrlKkqePoVZMo+DkyjhexUMoXjy6/iVWskRFEBVboBAsGLMRlVRVmfwcSFBRDHR4FIwC2GGzCM7zRH2Ip1lYnSo3YlG74vWjdkP2xTsH1vIcF7sAhYk3cIZ2PKXYNuFQEtX1r8ggsmw6ADyjol4XIZbGR+SQLFe2P+dgFK4LkaFSW2Y/MtXDmFJGfLtGNFtnzpJOtTi7OCPlK3GPKj7Hb4pRsdsUF+JLsXE2talEpvOGw47JemxjNlW5nCU5RVIpSKUtlcgwibSi1wumJiSvbcX23EMwDHXjGUI6rOle5fEprNQqyeqbWmqcowovG8JGaZ/PYQeRAuJ0+ZuY/vAzlpDRh99ImyiN1ej6xwzaRhGdWD2ivJ/OImRUNL1KtoAIwyS5KC88SCIIw/Dr90h0UOU2daMU18yFa+APrEyGDO+8KBvG58+MzxdaTFQe9/LXcIBbWA6acdRpa2P1bhqcEJgJalETwr4JzQ8x+ZuP5fBACO4cYIGjrZU9+dq0Qc2lrDG1996x2MaKo8AZamEoSk+rT2h0APhjqPFn5p65XUOdDVDkFRF4LqAmgDBUtqRTINPcIcWirHQVK/V1wLYrKp8+e35w2v13qPS7+ktLXbDZz5WG8GNPwUbxa6VRgUn4V/ePS3Pp1fJLxfD3pUWPPm1RN/jxNYsXly85Q4zc9KuzSfvI5LS0J2Y1DKKSs3mFt0VRxThgTICSqoZvCukHc9MDWtNUQGx9cHc3C0WaTSK1FyQyDHU8bD9ILsTGZnKxQzSbCAjDuN0nzdkdZql9ZzAY7NytbHtwtyf7b4bW5HHYvuRpVzsD5JfmACAxmrQvi8rGD6f6nWpvbCUXwf48pUvey9XuWhbu3y2DavareqazHLXPtDvb4ISTFkKTPcATSzGdd2f3d7PzG3dYYvrO4N7g/uDh3L783qjzbZuo1jIzauNFaj/rUFxOC8GrYP8ADJ/vL3v9/4uDR787BwfislDLlXDs4IclNbJykLQTttKg9drouJQzSVMBInjxpOQSK9z0GFN3gzS9QrOYasAPGG74Z+nKvHmHelzs3Be/bR1HOlbNXmT6b3YQHYc6bq+LjYdeIGJ9wbYR0ShRRd4ZMpMsWm15JbLw9v11EL0HepezxsCV96DLUhPjamE5x1Mh6vvkI5fO6uEUWGZx7YwTsl/daworRTJJVbv4sOMetDeW627BDQtXurp0eiOPpHX2PROFa7Jbpt1QdV3kL7UPnkuHRtyUSqxPQd1SKPTJfJqTHP4IXf5GxmgVlXeppzFltSJ2r8uNdVJieUNflsJHMHFt0gs9VHzmxThJBBP8KGWdleUvu1oXsu7JlFsLXFWZyrtH5XcrDPYjJW0bej0P9uc1CS0mFfv5qO1QS+oKxnvR/EGH2Xn7/j1seqdMzRsRbNVZ0KG2cipnbtnQ5hMZyZ4lCEZq2O3tT3bXevsun/MhQplkSqDVpX1IgeED/oIYOHd7VROEy0dwtJgkNSnomn9Ed4oaqBSQvlKcwwkADyhNfEPdo4XW7lpStsjPD7E6Nxz73UpTw50x2vpqaaFSOw5Pr3+VZS5EzxiYFZ3aTlgJo5zqeweS03yGgn1PruPO8rhm93BwduogVle5xlWrqyU8KSFcFJe27ZElKdv1zhwbUE+DkGIByRvu2BHlhqHhPp+NcsvE4VLfv6seAleabBzLCYN97/4kDFNsmbvcs/3ParPMTtqVGObAzx7LtD59pSEug263l+soQ0Xb7QZtcXl1hdGrapPt/wS/ZvVhUXN8KhRzcuyrJINq6BdtUXKp3l7kp1FZYdqyBgYK3FlVklPD1RkRNkbHC3WQWml1uzHwd7d7NVuWpdySCR1n1hdSCiGnBWx4n5IkAtNN2G4+iZSxXmub0EYFRnxqKiq3VKeKXPEQcYHzm7n2EbSMT39QenjO2Xo58+U9FoF4YVcLYPhTdnk5VaXbSBkplXLdJ6d3n/4eLoEwJaZu5KZAXg5HTZP9LQXVx8uuWdzKWdHRcefk4MnRSYfOcL5jsdSDExVxdwaZJy2dIcacwthT+QCADm2nXiyMcH7b8GQcqT16f49/fvgbnfFzl0Olf6IR6gliT2IrKNWJ9WCX+/Z/KL0EcIpwlJnSa39YeI3f2ncvjeTEHWfMXthfWIYe7fr5Kp6fv7tAnxna3yuvYNx+Sqss7oZf2N0rL7Tw2m7pte/973rw/Pr9UPdNoYMe0kxMIS6hJOaOEabyX1smfYJtgYxD+jUPQVRcYVpgGiIrCdpUN8CNcneYLCM6Ljcpt7qmhOhJj654MAhIgYehqiqN2PjFmBMl3A2NMZ/4+ydCJtb0zKIUDqh/ef1TNpWE64FR01nE179Sa2xu21/SNP4BjMz0twA2N6ssNWm8OT3OdrOa87NWaZx/zPryburq/NQ1Gl9zh/ZjDd3U+Zijr0dk8Cv04IFYE5tC/P0//0sIxMbq+3iLKZRJTFig0mYlSp5QiVKVzsiwvr/gY3nugeswL/qFRJcxiaPyxZTIRlnwRevFX/fqzoJQpXyZ9fKsGijLl35a4riMXLADd0mGe658Gyc0DSH5psX1+1Rnpl2cNJaQVgGl/B0NslsdajOS7ohEp6Pp8dmE7/qYeJDzvsTo+sfUNzv4kh5dGhoR6ONWG91lEGd7ZyBJvfKxxB97Z9MqCmUvgsAcuOIkp2WDYWGXb5w5COYvGsxi6AyT3FzMzxDEJ+Tm7aIP8Undon8y05bxg4Odnw4e/kUGS0Ub3U64+gwIyN2d+Sd3wFxW68p/XaawmFsR6OZSdLVQB38GUPlaIc4XVuGQWHU+vMRZDpyE5c83BhvCvznXgSp52k0Ae6FxMeX//rKu6A39JRLQP+wwLdm89/ibGkrziy0azYISPt7F2uYuVhFebq+PtcCG649srK9/UW6X3F9fKuPl7eLtj7RKSz3ySA2ypTWEU9XmnASD/UoaWdK7Xdq3vTVm2Kh/by74Jjt1ln5vRjrOrUst7IX26A1N0bnyrciOPhrQfZzPX00d083GWzhfJISUONolWPTSXfsnVNTTKh7TNxXo6r3057jT2xzUXwIKCnN30Oi/y+BuOxdXxOhegMc53NvjyyD5RDaKBpbrfNNtrDzlGylpS7xwd8KKa/H+Zjf1SEZ0I4BXMoy8/H3YPAa2Zn40308c0C2D6mmn4h6jP6eYXQhz65ShW7VVtawrvRhAl7YRSiZLJxLSNofAoHSlqr6xtR2qYaPcJb2zqTa3HjxozKDESvk8zkdLIpCn3APfKXyEjzY2NyuZ1Idnbo3AAebC4SL/M4ADuJxhIb55B+Kt9a/UyLlahgI6pfsa7Rwpz/YlYnWJpqE3skm79XAz2D91Kt8SH/5HvIr5vujWfFK4gYvN+wV28L2bB+vrRfODz28EJ1D/BZlPJOoy6Q6fDJ07shut7Z0pz1/dQ0xY1nVtVL8mAqMNddFZaO32bGVpCC0u1p6JI9h/uC6SnKp/EsdX6+LbY/qgY7pFoqX48PPXOcyCC6RQffhtd40I+ZBT6X0fRPylBKooSkacig2wtsl31ugbPMb1OYqvVUx7D64rDKN3XzS5I5pNvtEJ26ZLFdNI0Gy6nY7ZYbkX70uY4sp7zDdFi+8VgZL/ZlFLfJ37Wzv5iG+GUBXEDkadILpSX3a+CR21IJhkqnKLE/TcN2WOT54dojLx35hxZwaJAlSbxSVBgQABmprSE+ma0nyNRb1r1fzuUMT4vdX+F39cRbE=")).decode("utf-8"))

ava = Motor(titulo='Decidir', meta_xp=90, insignia='Quien decide')
quiz, ordenar, comprobar, pista = ava.quiz, ava.ordenar, ava.comprobar, ava.pista

# Pistas de los ejercicios. Van comprimidas para no destriparlas:
# se piden una a una con pista("E1"), o con el boton Pista.
import json as _json
for _c, _p in _json.loads(zlib.decompress(base64.b64decode("eNp9Vl1v3DYQ/CuEgcBJcVC+7NY4IClcw0EfUsSI8xIURcST1mcaEqmQ4sGXor8oPyFv/mWdWUp3dpHUD7Ylkcud2dlZ/n1w/vxgaf48eC8pS7cRk701TegHG23jgjfWdNZs5IvZGuvDaE3KZmO7EI3t8KkNlbkMXTBNtr4NeJPMKH7NP6G1aWFush+tvv+cxXSdbKw3h1h8aLI5DPGwOliYg3Nv6qbGIXVbm2u7NW1IZoiy5u6ErByCVea9NFlia5dTBLl1a9HgWL/QcCYJAPirEHsi8US0O6KtF2YlDdImCDFvbIflW2JUcJLMy+pooan6wGzXFhQcVceA6Yx3fk2C2t2ZBkSN2MVDcz90ssB5JQ/sH5144eo2ePy2axujjUmqg7+Y0Aul/kNgHEkjeBHPVJrr3MkIkHIDvEirnt58CoNEJB8lPX5Sc7WWS7oW+WcwCyirnBotWhO63JdkL7bjdfDKwlkGk4BPZrpQMBAtAACKNOKXpn5am5YbG9eDpcSyPMU7XS8IH0dbmY9GOoNMXJsDwWIdTm3dxiUIZ6HPQWNfvvvt/amefvd15g3nt44Ku0+ev/vGbWBwBooQa9eEBQ9o5b4yl9yGZKJbSUEjY6TgfAGXXL+CMNPE9Uvl+iwMzjJvSUgMO1Aj6Bji7QAy5TS6MW9FkQ6QHv7g5SA+uQ1CmXc3Ycnto9xiq/SDky8kfZTYO1/0NkY3kDTk6jrAVeBvC51kZlYFSbDYFchvcggWtQxjgHCzB2abXFEiypIc+6CGCKuqMudIzeOlPlwihv7zxvlLVxexA6LPvnGsNYgW88fpJRk8MdchspXOkMkJGsg249SsG1CjzbOK1gzgniv3WXQPjQGH16/NSa2E169f4d+J6iOl+i27aYy5GXO0kzh6l4CNwSDdIUluQxNat96ZiO1XznptrsF2FomkJUHXeEzsxNpd4cSakGs+CvoXL7YsSl3wU6K2aHnu+6J/XYP9xYRK95Yao/pXKB81MEmW38WPMRRGYqRELZ8lDSQADJ4jo2T9OoK380viSw5udetKyesogO7rnTiw4IolCdrnSAbUB/UVfBmi6yF5Ey0YurYo2Lx/wZzoHit7g3TIE5U/WcPku4BctvqwaycXpoIcF+2X/NGxQ6bbAH1IbkWpJrf27so1djIUC+xqHCNW37PhXq14Yn2ijIKgfEJJs58Efz71BnWkypnaXXWotu1HUfD78IUDqEm9GJJ6WT2r1VCPKkgV5Sjx7LVDpShvu7PoJNGcftfhViFCEUtmMdqVWhTYvfsKUzeBkO6+7UXOl9oT6menLHIoUn81K32W+c/T4Nx5kIVnyu1AS8TRQ/Do8zZ0k5tAkECcXJqbqbefM2WBhcKywlLPpiGK9RuhlxWHzd1IQIkTi/XfjbO4n5X7yaCTAG9uVNvLYt2zbT809Ufq83sbRwa2a3JnNWMwnYTlxdmWuVMRKfc27hJIbEbkXKskMVKznZYSZB1iXZnTvX/MkDvBxw8RL3HG4zKEEaWM4yfaz1aFhgBYbFY2jcVcGYgSTYJnJxxEpRq/aDUusaT5oc6/56gP9T57tKC9LyjFK3djR7WrAiHNLaHdis4JvMx0a6y0XRfKZeatkFbemzBjkDi7JAZk0HOme71O4bBZ+BFi3I1AbaVfKcmN7Krmec26mJbv2kRbq9gVkCIgU7BxMbkLZe8R4t4kILGt3GjJ775aJgXRYuhqE8vkgsW93BWvREYXYegVM+KYqYsQ9Q70Q0NjV6ZyeeycF1rjOjt2ALsJmTPrEtvQ9iiGfXm26mPqgsAEyJuAuykaf89ZN5vbiRb+d/sFVBAvlButCn/mCYndnzWL2Tu74hH1kKrpTfxUFrKtqDjmApwrJE7q1HtlHr1wozbj7lXa8YKXO433/PhRkVLCXdUXIxn0KxI5OX6EjuQcBeifzLPq5LguJgNmLAYuRwBVGznT4Lji9SIceAdBG4KsGOL/+tyLZ/yZXCyZj6fFOecaIb3FXID/WuJkh7xfPKBMJ6deMzhf7927LO/ltKzAlqvL9NWx2c5baJf//AtdP0Jd")).decode("utf-8")).items():
    ava.registrar_pistas(_c, _p)

ava.barra()
esconder_entrada()

```

```python
portada()

```

## Al terminar este cuadernillo vas a poder…

- Escribir una **expresión booleana** y saber si vale Verdadero o Falso antes de
  ejecutarla.
- Combinar condiciones con **y**, **o** y **no**, y predecir el resultado.
- Escribir un **Si**, un **Si-Sino** y una **cadena Si-Sino Si**, en pseudocódigo
  y en Python.
- Anidar condiciones sin perderte, y saber cuándo **no** conviene anidar.
- Usar los operadores aritméticos —incluidos `div` y `mod`, que son nuevos— y
  saber en qué orden los evalúa la máquina.
- Usar `round`, `abs`, `min`, `max` y `pow` sin buscarlas cada vez.

**Lo que NO se te pide todavía:** repetir algo muchas veces. Eso es la semana 4.
Si un ejercicio te pide «para cada uno de los cien estudiantes», te equivocaste
de cuadernillo.

Este cuadernillo tiene **80 puntos** y **90 XP**. La insignia se llama
«Quien decide».


---

## 1. Calentamiento · 8 min

Tres preguntas de la semana pasada. No tienen nota: dan XP y te dicen si
puedes seguir o te conviene volver un cuadernillo atrás.


```python
quiz_eps()

```

```python
quiz_traza()

```

```python
quiz_asignacion()

```

---

## 2. ¿Puedes matricular la materia? · 8 min

En la UIS, para matricular una materia tienen que cumplirse **tres** cosas al
mismo tiempo: haber aprobado el prerrequisito, tener cupo en el grupo y estar a
paz y salvo financiero.

Tres condiciones. Y basta con que falle **una** para que la respuesta sea no.

Antes de leer nada, ejecuta la celda de abajo y mira las cuatro filas. No hace
falta que entiendas la sintaxis todavía: fíjate solo en cuándo sale Verdadero.


```python
evaluar(
    "True and True and True",     # aprobó, hay cupo, está a paz y salvo
    "True and False and True",    # no hay cupo
    "True and True and False",    # debe plata
    "False and True and True",    # le falta el prerrequisito
)

```

Una sola vez salió Verdadero: cuando las tres se cumplían. Eso es la palabra
**y**, y es la mitad de lo que vas a aprender hoy. La otra mitad es **o**, que
se conforma con una.

El resto del cuadernillo es ponerle nombre a lo que acabas de ver y aprender a
escribirlo en los dos idiomas del curso.


---

## 3. Concepto en corto · 30 min

Lo que hay que entender antes de tocar nada. Tres ideas, y ninguna es larga.


### 3A. Una expresión booleana es una pregunta con dos respuestas posibles

`nota >= 3.0` no es un cálculo: es una **pregunta**. Y como toda pregunta de
sí o no, se contesta con uno de dos valores: **Verdadero** o **Falso**. Nada más.

Ese tipo ya lo conoces: es el `bool` de la semana 1, el de `True` y `False`.

Lo que cambia hoy es que ya no lo vas a escribir a mano, sino **calcularlo**
comparando. Ejecuta y compruébalo:


```python
nota = 4.2

evaluar(
    "4.2 >= 3.0",
    "4.2 == 3.0",
    "4.2 != 3.0",
    "4.2 < 3.0",
)

```

> **El error que va a cometer media clase.** En matemáticas `=` compara. En
> programación `=` **guarda** y `==` compara. Si escribes `if nota = 3.0` Python
> te va a parar en seco con un error de sintaxis. Es de los errores buenos: te
> avisa. El malo sería que funcionara.

### 3B. Los operadores, los tres grupos

Ejecuta la chuleta. No te la aprendas: ejecuta la celda cada vez que dudes.


```python
chuleta_operadores()

```

Dos de la tabla son nuevos y merecen un momento, porque resuelven problemas que
hasta ahora no sabías plantear:

- **`div`** (en Python `//`) divide y tira los decimales. `7 // 2` es `3`, no
  `3.5`. Sirve para «¿cuántos buses completos de 40 puestos necesito para 130
  personas?».
- **`mod`** (en Python `%`) da el **residuo**. `7 % 2` es `1`. Sirve para «¿este
  número es par?» — lo es cuando `numero % 2 == 0`— y para repartir en turnos.

### 3C. Y, o, no

Ejecuta la celda: el color hace el resumen mejor que cualquier párrafo.


```python
tablas_de_verdad()

```

Y **no** simplemente le da la vuelta: `no Verdadero` es Falso.

> **Cuidado con la trampa del español.** «Si la nota no es 3 ni 4» **no** se
> escribe `nota != 3 or nota != 4`. Esa condición es Verdadera siempre: si la
> nota es 3, no es 4, así que el `o` se conforma. Lo correcto es `y`:
> `nota != 3 and nota != 4`. Cuando el enunciado dice «ni… ni…», casi siempre va
> **y**.

Compruébalo tú, con `nota = 3`:


```python
evaluar(
    "3 != 3 or 3 != 4",     # la trampa: sale Verdadero y no debería
    "3 != 3 and 3 != 4",    # lo correcto: sale Falso, que es lo que queremos
)

```

---

## 4. Laboratorio · 55 min

De aquí en adelante todo se toca. Cada estructura aparece dos veces seguidas:
primero en pseudocódigo, y justo debajo en Python. Léelas juntas — es la forma
más rápida de que se te queden las dos.


### 4A. El Si simple: hacer algo, o no hacer nada

**Pseudocódigo**

```
Si nota >= 3.0 Entonces
    Escribir "Aprobaste"
FinSi
```

**Python**

```python
if nota >= 3.0:
    print("Aprobaste")
```

Tres diferencias que hay que ver de una vez:

| | Pseudocódigo | Python |
|---|---|---|
| La palabra | `Si … Entonces` | `if …:` — los dos puntos son obligatorios |
| El cierre | `FinSi` | No hay: cierra la **sangría** |
| Lo de dentro | Va indentado por claridad | Va indentado **por obligación** |

En Python la sangría no es decoración: **es** la sintaxis. Lo que está corrido
cuatro espacios pertenece al `if`; lo que vuelve al margen, no. Ejecuta y mira
la diferencia:


```python
nota = 2.5

if nota >= 3.0:
    print("Esta linea es del if: solo sale si aprobaste")
print("Esta linea NO es del if: sale siempre")

```

### 4B. Si-Sino: dos caminos, siempre se toma uno

**Pseudocódigo**

```
Si nota >= 3.0 Entonces
    Escribir "Aprobaste"
Sino
    Escribir "Reprobaste"
FinSi
```

**Python**

```python
if nota >= 3.0:
    print("Aprobaste")
else:
    print("Reprobaste")
```

Ejecútalo de verdad en el motor de pseudocódigo, que además te dibuja el
diagrama de flujo:


```python
r = ps.ejecutar_pseudo("""
Algoritmo Aprobar
    Definir nota Como Real
    Leer nota
    Si nota >= 3.0 Entonces
        Escribir "Aprobaste"
    Sino
        Escribir "Reprobaste"
    FinSi
FinAlgoritmo
""", entradas=["3.8"])
print(r.salida)

```

### 4C. La cadena Si-Sino Si: más de dos caminos

Cuando hay tres o más respuestas posibles, encadenas. **El orden importa**: se
prueba de arriba abajo y se queda con la **primera** que se cumple.

**Pseudocódigo**

```
Si nota >= 4.5 Entonces
    Escribir "Excelente"
Sino Si nota >= 3.0 Entonces
    Escribir "Aprobado"
Sino
    Escribir "Reprobado"
FinSi
```

**Python** — aquí `elif` es la abreviatura de `else if`:

```python
if nota >= 4.5:
    print("Excelente")
elif nota >= 3.0:
    print("Aprobado")
else:
    print("Reprobado")
```

> **Por qué el orden importa.** Si pusieras `nota >= 3.0` primero, un 4.8
> entraría por ahí y nunca llegaría a «Excelente»: 4.8 también es mayor que 3.
> En una cadena, **de lo más exigente a lo menos exigente**.

Compruébalo cambiando el 4.8 por otras notas:


```python
nota = 4.8

if nota >= 4.5:
    print("Excelente")
elif nota >= 3.0:
    print("Aprobado")
else:
    print("Reprobado")

```

### 4D. Anidar: una decisión dentro de otra

A veces una respuesta solo tiene sentido si antes se cumplió otra cosa. Eso es
anidar: un `if` **dentro** de otro `if`.

```python
if aprobo_prerrequisito:
    if hay_cupo:
        print("Puedes matricular")
    else:
        print("Aprobaste el prerrequisito, pero no hay cupo")
else:
    print("Primero tienes que aprobar el prerrequisito")
```

> **Cuándo NO anidar.** Si tus dos condiciones tienen que cumplirse a la vez y
> te da igual cuál falló, no anides: únelas con `and`. Estas dos hacen lo mismo,
> y la segunda se lee mejor:
>
> ```python
> if aprobo:
>     if hay_cupo:
>         print("Puedes matricular")
> ```
> ```python
> if aprobo and hay_cupo:
>     print("Puedes matricular")
> ```
>
> Anida solo cuando necesitas **decir algo distinto** en cada caso, como en el
> ejemplo de arriba.

### 4E. Precedencia: en qué orden se evalúa

`2 + 3 * 4` da 14, no 20: la multiplicación va primero. Lo mismo pasa con las
comparaciones y con `y` / `o`. El orden completo, de primero a último:

1. Paréntesis `( )`
2. `**` (potencia)
3. `*`, `/`, `//`, `%`
4. `+`, `-`
5. Comparaciones: `<`, `<=`, `>`, `>=`, `==`, `!=`
6. `not`
7. `and`
8. `or`

Lo que se te va a olvidar es que **`and` va antes que `or`**. Ejecuta:


```python
evaluar(
    "True or False and False",      # se lee: True or (False and False)
    "(True or False) and False",    # con parentesis cambia todo
)

```

> **Consejo de oficio:** si tienes que pararte a pensar la precedencia, ponle
> paréntesis. No cuestan nada y quien lea tu código —incluido tú dentro de tres
> semanas— no tendrá que pararse a pensar.

### 4F. Cinco funciones que te van a servir siempre

Ninguna hay que instalarla: vienen con Python.


```python
print("round(3.14159, 2) =", round(3.14159, 2))   # redondea a 2 decimales
print("abs(-7)            =", abs(-7))            # valor absoluto: quita el signo
print("min(4, 2, 9)       =", min(4, 2, 9))       # el menor
print("max(4, 2, 9)       =", max(4, 2, 9))       # el mayor
print("pow(2, 10)         =", pow(2, 10))         # 2 elevado a 10

```

Para las de matemáticas —raíz cuadrada, seno, logaritmo— hay que pedirle a
Python un módulo aparte. Se hace una sola vez, al principio del programa:


```python
import math

print("math.sqrt(16) =", math.sqrt(16))    # raiz cuadrada
print("math.pi       =", round(math.pi, 4))

```

---

## 5. Ocho ejercicios · 50 min

Aquí se juega tu nota: **80 puntos** en ocho ejercicios, de menos a más.

Dos cosas antes de empezar. Si ejecutas una celda de ejercicio sin haberla
tocado te va a salir `NotImplementedError`: no es un fallo, es la marca de «aquí
falta tu parte». Y si te atascas, `pista("E4")` te da hasta tres ayudas
escalonadas — pedirlas no resta puntos.


### Ejercicio 1 — Verdadero o falso

*Dificultad ★☆☆☆ · 5 puntos*

Con estos valores:

```python
nota = 3.4
faltas = 2
becado = False
```

Di si cada expresión vale `True` o `False`. **Sin ejecutarla**: piénsala,
escribe tu respuesta y deja que la prueba te corrija.

| | Expresión |
|---|---|
| `a` | `nota >= 3.0` |
| `b` | `faltas > 3` |
| `c` | `nota >= 3.0 and faltas <= 3` |
| `d` | `becado or nota >= 4.5` |
| `e` | `not becado` |

> ¿Atascado? Ejecuta `pista("E1")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_1, points: 0, grade: false, solution: true -->
```python
RESPUESTAS = {
    "a": ...,
    "b": ...,
    "c": ...,
    "d": ...,
    "e": ...,
}
### INICIO SOLUCION
RESPUESTAS = {
    "a": True,
    "b": False,
    "c": True,
    "d": False,
    "e": True,
}
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_1, points: 5, grade: true, solution: false -->
```python
assert isinstance(RESPUESTAS, dict), "RESPUESTAS debe seguir siendo un diccionario"
assert set(RESPUESTAS) == set("abcde"), "Deja las cinco llaves: a, b, c, d, e"
assert all(isinstance(v, bool) for v in RESPUESTAS.values()), \
    "Cada respuesta es True o False, con mayuscula inicial y sin comillas"
print("Formato correcto. Las respuestas se revisan al calificar.")
### INICIO PRUEBAS OCULTAS
assert RESPUESTAS["a"] is True, "3.4 si es mayor o igual que 3.0"
assert RESPUESTAS["b"] is False, "2 no es mayor que 3"
assert RESPUESTAS["c"] is True, "las dos se cumplen, asi que el 'and' se cumple"
assert RESPUESTAS["d"] is False, "becado es False y 3.4 no llega a 4.5: ninguna de las dos"
assert RESPUESTAS["e"] is True, "'not' le da la vuelta a False"
### FIN PRUEBAS OCULTAS

```

### Ejercicio 2 — Del español al operador

*Dificultad ★☆☆☆ · 5 puntos*

Cada frase de la izquierda se escribe con **un** operador de Python. Completa
el diccionario.

| Llave | La frase |
|---|---|
| `igual` | «¿son el mismo valor?» |
| `distinto` | «¿son valores diferentes?» |
| `residuo` | «lo que sobra al dividir» |
| `division_entera` | «divide y descarta los decimales» |
| `ambas` | «las dos condiciones se cumplen» |

Escribe el operador entre comillas, tal cual se teclea: por ejemplo `">="`.

> ¿Atascado? Ejecuta `pista("E2")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_2, points: 0, grade: false, solution: true -->
```python
OPERADORES = {
    "igual": ...,
    "distinto": ...,
    "residuo": ...,
    "division_entera": ...,
    "ambas": ...,
}
### INICIO SOLUCION
OPERADORES = {
    "igual": "==",
    "distinto": "!=",
    "residuo": "%",
    "division_entera": "//",
    "ambas": "and",
}
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_2, points: 5, grade: true, solution: false -->
```python
assert isinstance(OPERADORES, dict), "OPERADORES debe seguir siendo un diccionario"
assert set(OPERADORES) == {"igual", "distinto", "residuo", "division_entera", "ambas"}, \
    "No cambies las cinco llaves"
assert all(isinstance(v, str) for v in OPERADORES.values()), \
    "Cada operador va entre comillas, como texto"
assert OPERADORES["igual"] != "=", \
    "Ojo: un solo = GUARDA un valor. Para comparar hacen falta dos"
print("Ejercicio 2 verificado si los cinco operadores son los correctos.")
### INICIO PRUEBAS OCULTAS
assert OPERADORES["igual"] == "=="
assert OPERADORES["distinto"] == "!="
assert OPERADORES["residuo"] == "%"
assert OPERADORES["division_entera"] == "//"
assert OPERADORES["ambas"] == "and"
### FIN PRUEBAS OCULTAS

```

### Ejercicio 3 — Completa el Si en pseudocódigo

*Dificultad ★★☆☆ · 10 puntos*

El parqueadero de la UIS cobra **$1.200 por hora**, pero si el vehículo estuvo
**más de 8 horas** hay tarifa plana de **$8.000**.

Completa el algoritmo. Guarda el pseudocódigo **completo** en la variable
`ALGORITMO_E3`, como texto entre triples comillas. La prueba lo va a **ejecutar**
de verdad, con el motor del curso.

```
Algoritmo Parqueadero
    Definir horas Como Entero
    Definir total Como Real
    Leer horas
    ...   <- aquí va tu decisión
    Escribir total
FinAlgoritmo
```

Con 5 horas debe salir 6000. Con 10 horas, 8000.

> ¿Atascado? Ejecuta `pista("E3")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_3, points: 0, grade: false, solution: true -->
```python
ALGORITMO_E3 = """
"""
### INICIO SOLUCION
ALGORITMO_E3 = """
Algoritmo Parqueadero
    Definir horas Como Entero
    Definir total Como Real
    Leer horas
    Si horas > 8 Entonces
        total <- 8000
    Sino
        total <- horas * 1200
    FinSi
    Escribir total
FinAlgoritmo
"""
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_3, points: 10, grade: true, solution: false -->
```python
assert isinstance(ALGORITMO_E3, str) and ALGORITMO_E3.strip(), \
    "ALGORITMO_E3 debe ser el pseudocodigo completo, como texto"
_r5 = ps.ejecutar_pseudo(ALGORITMO_E3, entradas=["5"])
assert _r5.ok, "Tu algoritmo no se puede ejecutar. El motor dice: " + _r5.error_corto
assert "6000" in _r5.salida, "Con 5 horas deben salir 6000 (5 x 1200)"
print("Con 5 horas ->", _r5.salida.strip())
### INICIO PRUEBAS OCULTAS
_r10 = ps.ejecutar_pseudo(ALGORITMO_E3, entradas=["10"])
assert _r10.ok, "Con 10 horas tu algoritmo falla: " + _r10.error_corto
assert "8000" in _r10.salida, "Con 10 horas es tarifa plana: 8000"
_r8 = ps.ejecutar_pseudo(ALGORITMO_E3, entradas=["8"])
assert "9600" in _r8.salida, \
    "Con 8 exactas todavia se cobra por hora: 8 x 1200 = 9600. La tarifa plana es a partir de MAS de 8"
### FIN PRUEBAS OCULTAS

```

### Ejercicio 4 — El mismo, en Python

*Dificultad ★★☆☆ · 10 puntos*

Traduce el algoritmo del ejercicio 3 a Python, pero como **función**.

`cobro(horas)` recibe el número de horas y **devuelve** el total a pagar:
$1.200 por hora, o $8.000 planos si estuvo más de 8 horas.

Devuelve un número, no un texto. Y usa `return`, no `print`.

> ¿Atascado? Ejecuta `pista("E4")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_4, points: 0, grade: false, solution: true -->
```python
def cobro(horas):
    ...
### INICIO SOLUCION
def cobro(horas):
    if horas > 8:
        return 8000
    return horas * 1200
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_4, points: 10, grade: true, solution: false -->
```python
assert callable(cobro), "cobro debe ser una funcion definida con def"
assert cobro(5) == 6000, "5 horas son 5 x 1200 = 6000"
assert cobro(10) == 8000, "Mas de 8 horas es tarifa plana: 8000"
print("cobro(5) =", cobro(5), "· cobro(10) =", cobro(10))
### INICIO PRUEBAS OCULTAS
assert cobro(8) == 9600, "Con 8 exactas todavia se cobra por hora: 9600"
assert cobro(9) == 8000, "9 ya es mas de 8"
assert cobro(1) == 1200
assert cobro(0) == 0, "Cero horas, cero pesos"
assert isinstance(cobro(5), (int, float)) and not isinstance(cobro(5), bool), \
    "cobro debe devolver un numero, no un texto ni un booleano"
### FIN PRUEBAS OCULTAS

```

### Ejercicio 5 — La cadena de notas

*Dificultad ★★☆☆ · 10 puntos*

En algunos sistemas la nota numérica se traduce a letra:

| Nota | Letra |
|---|---|
| 4.5 o más | `"A"` |
| entre 4.0 y 4.49 | `"B"` |
| entre 3.0 y 3.99 | `"C"` |
| menos de 3.0 | `"F"` |

Escribe `letra(nota)` que devuelva la letra correspondiente.

**El orden de la cadena importa.** Piensa por dónde empezar: si pruebas primero
la condición menos exigente, todas las notas altas se van a colar por ahí.

> ¿Atascado? Ejecuta `pista("E5")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_5, points: 0, grade: false, solution: true -->
```python
def letra(nota):
    ...
### INICIO SOLUCION
def letra(nota):
    if nota >= 4.5:
        return "A"
    elif nota >= 4.0:
        return "B"
    elif nota >= 3.0:
        return "C"
    else:
        return "F"
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_5, points: 10, grade: true, solution: false -->
```python
assert callable(letra), "letra debe ser una funcion"
assert letra(4.8) == "A", "4.8 es A"
assert letra(4.2) == "B", "4.2 es B"
assert letra(3.5) == "C", "3.5 es C"
assert letra(2.0) == "F", "2.0 es F"
print("Las cuatro notas de ejemplo dan la letra correcta.")
### INICIO PRUEBAS OCULTAS
assert letra(4.5) == "A", "4.5 exacto ya es A: la tabla dice '4.5 o mas'"
assert letra(4.49) == "B", "4.49 todavia no llega a A"
assert letra(4.0) == "B", "4.0 exacto es B"
assert letra(3.99) == "C"
assert letra(3.0) == "C", "3.0 exacto es C, no F"
assert letra(2.99) == "F"
assert letra(0.0) == "F"
assert letra(5.0) == "A"
### FIN PRUEBAS OCULTAS

```

### Ejercicio 6 — Precedencia

*Dificultad ★★★☆ · 10 puntos*

Cuatro expresiones. Predice el valor de cada una **sin ejecutarla** y guárdalo
en `PRECEDENCIA`.

| | Expresión | Tipo de respuesta |
|---|---|---|
| `p` | `2 + 3 * 4` | un número |
| `q` | `(2 + 3) * 4` | un número |
| `r` | `10 // 3 + 10 % 3` | un número |
| `s` | `True or False and False` | `True` o `False` |

Recuerda el orden: paréntesis, potencia, multiplicar y dividir, sumar y restar,
comparar, `not`, `and`, `or`.

> ¿Atascado? Ejecuta `pista("E6")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_6, points: 0, grade: false, solution: true -->
```python
PRECEDENCIA = {
    "p": ...,
    "q": ...,
    "r": ...,
    "s": ...,
}
### INICIO SOLUCION
PRECEDENCIA = {
    "p": 14,
    "q": 20,
    "r": 4,
    "s": True,
}
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_6, points: 10, grade: true, solution: false -->
```python
assert isinstance(PRECEDENCIA, dict) and set(PRECEDENCIA) == set("pqrs"), \
    "PRECEDENCIA debe tener exactamente las llaves p, q, r y s"
assert isinstance(PRECEDENCIA["s"], bool), "s es una expresion logica: True o False"
assert not isinstance(PRECEDENCIA["p"], bool), "p es un numero, no un booleano"
print("Formato correcto. Los valores se revisan al calificar.")
### INICIO PRUEBAS OCULTAS
assert PRECEDENCIA["p"] == 14, "la multiplicacion va antes que la suma: 2 + 12"
assert PRECEDENCIA["q"] == 20, "el parentesis manda: 5 * 4"
assert PRECEDENCIA["r"] == 4, "10 // 3 es 3 y 10 % 3 es 1, asi que 3 + 1"
assert PRECEDENCIA["s"] is True, \
    "'and' va antes que 'or', asi que se lee True or (False and False) = True or False"
### FIN PRUEBAS OCULTAS

```

### Ejercicio 7 — ¿Puede matricular?

*Dificultad ★★★☆ · 15 puntos*

El del gancho, ahora en serio.

`matricula(aprobo, cupo, paz_y_salvo)` recibe tres booleanos y **devuelve un
texto** explicando la situación. No basta con decir sí o no: hay que decir
**qué** falló, porque un sistema que solo dice «no» es un sistema que genera una
fila en la oficina de registro.

| Situación | Devuelve |
|---|---|
| Las tres se cumplen | `"Matricula aprobada"` |
| No aprobó el prerrequisito | `"Falta el prerrequisito"` |
| Aprobó, pero no hay cupo | `"Sin cupo"` |
| Aprobó, hay cupo, pero debe plata | `"Pendiente financiero"` |

**El orden manda.** Si le falta el prerrequisito, eso es lo primero que hay que
decirle, aunque además deba plata. Y si aprobó pero no hay cupo, el estado
financiero da igual.

> ¿Atascado? Ejecuta `pista("E7")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_7, points: 0, grade: false, solution: true -->
```python
def matricula(aprobo, cupo, paz_y_salvo):
    ...
### INICIO SOLUCION
def matricula(aprobo, cupo, paz_y_salvo):
    if not aprobo:
        return "Falta el prerrequisito"
    if not cupo:
        return "Sin cupo"
    if not paz_y_salvo:
        return "Pendiente financiero"
    return "Matricula aprobada"
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_7, points: 15, grade: true, solution: false -->
```python
assert callable(matricula), "matricula debe ser una funcion"
assert matricula(True, True, True) == "Matricula aprobada"
assert matricula(False, True, True) == "Falta el prerrequisito"
assert matricula(True, False, True) == "Sin cupo"
assert matricula(True, True, False) == "Pendiente financiero"
print("Los cuatro casos principales dan el texto correcto.")
### INICIO PRUEBAS OCULTAS
assert matricula(False, False, False) == "Falta el prerrequisito", \
    "Si falla todo, lo primero que hay que decirle es lo del prerrequisito"
assert matricula(False, True, False) == "Falta el prerrequisito"
assert matricula(True, False, False) == "Sin cupo", \
    "Si no hay cupo, el estado financiero da igual"
assert isinstance(matricula(True, True, True), str), "Debe devolver texto"
### FIN PRUEBAS OCULTAS

```

### Ejercicio 8 — El mismo algoritmo, en los dos idiomas

*Dificultad ★★★★ · 15 puntos*

El de cierre, y el más largo. Una tienda aplica descuentos por compra:

| Compra | Descuento |
|---|---|
| $200.000 o más | 15 % |
| entre $100.000 y $199.999 | 10 % |
| menos de $100.000 | ninguno |

Tienes que entregarlo **en los dos idiomas**, porque en un proyecto real primero
se acuerda el algoritmo con quien no programa y después se escribe:

1. `ALGORITMO_E8` — el pseudocódigo completo. Lee la compra, calcula el total a
   pagar y lo escribe. La prueba lo ejecuta con el motor.
2. `total_a_pagar(compra)` — la función de Python. Devuelve el número.

Con 250000 se pagan 212500. Con 150000, 135000. Con 50000, 50000.

> ¿Atascado? Ejecuta `pista("E8")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_8, points: 0, grade: false, solution: true -->
```python
ALGORITMO_E8 = """
"""


def total_a_pagar(compra):
    ...
### INICIO SOLUCION
ALGORITMO_E8 = """
Algoritmo Descuento
    Definir compra Como Real
    Definir total Como Real
    Leer compra
    Si compra >= 200000 Entonces
        total <- compra * 0.85
    Sino
        Si compra >= 100000 Entonces
            total <- compra * 0.90
        Sino
            total <- compra
        FinSi
    FinSi
    Escribir total
FinAlgoritmo
"""


def total_a_pagar(compra):
    if compra >= 200000:
        return compra * 0.85
    elif compra >= 100000:
        return compra * 0.90
    return compra
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_8, points: 15, grade: true, solution: false -->
```python
assert isinstance(ALGORITMO_E8, str) and ALGORITMO_E8.strip(), \
    "ALGORITMO_E8 debe traer el pseudocodigo completo"
_a = ps.ejecutar_pseudo(ALGORITMO_E8, entradas=["250000"])
assert _a.ok, "Tu pseudocodigo no ejecuta. El motor dice: " + _a.error_corto
assert "212500" in _a.salida, "Con 250000 y 15% de descuento se pagan 212500"

assert callable(total_a_pagar), "total_a_pagar debe ser una funcion"
assert abs(total_a_pagar(250000) - 212500) < 0.01, "250000 con 15% son 212500"
assert abs(total_a_pagar(150000) - 135000) < 0.01, "150000 con 10% son 135000"
assert abs(total_a_pagar(50000) - 50000) < 0.01, "Menos de 100000 no lleva descuento"
print("Los dos idiomas coinciden. Ejercicio 8 verificado.")
### INICIO PRUEBAS OCULTAS
_b = ps.ejecutar_pseudo(ALGORITMO_E8, entradas=["150000"])
assert "135000" in _b.salida, "Con 150000 el pseudocodigo debe dar 135000"
_c = ps.ejecutar_pseudo(ALGORITMO_E8, entradas=["50000"])
assert "50000" in _c.salida, "Con 50000 no hay descuento"
assert abs(total_a_pagar(200000) - 170000) < 0.01, \
    "200000 exactos YA llevan el 15%: la tabla dice '200.000 o mas'"
assert abs(total_a_pagar(100000) - 90000) < 0.01, "100000 exactos llevan el 10%"
assert abs(total_a_pagar(99999) - 99999) < 0.01, "99999 no llega al primer escalon"
### FIN PRUEBAS OCULTAS

```

---

## 6. Habla con el asistente · 5 min

Tienes **cinco preguntas** para todo este cuadernillo. Cinco, no cinco por
ejercicio. Gástalas donde de verdad te atasques.


### Presupuesto sugerido

| Ejercicio | Preguntas | Por qué |
|---|---|---|
| 1 y 2 (booleanos, operadores) | **0** | La chuleta y las tablas de verdad lo resuelven; ejecútalas otra vez |
| 3 y 4 (parqueadero) | **0–1** | Si el motor te da un error de pseudocódigo, léelo: dice la línea y qué esperaba |
| 5 (cadena de notas) | **0–1** | Si te falla, casi siempre es el orden de la cadena. Pregunta por eso, no por el código |
| 6 (precedencia) | **0** | Ponle paréntesis y ejecútalo tú: la respuesta la tienes en una celda |
| 7 (matrícula) | **1** | Aquí sí: el truco de descartar en vez de confirmar no es evidente |
| 8 (los dos idiomas) | **1–2** | Guarda estas. Es el más largo y el único con dos entregas |

### Cómo se pregunta bien

Mal: «no me sale el 7».
Bien: «en el ejercicio 7, con `matricula(False, True, False)` mi función
devuelve "Pendiente financiero" y esperaba "Falta el prerrequisito". ¿Qué tiene
que ver el orden de mis `if`?»

La segunda le dice al tutor qué probaste, qué salió y qué esperabas. Con eso te
puede responder de verdad; con la primera solo puede adivinar.


---

## 7. Cierre · 7 min

Tres preguntas que solo te interesan a ti. Nadie las corrige.


- ¿Sabrías explicarle a alguien de tu casa **por qué** `=` y `==` no son lo
  mismo, sin usar la palabra «programación»?
- De los ocho ejercicios, ¿cuál te costó más? ¿Fue por la lógica o por la
  sintaxis? No es lo mismo, y saber cuál de los dos te frena cambia cómo
  estudias la semana que viene.
- ¿Pusiste paréntesis donde dudabas, o los dejaste al azar y confiaste?

### Lo que viene

La semana 4 te da lo único que te falta para escribir cualquier programa:
**repetir**. Con decisiones y repeticiones ya se puede escribir, literalmente,
cualquier algoritmo. Lo demás del semestre es hacerlo bien.

### Glosario de esta semana

| Palabra | Qué significa |
|---|---|
| **Expresión booleana** | Una pregunta que se contesta con Verdadero o Falso |
| **Operador relacional** | El que compara dos valores: `==`, `!=`, `<`, `>`, `<=`, `>=` |
| **Operador lógico** | El que combina respuestas: `and`, `or`, `not` |
| **Precedencia** | El orden en que la máquina evalúa una expresión larga |
| **Anidar** | Meter una decisión dentro de otra |
| **`div` / `//`** | División que descarta los decimales |
| **`mod` / `%`** | El residuo de una división |
