# Buscar y ordenar
### Semana 6 · Unidad 6 · Colecciones lineales, búsqueda y ordenamiento

Hasta ahora cada variable guardaba **un** dato. Esta semana guardas muchos en
una sola, y haces con ellos las dos operaciones que sostienen media informática:
**buscar** y **ordenar**.

**Empieza ejecutando la celda de abajo.**


> ### Antes de empezar, dos cosas
>
> **1. Ejecuta las celdas en orden, de arriba abajo.** Una por una, con
> `Shift+Enter`. Cada celda usa lo que dejaron las de arriba, así que saltarte
> una hace que las siguientes fallen aunque estén bien escritas.
>
> **2. Al terminar, entrega.** Tu trabajo **no le llega a tu profesor** hasta que
> pulses **Guardar y entregar** — el botón está arriba y también al final del
> cuadernillo. Puedes entregar las veces que quieras: siempre cuenta la última.


### Qué mide este cuadernillo

Cada ejercicio calificable está asociado a una microcompetencia del programa. Estas son las de esta sesión:

| Microcompetencia | Qué significa | Ejercicios |
|---|---|---:|
| **mCP17** | Aplica conocimientos matemáticos para la solución de problemas usando programación. | 1 |
| **mCC87** | Identifica variables, conceptos y aspectos relevantes de un problema para desarrollar algoritmos que permitan solucionarlo. | 3 |

No hace falta que hagas nada con esto: está aquí para que sepas qué se está midiendo y por qué estos ejercicios y no otros.


```python
# Motor del cuadernillo: barra de progreso, pistas y verificadores.
# No necesitas leer este codigo. Ejecutalo una vez (Shift+Enter) y sigue.
# Codigo fuente legible en el repositorio del curso:
#   notebook/cuadernillos/**/ava_motor.py
#   notebook/cuadernillos/**/contenido.py
import base64, zlib
exec(zlib.decompress(base64.b64decode("eNrlfd1y20iW5j2fIhsOr0kXBcvyT7loSxuyre5yjWzXWq7y9GoUVBJIUpBBgA2AtGQuI2beYV6g9q6io67qYiK6LybCepN5kvnOyUwgAZKS5aqdiIlVd1kkfk5mnt/vnPyR53kv0yLNRHzxtzAKUhEqEae5CKYyVFkSxfQlVLHY/XHXb7X+1/TiZ6FycS7+Qp9evcaX1sbST2svL5QYM+FJlBRSxFIEciKJ/CS7+CkookAyXaehHohGgcq7eCQdxGosqelWmoUqkRmuRnmBSyoPZJwmMpR4ciCzzFBNR5nKU6FElOTRKIlk7ovnUmSqyFIZR2OFfgTRxa9JK0rGKowkunUu8ihBz/IiGqVoQYLYX6ZKoGMqL6Z4KMFAciVkkamZFFIMZRzLDKy4fXtfiiQt6B+Ry1hRLyTY8ot/+7aw985p5IXCYNCNi18kP5kwlzGUQMWhHmQyGGXECNE+5g/9KOzhvbzoq1OVBVEQpf1Xx52u7l0OGfGnWIHYcTDNi3Tsn+bH4rx1LCdRX51N0qzwJ+fHvjiIhIxHF39Dq9QBVYkmVDlkkqlAdUVdFGDLCOTzSCVh2gK/oyEEBpn4Yo/1IRoodBMis4LV+oMvaZahhwkRfA9q+EVilvF0DD6di4uf49ZkqsCATEGSWTSIslg+FuivIPElpIJBhPskYR5kMIXoFFj+PZ5hvZtFM25jmojxxa/hNCbKJIapiotU35Gauyu086qf1m7Z3zwFaYg2jtVIiuMPafb+jsMmP5qcJ4PjLnVEZsFJNEsFGA076lF3TuQ5eP5rAvlO0gTCxWPRmCUj9mIxUgkzkQQBJmK0rgRu346SIINkJcsGw4SK0pMkwJYZNjQNg4UMJhkUPDNj9sVL0ltQBx1WSTFR2TiC4KH3hcRlbWP0CSoD1Y5y/hyk47RsK0mzsYypgZn6yH2ku8xZSSYpRtFIaskcwIrAiw9ROFJFTh3WGkYspLGpHt4nxSYTIgozleVki0TWMDkoptr6WuMpxpjxuFUyTYIITIKAodQxq9jFr+gHjB1MBbt8sZtjhA0FpveYGlQAIxjEMoF3Iy2DEKx/YXYm6GiS+i3P81qtGwIKIL6HkRLfBfQwKaJQhlC5fApmXFOTboiDdJBFNGq0Obn4NY8K6AbxRxYKN2JLOE8TdrxwOeDTmHwbmEdX1Dg9jXK/tfu/f9jvCvq3//rgmdgW3o0t+fWj8KHXxce7m1vf3A+91o97b57vdQX/Kp/bHD6S9yU/tykfBPfueq03r797DWovn+6+4UeCR/e+2RrwI4MHX29tbnqtP715cdAV9G//2f7uGzz+9DWo8uMPth7cva/48eHD4dfDb/hjOFT31CN04sXr/b23u/rJwf1BcB+s7T87OKArYPOTvDiP1U7LlzO5EchTOR+Ql896dydnpA5RKG62+VInf6xvbcRqWPTuO/flx2lc3YZuRdO893By9rgl4MjDMEpGvbtbeOEu3no8ltkoSnqP8H3z8UAG70dZOk3C3s12AH+egtAwTYqNoRxH8XkvP4eRjTemUXdDTiax2tAXuiB960CNUiV+eHGr+yYdQMu7uUzyjRzSHGoaefRR9e7e9x+g2ThK1MaJikYnRe+u/+DB4wDKnvVu3A3of48XFQv89P3cGeqGfvBmG5ai2VD1+cZwayiHD2pvw1JXvp6lp2nz7XB4f3iv9jaH1pXvyzG8xDKBb9SwRgBRNl3d/SglW2pS+Hr4YBhaCkVUzJlzHzSjHj7YfFy2DymneTBlEWkhboDriHekDA0GipJcjXvl+w2GLT9uuLXEmuUnLV/Mo4xDNjCIMJ1bcUP3aoNW3yDSDlcobIpeDuP0Q+8kCuFw6jQzheCTVFQ3N2/edMmShslsg5wronXR/mYTnrZrzaPrSKBDlrHcOjwtAFMRpUkP/rs4Ef69B/A6Mle2I8Aucr7aPFbr/j2QtYwaZVFeia5IJ67cCmCKUu/wRiwnuerZD4+v06Zj5A/ZyN0mRHHSdb+Flzkc6zyIzN1N4pA6KzaAgUZJj3S7SXm+0p2sVOCalm+Wfcxno/lYnm0w/42IjbzltEgfh1E+ieV5bxCnwXs7Suqa4NuGyiRS15LSFnuohpygIs5wA0JemSu8jYrP1g5Z0TVdcus9IAzEafTqzm3xzCJcIZMQ3ZKngAI5gZiAgjcQm4UbgEQM6JADfHQRD7R/RvkDOsYxn7CBgZkBfwAWQHemhO8BayRhAUKNoTplSMO5x4yNGLg5A46DgNAA0Sux9WOLGYfAvUAeFI23Hoh/eEo9P/jxT0wR8VgSeJMJY1Nf3L7T8gG44t6JzNs7wk+nxWRa9D9kiBlAeyyUNJjGMD6B21GCu3MrywSI0EpOP9S89eSOCZMImeKmmFOPhcda6vV0LEbQJf3CV40OPNY+fHWjtsdOEBc1NtBkyNPhkoYBHnszokJ4gN7QPoPe0aHcNKQV2TRH8MJQc9ysbYfvCo/UyvSn21oQwHoJpGxUIohUVqR3VJYBl7EkwWsGPvyFkQ+rywzigWQ0Gi3SMKUMKAexPCIdl7nG/wHGgZTNoOSI8hjFmQ9UYMhZBCcjiZxBDQC8/Vb/xbPXr173X/8DkEmbx3LrCaxRsCVue3cfeEIbov48i9SHp+nZtrcpNsXdh/i/J1hI28QCSmpjYzkbW5WdZNreJ2fezi3TRhBlAfKKAKQeeSI451/Ztve1v+VBCeN427uZe3fK5ycSjjnc9l7e9x+KR/69eMvfEvQfvm/c97+xL5HqUJey9D36dGM4HNpvG3ZI9LAha+5QCEESt+2xDytfoMunaZTY63d2oJGz0c4tKCOLuNWx/Hu5u//fhYEP/PsC/8UPwDz89/LuJjhKFzbMld+TlS7LyNTAsVaryM57QtygFFBnSVB2YEwkrnhXIRRLJ5d6zBUKZM1kBiaHYgvhQoTPjeuc0k3AYBDv+Na3u3/uv3vx/E97bwl7v82QXamzQE0K8YJf2iPT495MMjkay542PshCbKC5iSInHXIixmkSE30HUq/AnRUt/FHGuWq1TI9kXtiPcJEncTSwX6PUftIcC06UvQC7bg2zdCxefH9enKSJb9yiHea3b1/C05mLYGeohqJ/UozjNgWutNPjbpn7bXra3CDeg+s68mRcuaDohOEmgnIT7ULw5Tv45IMgi8AlVa92gBHfHZB7Qawg7wNVDW3Ng8IHoITkbPnclidSjks2foA69CSWYziwNj0JeQvKobkIdEPXdGy/4MKo2sallOevX3a6YkwIL4NwzXOIZviX+k4lKp0d55EaTzLVRV5/QxTkEC9+TgQ8J+IlrEnywCcXP8GuEML20RGCd6Es8InS9SiDs+2J49Pp5Byx38/TaRaovoamx6DJfRXf6dv7coCWMMykAGXzTp+HGQ36yQCMV0CXaZLrNF+XeUjVU5Aixcdt9OOYY2jnmApbeWqibOnoK4/Njedof2tz6x7cN4m3v3cAJ/R8700zu7wqNrP1bYwpHl0SoIkKPz5jjEqg5MuJDmNVgl/nScZtBpGJdcmng+JWguse+dSHk1X0hZy7eBSvMMgLFVRXMvTn0RogbnEd+hLio3KybUT4HGQmCAsEDB2Q8iSMZgLwI8+3vUbr3s5LU2+s12hkUESzlHzYEynSJID+vt/2ipMo9wNIHXrSvsXsv9XxmfI+or1fpKNRrNq36iK51XmMHHSaJeQYkTp47AXEDnkyapYfenJHoqN30NMSVNUASg1NLYxvMVA167PthbJtPAwIGF/SKNCVBmysFH56DJ9NgBTpIXkFlcQyULrcPKMqnd9ikvtSTHNZESCMhBwwAaEeY9KAHABX5EV88QsSP/5IyiMTxfVlfDVBgCmuwK2COhaRW4F/+Qm3Mjh5LvTqKhWXnzIu8cOL5BF8C5MyVbmkKwBXJ7JWdsTXguKR4HxhqbLvW4bxb+2s66ZLvplFLFhVHA4XuvhHrgIqalwlMbFGn5//HvyEM1OZHrtTh6OvyN6pUNyjIK6VAz+v0vEgU83+mpkAjGd5foEr+naKQQ+MHGj/bALSsIqS9D9+z8EElEkIVAaFF4xCrtJOEfApcjALq+kQ+g61p8KoieymnXW9dvtimyH+jSR5fK5aZzoDoj5WYtAhkgJnP0qiot9v5yoedg2PunZE20hFu2UD296B7psFLFAmzxgD/RAJ33B525Cq37SM2rYN1G+XI9ku26w/wK9u1q/1IRUu++e4l6ui3eF7N8gXzWiqCqwIAC1pvqinJwnGUF5Ih4ILR+AGQTPPtC3mC1H9GIJiY8dORIEy4jj7hCaJAVWblwg5JGJOSEmv+Jml17OMXzw8qvXAwju8pxVzBv+ZG23N1DCjubGMMiP7jrE1QISOFjoVuGnw55VGb3zRT6VCxNDMKBAPsCvOJo5eQOEO8Ai1WsNIhAX4eV+8URRPGEPown1GWb5vPQabwtBwD3C4KXgQgmo8gW5UrQquF3BEWKsxvgzDNlPtLCnaV9uiqZ/9ksVGQu3qNdMUXqkYMyRITQ5bP84scvgyCQqyhChps5lxxkAfxe2yE3dqdtPpEBdqlqQQ65zxxRBpWmZi9meIlIwctkm0NBKYcwq3aJakEKtp6kPGgUw+QrN7Yl6zzQUCPmiZvMoRju3yzvZyD+u98ea1BzbsqwvSEOt4S1/w6a+NDnz6e6UUS/xfGngDlTh1Wm9n9U1TcLWZXUXLsE8X6eaQ3uImaDCa0P/euqJxGjHeGOzM7Yif3BnsQMg1hjAb/kcyyCePP/2b/i3mLNhFs5lOpWwrNQz283QpgJkwAfA1Rfry0RgkJwj4xHjJlsXypgUmaeEmfHVr075GW0pT9TudFYZZXvoAlX3nc6p25euuh/QJd8NoPlS3bd73weHNst02+DSEA/lQORb9UH1wH/yZjME5o93LPXTc69M4BZNzPWVKfi0Kr+tlq77r2iDlVIQvV0iY62kVZtRyNXVRp+Y6dSDqKWuEDl5ZrmpS1lJc0l3dDavvtxzu0ixJAz4EU5VNUvo+Sbc9r97fN4qQVsZQGbj0nOfKCwBkGVIJSNF8b5c0JO2uWLxR6yv1j0TS9ux0jfDEV9xsx8fgoonjodE5PLpslbju7cx1162Fka4bJMMedok/dTJz7shCk1nM9fAXy6zS1WXDrHw2ojUtaolBL830NzCunQ0mrEEwHcnfR8LDbapJ2/UBYdphpAkRg1yNQVpXGyrUuXQsZj4CY8G/i1WOTYDFK15C23gJ/7pMxFfLwY7DCJ46MXyAwOVA0bCAZIYRCDrswC3O6n0qRLaH3pPiZGceoAX89thuA7Jbh0Y1Oi0G5/XaKEAq2yFtcYmHhnjoEqc+dehB3MjMDbpm7+UrwlHJWRqpqqkbDR12BFJzdHrBRCuV4edLpWF38r2GkRu/5afkfKZGEZu9Qad13Kav1dXxT1OZcaqSr1iBZXKxMmKINk/qEE6Dj4l4jqazQiVdgHzIjR9BUISK26YPjvumWfZEP1zvrq6/uZDKIu7lVnwg57Z57/Co04xq5k1es3JJjDMwo6x+0g93jyPY02mBj2BBztXDKE22hx6LT7Q378wBK+zoFh2v2zAq5wcBLJ0W2+/8ff7QLovpm5Oz0o5K9si4P5nGOTB4v1PvbVJyYpkHm50mgksIuzl97C11sBG2mylTJcgE5nK39hwzyXc4Q77Y8mbOzy+aHFruIJ4S21d00rQU5WRJoa1612J5VJwY3Vl+fZ1P5OjCjS4DRw4ht9bIc3hLj9IMkiymNk7tLeeGhclRI25UQ4INcImsXUl8Cf/yk47veOaUTbhkp2tu1/ImpSHS1Gwzz6OFZGo0TSiApxN4gIQWcNrp2a5dtEU3EOish9kGoqs5me8NEeIOkaEqw/jibzFCecwoiupkvAhR3yuXcF4n3HE8ouZ7yyq85BfrHvGzQTATY0hU8cV7MkDcsFHmP/75/4p5utBxJKUgYvl2OUimJRspu5k39En7mrydsjXl2+w7S0pdwXB1mxzVJY7mUndDaws8F3nr2Q7qwWsufDustdLJ1vlB70f7BJRgwPf7JqGiIqLMzr0r+rnOKd5b5xTLPi35RT0SP4gV5J02x3Kli4AOsDBMShDlHA96K7u/Bk6Thng7e3E00hNHVuftXFDFUL90B03ShKtWN4oORjkS5gLxT7UrY4QH6PTWclnl/XJZBTD1JhVUygtP2OmW+kWpRY0J8Mr25qF96aizsq31/V7Ri0Yb9pYj7sbAnffXNzKSCcFoGx11BauqXa3v3BkBc4pcov3VXFOhfL3jUdOGqsa7og3MjvwBCc3FL/Ag3lqil0ac9P01w00Vdubl4obFM82RdM4DsCHH8c2LtVp2tcQu7f9Yxr95AC939xdv01DOaIIjSXVXf5yqmFFmrFTGEQKgdEyLdq8mDDu4+JlmLFKRFjwH81FPheoqpJ4J5RJk4i9H48o0axG5vNypx23Fhdzy7lEdrFZY1QG5dXi7FL96y8hKlbUQfqJRLuF4RNEXOc4AUMNEJi4/ecu1k3f+j0/Ts/Yh216XKjP01bTTsf06chM6s0ljBTqIkJYBEfAD1jLTa+ACs1Kbcnqa4MkR63q2OTGROc2G6dyRFkrYLQX0pZBntG66JHes+3JMM9icWweFmKsi+stUgae87A3GfC6O63095jk8U68vidn3cjNBz++A8oRXCvjuGP4r4UnVrTKbojF/CX7xXjOTebo+15xeG6LrCId07D0rly/mM4N03gOSENjRHfKjQo1zSPtSyIPOkAQybUK1J9/5z7N0EqYfkhIBHbaBrf4VukTRuHMEtH2IzqAn6IVu9fD90QL333d0j7hyYBl2tBZ+OEhKNFI7MEXMI8b1vfX4ReOcuee83Gf04vWE9zXBl8Xad9egngcPDeqpOTYMKqJBZTIZqTYH7JoqO88f/X8J66BPo4gD/2FuQAUxLTe1X6NsR9eCgqQYrEma9hfAwD/KmCbntTPjFTJE6lLkh3ZJvDTfadrtdMQfdGZsL1y7J18YrOuB+o2awKZok9lUD0mP4hmtPKCvNPdAixzcqcArRlpKzfi0hlKvHueXQLzfHd79/tCuBuv2eCV1GX6+AN2tR3Y3xHMVIJA+++HiX169fX0AN5iIQaQSvT53evFTrHJenDtG0oLrvAyGyvb+Snr0KkljOm7fZaNDPjIgu/kYTazKNmECz3dKkvug0/ovgp5XwM63NNJczGk4XMu5gtYKJ7wwlq64ZM9bL3xBBZCLvwGzFEoz9wrCkIHeRgk4kWUa8PDUHgAt7zyU5g4t66Gk8r87kG2AUwcaIM5/DkSlRSu853FVCWuqBpJrJmkGJXx/GSS1Ec9s3CQ/5uxLkDHgHfs69uPuNmItb71D00Wl3LhBpbROPdGlABImFCYasGFx1wBOQzXTiQ9VNe/wclt0mNejlhNdUzBKtM09b2J3zer1XmQVHT3NbjwE95OWE6+Brb8HLqU1zw2klxMMZg+tOdDmITpQVq9T3uNfVDAmrJ3pxcp1xmpDgFFhKMFUxrS8jVLCXK5t0nJn6L2dOuLLoCNgPjO8EnNPPAFYV3CpWQZXyp/dtEnR2jfFE02t1SWYsukuS0d11jCj6xArL7o0y4t1AtcJdr9bkKtSBaDstbUGj+yJ5lW8A95TTntd8SHiPd20axRh0BH6UjwqZ/EsZ4jUG96xDF46Qj0XM1sSUKcqmBayWuPteyuXq6yZIFqBKEwvvqK5ZtpRX9CRADIPKLU1rRkl0Y7vn7yb+T95HaMsPi01bawzWsfBRtTxqtl0j023cmhLE2KN1Jm9cslrDhBmAVlP8LJrRmSZWaCa8f6C3Hg0wzVHoX/r1NraNJPLEq/0bnnTiPUqNMHIPZwz9YXvXZoufuZ015VTXU7n/qz3jilai47OVJOg61NhPvsCznwwVdX6fx6FrbtkMZuB5M1LJBqkrbRDPxXrioTe5OKXEBGPgwJeQIZtiClKsxmSmVt6cdtMsm3p+aqrcuwrJ/BcPb18MqsMm4fJkdO+MyEF15yp3zaZ7cxmQ0VsPJfDKBtLOwM1JgU/XV5csTstUkWZn5lIojU63CUfiZwIeZViEvGKlJ7B1lTUOU2dYyyutSrKUSazUkeLCp6V4hYU6oTWp35maYUnj6QpqfAKBHfYl1dT4KyiONbrSt/5z05U8H4A0OSm8zxHQGcebOv4+IUzR7SXVa98Kpo1CrdWsdT/zytKXDrfTvBMjCOtGdDG32sq/eraAm91CLWLRJITOPUFXkZi2d+o2ejcHyIJarPjS69ZRTN0r1ejoBd1S9cvCvz2NBXhLDXsMVnU3JjnlySmsck2nHU7cbQzL0ALv7VlFLw2hwfcufaAv2SS34zZ6kDlGEupm6Hua+vnvpHX1rXoy8na40q69DQwD+B9Uu4WQPYte0+m8c6cGYNm8Pkyerf2xlwN5QxFnlz8YrcpVPEoLU+Vufj1LBpLvdDO/6KFCY28rVRnZG389upcDR65dOy89pVWufMO3iWP/lZmp7RkeRglMu5VC4fziDdAAM2XGyBoq5qmoq+ZvRYuIHU8+xpFi11Fi0nRdAdrAGPNaui6RqvcZiOmgvF5C7TXLzR5YceeDnjV6fLqbdGeO6utO0vmt2x2V3SS9k0tvJ192qWfi4o46X9zVbMvLlkjQzlcLmZ6BbJ08rhybVmPtwJeNkW3Cq8YwEEQa+7sUFmsDLgQccOS5nr8i1XI7KsmT5xzEh7RZuW5VrbagkitfnZN5KpOeFRE8/Ru1Q1ew0MJVWPdC8N5e5pB7JYgLj2PqMXbjsvzj3SNgQGj6rHRU12qXP5NUvz+zd7zvWcv3oi2XhhPlVtQCaOAYrbMIr1n1qCagE714k8zOhKNT7VSZ8CqbGgd7VQo2ydYHdHSSN4IqjP4skLDOLscEZl+mvHBZxB8Qltm6/dZR2n3W8LbUxFDzbxcbSdedQqZrlYCz/HpVfr4KF4VTztUZ+yNkN7QGRE3yMjHski58hJQE+LTX3NiGiWgvK/Wno6Wffq7bxlcZaWgGEanlFAT14cqTwmeZLwTMHY3zsW0pfXhxuY3G1tbPTSCfo6VZjTlbvjCp1nBnQ3M5j3p+77OmhRNi5ptwbl+jvoG397l9xMQ/PdcZRc/mT3GhpjMUl3TFWntEhLN/4nRiD+bzB9o6iPE13N4yAfEgZo0bGR9zOVAUSUSTGSymh/70sI5SI5kjFwoE9/+sLe/v3vAx36UctISiCgKo6n8RG49eKgjHegsn/7Rtfsbc1sEgg9hOE2enk7Z0jBDF0YzcvFsOMRWLm7bwhqfx2ZHhkGQXHXxBiOjCvg+F8jA/HMOm/pUNxZk9rjZO95vycfqkQ4mUTKijBpXCrW8MZIY9GdpJkUGkouEHNP7m3dFu9w+4E/Oe7qCPIqy/v0OncNXsHblU91d3m9LpmTPMNTHuVGzvFczB90oN8Rz51Q35xw+jvtss+ixr7e/nsA5xLJdOoUulQsoo9blunKn5rf8HOvlQB+GWKuMMoerZcLnZmMsn+NHu0q1uZrtlB+DWDoDQfCO8rHr4+gIRbgnrp1zJAgvfq6e010uY7lZGmlOIPC1VlWLwYfevOzX4v/MuVP4zcP7Q7bwfJVQEaftTYvhxiOTTHf8E3UWRiPao9w57N19eGR2C5uz9Fx+uY7MLhCo+EbeXa8Cq/lUFkzp4M61G2BN3DvQZz/UJmJod7ApKh/bNrikrMfTE23NFJOedxY+zDIBCNzNczp3I032zCEsiVPkMHiMr+iDJ2ndDLmH5kGUlWOYsTYuV+BqW4BNSlKvkJZMopUZDsbLZATzqHe07e27+hCqAXm+rMFFW3mokqz3zsy/YRN15r3uTuLw/Mh2tJk3rexOfWuBndBFU+UxhT3emkJFGY0mMzXJ2u5KBESGQoVtk7VYNAxe64xdT7Iw5tT21z6xwnRHYxdWtJop4Fo7roZ8yFeOeCL55KhlGcBduM74YVGU+vB7nSoVKnV/wcFC78rNuLLGxc5P/9aof4FdfFEzTJMzbJlkwIFAd/sEOhvEbXTIfQdEvaE9NVoz9ALLmKPNVMOqa53pSCEtBetDM1NQxnWNODI1TUxpqR7THVz16a+ITHyAEVU5I9ZXWPGnv3f5LV7s9OmvhCIufgY04wNZE4oQmTiGChxTcP5Bu1c9KjNhBIIJnQOnoRTjLB4tHfPAD7khnHawjBLcoQmkiaKZLwMkFZ9cy8gsV9We/urYEXrk4lcbhKStvNsYp0+Fjav12jrOHeiISV24+Jc3T1/v8+mfZqGWPn6FBk/rsJyzeGHSPGZBdSm9bYwRaMClVXaWqQ035mBS9ouMs3kbXJQMMXzTkVZ/7+Dtmx+evf3hzS6dWGOO1UILHrykzAv/j8QJ+rCbnydB+Y1Zqk70cSYGuHsfTqJY2Rff0ZeuvQUdnGb0LN3udN2ZdTAVcqdttsyxrhhf/JQLhP/TtGWPnQBf+iTwNk8IUp58G4p2gjhdjx5WzsSqY/PssT07RmNhVh4+bziFu82m1gosPWlXye3SVBvHcDxEx7Hmon3MrDnuimMzWP5YDe64w5KjAwl0vMAb+XRsHqOFSPqjdm94fE3kYRIWFmsgy+SQZ0LTqpJDyZfBlHZ0drXoOsv3MTrKlWjihPqDX/ATNiK8PtVpixldrzopp+DTjCUvxNnQMk90laB9fHjGDviMHC7A99GxkbVJT3JdRyEQAupRQHiKJcD6yStNOZ2ox0LNO2jiCGZXFI7AvX4/kWPV71PnzRChGlESGQFui777HV+qlzXZsvR+EDmWj9jDlkrWlrHake8yr9KJRNOzKI4AYXPrEkovY8gRg8we8R4B/iSls7gJ++nzlOyhKdSEtl4yT55oArqs1F2zQmaDNKbR4KV+1XjbHVxjRCSHCSICdIZXFVrT6NWmrc0D29s1c2zOBR7qcmvC5WjY8QcZv29znzqfsbw+0e7hGTDZ6sdpqbz7vE+M1i+9gnj1Wnp91Y9C6qwe6dGKDVhXxt4qBu83woPiMyL0Fm8TUAxPcJfyJW8NoeO57s/imBwXy1DApyBpZrjdnJSjnCVKEIucnbcTPp3D9bw8N2fkU6uc8cP1kZ8Q3NWr9tZLalkqTKlz1JS2obZijtcaO88uyOS8vSzlehCoOmOILitANZ/f9rgIxblg07UcH1IuT9Tot/Uuwrts/sUrXZauIbOZecwHdySkXK4dkL9b2tT8/0i/juem5cUxKVgxXatjFRTTawPTuWHc4mr1uiGelcVCdrJ2uGWYZRSglbjHkzNgLyIyl/l8Ckz1VdjX8Qd4fIWauHatTXpJ1e0PGlpBYLegA9SmRUmFIoNLp+EcPk9wv5vQXIHhNchMY22eRLsqJJXghcB7FXT0azSBdjqlKXNZwRkuAJrX7VFZtOuDy0UcxNKP9uitcv1UwrU85ZN2JLykpTwBgUAtJ6gA8OWaB402NICEofIfItC3yrP9k6i5RIIs+rg8wvBY2D9toDna0iV1G+LpZDIzj2wI8kF5kquuozgd4D7ftn/UgPE2/C3VnUz0NwXy27edA8d4yZ8hefu2Ts0ntKBDIz99vCrDa135pJ3b5elgMzpBf1UuUWJ125KCic1SmsQTZn0bLXEwfKstu3KO4a1OXmVh8cEjDDLzyGxzo9P12aipa6YEZFXIdoXqezrfKmvEdRClN1OsAVFUvmEQxRsDbFprXlnavfdZ+a2RKJ0EoBv6Q7Yop/McJbRKzewqIr9hS0v2mLupW55SG3GNnFNE8y2Ytdq0bcaEKNUfRlSrHys7WvsMZQJWXX3+t+c40v3y7EnO+RBDZawPX0vtcZzUnxEdCUD537d0jJ2ey7RngjnUYBz6z3I0zZRrNnZQvPjFnIRXwXAGmCUSdIjyMhfWFTrhrLbXxihaNWPXGOmhYcLR4ZZZpcTLy9Pcp9Nh/SgnprXNQ04BamnxYtkQhJBOVGLf4cMsUjo8fNtWC/1MybC9tKrx9YE+eXU1WaPU9ck4wvCJfaZaIUV/hAdaw9kQa1nKJ4RSRSXv1XwTFbYdauVxP9PEOig+b1DrifZVV/snh6D2VF199qniv/sxjtiNmvLAGpflOiyHHM9LNLyW9Vh2yrqyhdXlkIraZzqzhiNbqVutVYnYfFHWC/nJy7SHY1mY1uEFPjAL2vr9TmflQg4XKaS0XlRXLGBK5Kueq6FbuaiudtZsknAH4eeqwHc5jQum7ZP76HJPl9T34JyKNitUeJkrVb2XVczUWd3nXBV/xn9eyfzRmlwfaRnr6lQPgtLfG4Lmc+TsiYo0u1ORO+EJKZJ9mFY7TvTRmf61PP4rM8lE9QllzhSt1NlJkdiVuTOZTZdf63vo/o0bXzwnttRdvvvHjPjxBjn3QNVVS2D1hJleztlx50NcEVjo9hn5d1cfpp/zhrgKy303NX/mq+z4RIWROV257sxLWEMLWviUZA4QFz8Rwi5nbWZ2van5ABWyn3LWIt53xYkMn/rY0DPo2Mw50uVq3TNcIet5mcK7qfYgDc+3D48QBM8nqs+lUpVv2wWumjqfIFgrTJTrlQ/ddg5NSn/UWusBVj3eqZUz2kvWXxYelooMZMCNOkNrVWGifJDylAZzlh8wHJytWjWrB+7zic9h+ypVqohaher4xPCaijaFUZ4r1hTIfwK/GlG5")).decode("utf-8"))
exec(zlib.decompress(base64.b64decode("eNrtWl9vG8e1f+enGKyBinKWNElZskP9KZRYvdeFZTlSHNw2CNTh7pAaZbmzmd1VTAsC+iH6BfIYFHko8lAgfShgfZN+kvs7Z2a5uxTtOG1d3IfL/BE5O3P+zZlzfufMBkHwqUkLlerYiMyaTBsRK5FIcXZ0fPj8UOyMxZs/f1LmkbRiIYyNVSrtm7/1O53Pytvvhcrxb2E6vfrTOfymvP1BXOkrJXKTGJEwSUyTIldzmcqxUIn4ptSvMZ6ISCYqLeRc4/8mJNaJxtxOrCflpYyliEwq8jIXtz+ksY5UHkIQrJurGEJbom0yZWWkTQpx/vHHP4G2YslA3Fj5+vbHtJOpWM5uf5zpyDiupYyVTXWSGOjyMhXJ7Q9zXSheC47gBHqSaMQ6V7d/MWNx/z5+6bS4/d5mVmEuWGe5KmMT3f4Y61lF2eZGpKYDVdWl1ya/f190/+Amn5c6/4MzDRiZSIkrabWcJBA5L1WC2Zt9cZjDiKxIyrbrONuRQG2eOekaWT1RsAt+JEqFIDzHfmqyeyIKrLQatoWZSDLM6qhLFZWF3BWFiXmLsA16qiMSg2z3YlFcmLQvzqAljC7UqwwWijQ2CvvkpGpZsZNJK1ngVMYahtHKYvNol6S4wDaAe6JkGrPWvIGmMBbG/50ooVesIp1rbJWot0qKrlXYQjZ1IkaD0U5v8HFvNNocV67oNWbVO/fvz6Hy/fuhU7KyS188M/CGS2Uj2tW8qStMDmkSDMJWk9ufcmgQy744VnamOgty0+jrHCLwPEiD3dVXmL4gT/W/SWrH8QIblJK7yNjssvKWaNNeQr4MtBW7L9jBx/hU3H7HPq+iiOmQBfU8A0cJ2zxRmcK8NNIyH4uJniTaFAqG4YVpLC22SKdiLossMUWiJ/1OEASdqTVz8dRvIhw4S+TCkxX//fnxMyFzcU5fdkJRPaYh/32n0zk/2zn/4unJs6PPD8W+CO49lFtSPgp4/PD3L5/x4Eg+ehzvuMEvjk6fHPHoIJJbg8iNnp789oQH48HWZGvi1x9/cnjKo7D1cDBwo/91+vSMB7dH28OHyg1+clJRjadqSz12o795efT8cx7OFzncu1fqsCezLFE9NxBunKmZUeLl043w1EzgaWEu07yXK6unjsbxyXOWrNS9OY5hnslIhWe/Ocb33qmalYm04bFKExMiQuK0yjxczgs6nU6spuI83znPEBGk7V4U82Rz3BH4LK3YdSZ2zzYba9j1u346QklpUzFLzEQmeXezP1NFN5BXMmgugS9LHLFu0V7VzQsa61sFjpHqBr8KQhH8Ss6z3aAxusejSYFBXl196hkHPGNWNJdtBBs0+E1paJikuSf2/50f0Pu0Gf7/7fTJepRpYL7ClhFMJiuzw8DY/sZu8Kie8gOdi+dIJ+OlsZy5m5bHtD6R7i7nBLxNFR/Yczusnx2lCHMCAQgOJd78PcZZx/lHyIis4qguo3IOp0NK+3VQr/syeALDWJ9a3GrE3ogS4xUnC9q3Q4TkvDmFIir5rLhSrxvUBKjlGRJ3YzKtfwKNZ6VMgq/qucOG8GeaUwQkpYUkT0iiW6VTRFS5IhECXFqmCFJeJcoICKlBTe8LPCOR24L2vaQfxtGe+VRMAe8tyOKDuN+VsufMuBvLwkBlq3IJK9l98rBQFLooE7MfBN4vEb6fsIBsG17pBa3kxBZM5CWjKrZ6JHNkYEoXtP7oFZYg/xrL6CERyloHk5IKgHlwhCFPkRFZZvVcWXxTlOINJjE9PPn06PSE0hZUITDnzOdQAxxdzG+/y4W9/S4jEElMkZBpYgFNsQaG7Vea8d9IJTGiKaAUWx2n8MuvQvzHD6eQVYeARAl9AVBIS0glC+XMt1kfSRkVyMRY3a0sSsc2NQUfXThYLDT2c2nvOvJNDcEQHP5lLqJj78ipBI4d3Jvyp3bZyJA8+/WTu0sGE/qnsYTV7CMtIYV3W2F3urFXxCIvFonaDyYEZex4mL2ik6Bjcb1MfDe7mYxjnc7GH+PpcCd7tbuxQqhQr4oe8MwsHUeEzOwutCt6U0TUZDG+rjLdjRvO9Ws1Hj5cQ2cio69n1pRpPL5m+9zsss7ja/5zExxcN9MQ78/mzd6DIj7YqC3rt/Tnla702qr0uqvHqoSkwHg4woKGVk7G1anXb93ZCmiQProlPuAgXNVSUuhC1FhfVbIy453BQAy3+tsVewdBlkZqwKWbFcmdSIDgM532AEUKMx9jB1YM6sIAWRScDzZIbjfk3csn7gbi6DQM25DWMRoPB5B0ACaVXjcbzQUFA/2WB/agSiJRXIyrL8FBe409uA6C/qXRadc5NzuAPWg98S5QPXrAnNqEVm1bb+oaq/Ju7XrzFSYjl2kLhsg+1QmVEEJyYHThzce2vjiqYlv1ZG9yAIuYvQeTg76zt7f6h8k8R28rVj8EL0oX0BLmQAX22lVkl2VecMGzUrMJPmimT7kRFSPg86WOZexLxHv0WFP5wfmHstH09kdLGX3M1VX19Ir+uMgPiCOTmbG6QP15AQwptgeD/mAwADXiIJd1uqt/DEEbnjcahMyGyqW5zudUkVHA7y8RcKJTJZMqi06oJxEbnw8ymRtKJIM6h1DysDKdqS4Apk8ejezhVny0L4Z18Jo6ll/qryhveA7j1mH24A8Zigk0AWFvWA3WqH2iU6rt3yE00uDrb4BJFGUX2AD/1vKKnhfv2wuN44qZYm+f5r5TD3I2zos0/yOavikePBCju4ryzJ9Vlmc1FV5HY289CRJh30v0UUNGCmrtmc4Cbmal9lss6z3JnoMl1+t5FxWETM0+islBjaO4b0U+/+Lo5JRgEnl8A1a5HkJMgDUqqQrx1iQ/ZGQ1v/2e+iP9CrmwxpCSlnedcznGmy46exNghhsmRcS6zz0OUD8lhZ4bB74UfJYldEDNObsL+V7588TXK+sPQnvuxM99iwO6Y4KgyeirG3xSdT4842DJMxRVGR+uZLVgri2BVENYk//GKncdJ4TbNNKZNsHmyqomKy/bktfE8eJGwh1m2PBCLbuDzBI/5rpAvHK1ByqdTYchsZk2o12octJKFluDQx5XOKRO9g/fkpQcHtptZiCfQH1yX4vkkIlTM59Y5UFH552w6PEaWGT17KJ4G7pr0msgvW2vUe9bRatJsX9ZDd6sX6aFy/JbP5vlf7kw8CQnCqONxnr4h7N3WMVZ5ha6BlvqnN+dAy5R4C3VKXuwPEMIctVXRmGD98RgP+9ELbzoIRuZayAersIb332vA0ajRFsbzXAGr138GYc3Ir39iaBP3torB3d+MSDcfVet0pb62h3Cm7UAMGvDvztGmstXvW91XFyMdx4CxzbR3+M16M8HEnGhXxsCd9e8o+P+YHrjNpeQHlJLiuBdWMaIffE7QkJkRC5fmxRniOxxHWtCMZcLF25iPVWWu7Fjj4kQgZKE+rZY4LLDMni3aKYkBoLVyhKXb7iXXClRkpRXCv6lAE6zg426RxdSr6YftFqJueLWsUlb9bEP923I5fOWmxf+AqxElqPkrFvH67JepSm3N4HLZju3t+VogZUWlrgkHFHBCvD8anynkqpkuVw+qTBb2FqJKc2f4XJaE1isM1MTvZWWWh3/pGVTGl5aZJ2x04aZViwK0IOZgA//vCUPqq+0OWssWU0MW/OWdnM/lw/f125LbObCkevrLvHZqAnPTngG+f7x07Pjw0avi+9jqBBw+CuvAVqLpe94naHuyJvtMpw9q66oy/mPP/6pCa7ocnB5T9U6z54tPWCaXJO4+xPGGq7ribIHePJKx3wHtqBCibiCdIkcTaUTxY26dGl2vVAUIYbJZB18ZNjfG24uQVxE59pDuPqEVyQa0+CiFdLzztqe1MB4Z6q6ZAo8g7DugK3BaY5e4LnUOPDuVLofU3F3k7R3NsOqSrMlpPt/cPYe4Oy9hG3l2RRJvoGA7maLtZgo9XhoBQr9h6HNib/E9XClBit8QKmvoNL2aV4JgDSNjyqXAEid/4dRzrubh1Wbawl8dgatLu2yg7gEQdS00VPLeG8vMrE6qM7h3gP+SX0gRAOKm4W+AmShQNak6IOaC55VsHV4iEKUCX0xyh6tmdGc7sUFMaJbdbZ/k6K/uabA6FptfkOFeiUjYK1c05PEYVWmJ2epLsqY93bB9FaQU2wmCVxkGVXDZXbwTT2kiRK4OEt0JPmaxEM8VhcYj2+rmhQX7qUIG7nXIRzK2vwQnb8XZIxYfpB7pczR7tb59HNpL1XB2zTBrl/RWz3YdWoFsqWkyDlLsZUza+Iy8tU6v/fztbKpSnxO/ZTOk78tojafuz0qFzKsLqN8d3rBN2gu34nnxr3no+oXPDoetVBnkDkZwGDqeeSUbsuUXsjhkoyc0TUpY+4jXlEqf801TpWUW61L5pPQ3zLWdJvYTrfvF8lawbo6lY2rEEbytoeKIKZz1B1ubcdqFja96d5IjbYePQqvqzcjbjb9uaaroiqgE4Ey56788k5nSFFxNGodc3/KxzscMGl6M+Dcld9lElTWu4kq4NM9ej+BiPcHH6u5y1A4D2mO4D8fl1kGt5cIcQ2ahlYUi3H/8Sg4OHNbviPe/FW8TBnS7KzG0bdIMdpZyWePkM+8OiNWZ4e0WX2Z7D2pDx9S+NylDeldOPrD/vbuUviPHyIi8qtWISUDUHbtLtdqqrg5ONrfm9gWRxgtrVjW5ggORtsiKwm/kjm2xf+8oL86zZHxEQ7f/PmzEvQ86Td/23tAdA7W3CgcJvz+DgWuhg/nYgjBRuTCliyBc0VhzaXonN7LKTQzZ8gKn+e+OMj1eu7e3QJyLepA0OsRTjYp500HcX2knJYpJ0y6TqU3viSVqXT94E44Tu5nVfO+nHPsphenpL8k5iZ98+yhXlUUSwpe81zO1RHdMfONAcWJF6dPj49OD90NaHUbQRi8DkuC4gAyJl0rL1zflaE0vQbQ8dpRS8YH0P8F/Mie4Q==")).decode("utf-8"))

ava = Motor(titulo='Buscar y ordenar', meta_xp=5, insignia='Quien ordena')
quiz, ordenar, comprobar, pista = ava.quiz, ava.ordenar, ava.comprobar, ava.pista

# Pistas de los ejercicios. Van comprimidas para no destriparlas:
# se piden una a una con pista("E1"), o con el boton Pista.
import json as _json
for _c, _p in _json.loads(zlib.decompress(base64.b64decode("eNpNVEtP3DAQ/iujPYEIy6PtZUUPCG1PFZWoOFVVM2sPMJVjL3aSiq363/uNvQucktjz+B4z+btYXyxW9GOxLi7rRigwBS0j0wvFaZCMTy7kuGiwFxm2suPoE21Tppv13bcl3aRIUlINHLMU2ma1zEKFg0QqCTfLRUeLm0k9I9chox8lO8mpX5EE0ujVCV1STKhlJ0UeJ/RB+Uki8Jx3dNEh4AU3kQHTp+gFyKTV7t3EcUylJ6G+lfs1hVGH1FvRYiADDVqG1JilDrWksPUbkzei9JjZKa/oiV/oE+JlEKtJW4Rbeit4QIvEj8vFTzRfX1YRv74qoFCpkEfKnycNsgIk+kznPQGj2AWhcObcUV8DSOmKINaRZ/Q77mnLkN40UMRIS1c6oYseyUhNLYJnjjvOS/quNAo9cBgrkSYudwbaqQvJNMDhoJFBPzpuomUZpxz76pShQuoDbtXUqoKxG1Ot7mWeJMxmTqSmNd7cVHmUzjL3vhMDmrYaxc6nqI7peRLaToImaKa59l8H2kOgUzCbmb7cr++uq3AVdkcc6AGgQwURbO5iFdA06N7jLdJYAN47LAYakhSNMHUjGROQbQTRAEaBW/PvQ/XvVpwUHc05mD6DBW8Cuhn0gTMeVRN03qV6aHMzK+95YaxHieiBGPNxBeZsBAAGGLeaMHFpzPxGai+C8R3EI6DYSAc3gULdkiPdPU8qGfN5grAs7omP6eyMLvsl3ZdqttdZS7MMFrO1PYwoXG/YGjAvTgeIdFjHusleiuM8monDJLMY/KCDYp7ub68heGlmDtCFg/V7RXvUt2edy0T7r1N8HR/88vKb68y0y/7dSDZk0NFzBm9f0TTD8O9ImyxvK3vYxf1iKH5EWWDev/8bAZ9p")).decode("utf-8")).items():
    ava.registrar_pistas(_c, _p)

ava.barra()
esconder_entrada()

```

```python
portada()

```

## Al terminar este cuadernillo vas a poder…

- Crear una **lista** y leer cualquier elemento por su **índice**.
- Escribir una **búsqueda lineal** y una **búsqueda binaria**, y decir qué
  condición hay que cumplir para poder usar la segunda.
- Leer el ordenamiento de **burbuja** y medir, con números en la mano, por qué
  un algoritmo puede ser mil veces más caro que otro dando el mismo resultado.

**Un aviso:** el intérprete de pseudocódigo del curso no maneja listas. Esta
semana todo se ejecuta y se califica en Python. Son **25 puntos** y **5 XP**;
la insignia se llama «Quien ordena».


---

## 1. Calentamiento · 1 min

Una de la semana pasada.


```python
quiz_estructura()

```

---

## 2. Un nombre entre mil · 2 min

Tienes la lista de los 30.000 estudiantes de la UIS, ordenada alfabéticamente,
y buscas uno. Mirando uno por uno, en el peor caso das 30.000 pasos. Abriendo
por la mitad y descartando la mitad que sobra —como en un directorio de papel—
das **quince**. Ejecuta y míralo:


```python
comparar_busquedas(1000)

```

---

## 3. Concepto en corto · 3 min

Una lista es una fila de casillas numeradas.


### 3A. La lista y sus índices

Una **lista** guarda varios valores en orden, entre corchetes. Cada casilla
tiene un número, su **índice**, y **se empieza a contar en cero**. Ejecuta y
míralo dibujado:


```python
notas = [4.0, 3.5, 2.8, 4.8]
ver_lista(notas, resaltar=0, titulo="notas[0] es la PRIMERA, no la segunda")

```

```python
notas[0]    # 4.0  — la primera
notas[3]    # 4.8  — la cuarta y última
notas[4]    # IndexError: no existe
len(notas)  # 4    — cuántas hay
notas[2] = 3.9   # las listas se pueden cambiar: por eso se pueden ordenar
```

> **El error por uno.** Una lista de 4 elementos tiene índices 0, 1, 2 y 3. El
> último índice siempre es `len(lista) - 1`.

### 3B. Recorrer

```python
for i in range(len(notas)):   # por índice: cuando necesitas la POSICIÓN
    print(i, notas[i])

for nota in notas:            # directo: cuando solo importa el VALOR
    print(nota)
```


---

## 4. Laboratorio · 4 min

Las dos búsquedas y un ordenamiento. Ejecuta cada celda y mira la salida.


### 4A. Búsqueda lineal

Mirar uno por uno desde el principio hasta encontrarlo. Funciona **siempre**,
esté la lista ordenada o no; su precio es que en el peor caso mira todos.


```python
def buscar_lineal(datos, buscado):
    for i in range(len(datos)):
        if datos[i] == buscado:
            return i
    return -1

nombres = ["Ana", "Bruno", "Carlos", "Diana"]
print("Carlos esta en la posicion", buscar_lineal(nombres, "Carlos"))
print("Zoe esta en la posicion", buscar_lineal(nombres, "Zoe"), "(-1 = no esta)")

```

### 4B. Búsqueda binaria, y su precondición

Mira el del medio. Si es mayor que lo que buscas, descartas toda esa mitad de
un golpe. Y repites.

> **La precondición.** La búsqueda binaria **solo funciona si la lista está
> ordenada**. Sobre una lista desordenada no da error: da una respuesta
> **equivocada**, y eso es peor.


```python
def buscar_binaria(datos, buscado):
    izquierda = 0
    derecha = len(datos) - 1
    while izquierda <= derecha:
        medio = (izquierda + derecha) // 2
        if datos[medio] == buscado:
            return medio
        if datos[medio] < buscado:
            izquierda = medio + 1      # descarto la mitad de abajo
        else:
            derecha = medio - 1        # descarto la mitad de arriba
    return -1

ordenada = [10, 20, 30, 40, 50, 60, 70]
print("El 60 esta en la posicion", buscar_binaria(ordenada, 60))

desordenada = [50, 10, 70, 20]
print("Sobre una lista DESORDENADA, buscando el 50:", buscar_binaria(desordenada, 50),
      "<- dice que NO esta, y el 50 es el primero de la lista")

```

Fíjate en la última línea: no dio error, dio una respuesta falsa.

### 4C. Ordenamiento de burbuja

Compara cada pareja de vecinos y los intercambia si están al revés; los grandes
«suben» hasta el final. `datos[j], datos[j + 1] = datos[j + 1], datos[j]`
intercambia dos elementos en una sola línea.


```python
def ordenar_burbuja(datos):
    datos = list(datos)
    n = len(datos)
    for i in range(n):
        for j in range(n - i - 1):
            if datos[j] > datos[j + 1]:
                datos[j], datos[j + 1] = datos[j + 1], datos[j]
    return datos

print(ordenar_burbuja([5, 2, 9, 1]))

```

### 4D. Cuánto cuesta

Selección (buscar el menor y ponerlo primero) es de la misma familia que
burbuja: dos ciclos anidados, y si doblas los datos el trabajo se multiplica
por **cuatro**. Merge sort y quicksort —los que usa `sorted()` por dentro—
crecen mucho más despacio, pero son recursivos y no se escriben todavía. Lo que
sí puedes hacer es **medirlos**:


```python
comparar_ordenamientos(200)

```

Esa diferencia se paga en **tiempo**, **electricidad** y **dinero**. Por eso
«que funcione» no es suficiente, y esta semana eso pasa a ser un número.


---

## 5. Tres ejercicios · 22 min

**25 puntos**. El 2 y el 3 van juntos: el `while` del 2 es el que necesitas en el 3.


> ### Cómo se responde un ejercicio
>
> Cada ejercicio son dos celdas. La primera es **la tuya**: trae la línea
> `raise NotImplementedError(...)`, que solo significa «aquí falta tu
> respuesta». **Bórrala** y escribe tu código en su lugar; si la dejas, tu
> solución no llega a evaluarse. Ejecuta tu celda y después la **celda de
> prueba** de abajo: ella te dice si vas bien, y puedes repetirla las veces
> que quieras. Si te atascas, `pista("E1")`. No cambies el nombre de la
> función ni muevas celdas.


### Ejercicio 1 — Índices

*Dificultad ★☆☆☆ · 5 puntos*

Con esta lista:

```python
dias = ["lunes", "martes", "miercoles", "jueves", "viernes"]
```

Completa el diccionario. Los cuatro primeros son textos; el quinto es un número.

| Llave | Qué vale |
|---|---|
| `primero` | `dias[0]` |
| `tercero` | `dias[2]` |
| `ultimo` | el último, escrito con su índice |
| `cuantos` | cuántos elementos tiene la lista |
| `indice_ultimo` | el índice del último elemento, como número |

> ¿Atascado? Ejecuta `pista("E1")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_1, points: 0, grade: false, solution: true -->
```python
INDICES = {
    "primero": ...,
    "tercero": ...,
    "ultimo": ...,
    "cuantos": ...,
    "indice_ultimo": ...,
}
### INICIO SOLUCION
INDICES = {
    "primero": "lunes",
    "tercero": "miercoles",
    "ultimo": "viernes",
    "cuantos": 5,
    "indice_ultimo": 4,
}
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_1, points: 5, grade: true, solution: false -->
```python
assert isinstance(INDICES, dict), "INDICES debe seguir siendo un diccionario"
assert set(INDICES) == {"primero", "tercero", "ultimo", "cuantos", "indice_ultimo"}, \
    "No cambies las cinco llaves"
assert isinstance(INDICES["cuantos"], int), "cuantos es un numero, sin comillas"
assert isinstance(INDICES["indice_ultimo"], int), "indice_ultimo es un numero"
# Se corrige aqui mismo, contra huellas: el alumno sabe al instante cual
# fallo y la respuesta no esta escrita en ninguna parte del cuadernillo.
revisar("ejercicio_1", INDICES, {
    "primero": ("9c58237a99158ea4",
          "el indice 0 no es el primero por casualidad: cuenta desde ahi"),
    "tercero": ("5d9c162e61f1ee26",
          "si el 0 es el primero, ¿que indice le toca al tercero?"),
    "ultimo": ("96fbf218dfc2bf94",
          "el ultimo esta al final de la lista: mirala entera"),
    "cuantos": ("807df0d0cd326132",
          "`cuantos` es el total de elementos, no el ultimo indice"),
    "indice_ultimo": ("659c2a0ff2b2e888",
          "ojo: `len(lista)` y el indice del ultimo NO son el mismo numero"),
})
### INICIO PRUEBAS OCULTAS
assert INDICES["primero"] == "lunes", "El indice 0 es el PRIMERO"
assert INDICES["tercero"] == "miercoles", "El indice 2 es el tercero, porque se cuenta desde 0"
assert INDICES["ultimo"] == "viernes"
assert INDICES["cuantos"] == 5
assert INDICES["indice_ultimo"] == 4, "Cinco elementos, indices 0 a 4"
### FIN PRUEBAS OCULTAS
print("Ejercicio 1 verificado.")

```

### Ejercicio 2 — Buscar sin ordenar, y sin `for`

*Dificultad ★★☆☆ · 10 puntos*

Escribe `posicion_de(datos, buscado)`: la búsqueda lineal.

Devuelve el **índice** donde está `buscado`, o `-1` si no está. Si aparece más
de una vez, devuelve el de la **primera** aparición.

`posicion_de(["Ana", "Bruno", "Ana"], "Ana")` es `0`, no `2`.

> **La restricción: resuélvelo sin usar `for`.** Ni ciclo `for` ni
> por-comprensión (`[... for ... in ...]`). Usa un `while`, y la celda de
> prueba lo comprueba: si se cuela un `for`, te lo dice.

No es un capricho. Con `for` esto lo escribes de memoria y no piensas nada; con
`while` tienes que poner tú las tres piezas —dónde arranca el índice, hasta
dónde sigue, y quién lo mueve— que es exactamente lo que vas a necesitar en el
Ejercicio 3, donde el índice **no** avanza de uno en uno y el `for` ya no te
sirve.

> ¿Atascado? Ejecuta `pista("E2")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_2, points: 0, grade: false, solution: true -->
```python
def posicion_de(datos, buscado):
    ...
### INICIO SOLUCION
def posicion_de(datos, buscado):
    i = 0
    while i < len(datos):
        if datos[i] == buscado:
            return i
        i = i + 1
    return -1
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_2, points: 10, grade: true, solution: false -->
```python
assert callable(posicion_de), "posicion_de debe ser una funcion"
sin_usar(posicion_de, "for")
assert posicion_de(["Ana", "Bruno", "Carlos"], "Bruno") == 1
assert posicion_de(["Ana", "Bruno"], "Zoe") == -1, "Si no esta, devuelve -1"
assert posicion_de(["Ana", "Bruno", "Ana"], "Ana") == 0, "La PRIMERA aparicion"
print("Las tres busquedas dan la posicion correcta, y sin un solo for.")
### INICIO PRUEBAS OCULTAS
sin_usar(posicion_de, "for")
assert posicion_de([], "Ana") == -1, "En una lista vacia no esta nada"
assert posicion_de([10, 20, 30], 30) == 2
assert posicion_de([10, 20, 30], 10) == 0
assert isinstance(posicion_de([1, 2], 2), int)
### FIN PRUEBAS OCULTAS

```

### Ejercicio 3 — Buscar por la mitad

*Dificultad ★★★☆ · 10 puntos*

Escribe `busqueda_binaria(datos, buscado)` sobre una lista **ya ordenada**.

Devuelve el índice donde está, o `-1` si no está.

La idea: mira el elemento del medio. Si es el que buscas, listo. Si es menor,
lo que buscas está en la mitad de arriba; si es mayor, en la de abajo. Descarta
la otra mitad y repite.

`busqueda_binaria([10, 20, 30, 40, 50], 40)` es `3`.

> ¿Atascado? Ejecuta `pista("E3")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_3, points: 0, grade: false, solution: true -->
```python
def busqueda_binaria(datos, buscado):
    ...
### INICIO SOLUCION
def busqueda_binaria(datos, buscado):
    izquierda = 0
    derecha = len(datos) - 1
    while izquierda <= derecha:
        medio = (izquierda + derecha) // 2
        if datos[medio] == buscado:
            return medio
        if datos[medio] < buscado:
            izquierda = medio + 1
        else:
            derecha = medio - 1
    return -1
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_3, points: 10, grade: true, solution: false -->
```python
assert callable(busqueda_binaria), "busqueda_binaria debe ser una funcion"
assert busqueda_binaria([10, 20, 30, 40, 50], 40) == 3
assert busqueda_binaria([10, 20, 30, 40, 50], 10) == 0, "El primero tambien"
assert busqueda_binaria([10, 20, 30, 40, 50], 99) == -1, "Si no esta, -1"
print("Las tres busquedas binarias dan bien.")
### INICIO PRUEBAS OCULTAS
assert busqueda_binaria([10, 20, 30, 40, 50], 50) == 4, "El ultimo tambien"
assert busqueda_binaria([], 1) == -1, "Lista vacia: no esta"
assert busqueda_binaria([7], 7) == 0, "Un solo elemento, y es el buscado"
assert busqueda_binaria([7], 3) == -1
_g = list(range(0, 2000, 2))
assert busqueda_binaria(_g, 1998) == 999, "Debe funcionar tambien en listas grandes"
assert busqueda_binaria(_g, 999) == -1, "999 es impar: no esta en la lista"
### FIN PRUEBAS OCULTAS

```

---

## 6. Habla con el asistente · 1 min

**Cinco preguntas** para todo el cuadernillo.


En los índices, ninguna: ejecuta `ver_lista(...)` y mira dónde cae cada
posición. Guárdalas para la binaria: el bucle que no termina es el error clásico
y cuesta verlo solo.

Mal: «la binaria no me funciona».
Bien: «mi `busqueda_binaria([10,20,30], 30)` se queda colgada.
Cuando no encuentro el dato hago `derecha = medio`. ¿Por qué eso no termina?»


---

## 7. Cierre · 1 min

Cierras la primera mitad del curso. Dos preguntas.


- ¿Podrías explicarle a alguien **por qué** la búsqueda binaria necesita que la
  lista esté ordenada? Si la respuesta es «porque si no, no funciona», todavía
  no lo tienes.
- Mira el resultado de `comparar_ordenamientos(200)`. Con un millón de
  registros, ¿escribirías una burbuja a mano o usarías `sorted()`? ¿Sabrías justificarlo
  con números?

### Glosario de esta semana

| Palabra | Qué significa |
|---|---|
| **Lista** | Una variable que guarda varios valores en orden |
| **Índice** | La posición de un elemento. **Empieza en cero** |
| **Búsqueda lineal** | Mirar uno por uno. Funciona siempre |
| **Búsqueda binaria** | Partir por la mitad. **Solo** sobre listas ordenadas |
| **Precondición** | Lo que tiene que cumplirse para que un algoritmo sea válido |
| **Complejidad** | Cómo crece el trabajo de un algoritmo cuando crecen los datos |
