# Buscar y ordenar
### Semana 6 · Unidad 6 · Colecciones lineales, búsqueda y ordenamiento

Hasta ahora cada variable guardaba **un** dato. Esta semana aprendes a guardar
muchos en una sola, y a hacer con ellos las dos operaciones que sostienen media
informática: **buscar** y **ordenar**.

Y vas a ver, midiéndolo, por qué dos programas que hacen lo mismo pueden costar
uno veinte pasos y el otro un millón.

**Empieza ejecutando la celda de abajo.**


> ### Antes de empezar, dos cosas
>
> **1. Ejecuta las celdas en orden, de arriba abajo.** Una por una, con
> `Shift+Enter`. Cada celda usa lo que dejaron las de arriba, así que saltarte
> una hace que las siguientes fallen aunque estén bien escritas.
>
> **2. Al terminar, entrega.** Tu trabajo **no le llega a tu profesor** hasta que
> pulses **Guardar y entregar** — el botón está arriba y también al final del
> cuadernillo. Puedes entregar las veces que quieras: siempre cuenta la última.


### Qué mide este cuadernillo

Cada ejercicio calificable está asociado a una microcompetencia del programa. Estas son las de esta sesión:

| Microcompetencia | Qué significa | Ejercicios |
|---|---|---:|
| **mCC87** | Identifica variables, conceptos y aspectos relevantes de un problema para desarrollar algoritmos que permitan solucionarlo. | 6 |

No hace falta que hagas nada con esto: está aquí para que sepas qué se está midiendo y por qué estos ejercicios y no otros.


```python
# Motor del cuadernillo: barra de progreso, pistas y verificadores.
# No necesitas leer este codigo. Ejecutalo una vez (Shift+Enter) y sigue.
# Codigo fuente legible en el repositorio del curso:
#   notebook/cuadernillos/**/ava_motor.py
#   notebook/cuadernillos/**/contenido.py
import base64, zlib
exec(zlib.decompress(base64.b64decode("eNrlfd1y20iW5j2fIhsOr0kXBcvyT7loSxuyre5yjWzXWq7y9GoUVBJIUpBBgA2AtGQuI2beYV6g9q6io67qYiK6LybCepN5kvnOyUwgAZKS5aqdiIlVd1kkfk5mnt/vnPyR53kv0yLNRHzxtzAKUhEqEae5CKYyVFkSxfQlVLHY/XHXb7X+1/TiZ6FycS7+Qp9evcaX1sbST2svL5QYM+FJlBRSxFIEciKJ/CS7+CkookAyXaehHohGgcq7eCQdxGosqelWmoUqkRmuRnmBSyoPZJwmMpR4ciCzzFBNR5nKU6FElOTRKIlk7ovnUmSqyFIZR2OFfgTRxa9JK0rGKowkunUu8ihBz/IiGqVoQYLYX6ZKoGMqL6Z4KMFAciVkkamZFFIMZRzLDKy4fXtfiiQt6B+Ry1hRLyTY8ot/+7aw985p5IXCYNCNi18kP5kwlzGUQMWhHmQyGGXECNE+5g/9KOzhvbzoq1OVBVEQpf1Xx52u7l0OGfGnWIHYcTDNi3Tsn+bH4rx1LCdRX51N0qzwJ+fHvjiIhIxHF39Dq9QBVYkmVDlkkqlAdUVdFGDLCOTzSCVh2gK/oyEEBpn4Yo/1IRoodBMis4LV+oMvaZahhwkRfA9q+EVilvF0DD6di4uf49ZkqsCATEGSWTSIslg+FuivIPElpIJBhPskYR5kMIXoFFj+PZ5hvZtFM25jmojxxa/hNCbKJIapiotU35Gauyu086qf1m7Z3zwFaYg2jtVIiuMPafb+jsMmP5qcJ4PjLnVEZsFJNEsFGA076lF3TuQ5eP5rAvlO0gTCxWPRmCUj9mIxUgkzkQQBJmK0rgRu346SIINkJcsGw4SK0pMkwJYZNjQNg4UMJhkUPDNj9sVL0ltQBx1WSTFR2TiC4KH3hcRlbWP0CSoD1Y5y/hyk47RsK0mzsYypgZn6yH2ku8xZSSYpRtFIaskcwIrAiw9ROFJFTh3WGkYspLGpHt4nxSYTIgozleVki0TWMDkoptr6WuMpxpjxuFUyTYIITIKAodQxq9jFr+gHjB1MBbt8sZtjhA0FpveYGlQAIxjEMoF3Iy2DEKx/YXYm6GiS+i3P81qtGwIKIL6HkRLfBfQwKaJQhlC5fApmXFOTboiDdJBFNGq0Obn4NY8K6AbxRxYKN2JLOE8TdrxwOeDTmHwbmEdX1Dg9jXK/tfu/f9jvCvq3//rgmdgW3o0t+fWj8KHXxce7m1vf3A+91o97b57vdQX/Kp/bHD6S9yU/tykfBPfueq03r797DWovn+6+4UeCR/e+2RrwI4MHX29tbnqtP715cdAV9G//2f7uGzz+9DWo8uMPth7cva/48eHD4dfDb/hjOFT31CN04sXr/b23u/rJwf1BcB+s7T87OKArYPOTvDiP1U7LlzO5EchTOR+Ql896dydnpA5RKG62+VInf6xvbcRqWPTuO/flx2lc3YZuRdO893By9rgl4MjDMEpGvbtbeOEu3no8ltkoSnqP8H3z8UAG70dZOk3C3s12AH+egtAwTYqNoRxH8XkvP4eRjTemUXdDTiax2tAXuiB960CNUiV+eHGr+yYdQMu7uUzyjRzSHGoaefRR9e7e9x+g2ThK1MaJikYnRe+u/+DB4wDKnvVu3A3of48XFQv89P3cGeqGfvBmG5ai2VD1+cZwayiHD2pvw1JXvp6lp2nz7XB4f3iv9jaH1pXvyzG8xDKBb9SwRgBRNl3d/SglW2pS+Hr4YBhaCkVUzJlzHzSjHj7YfFy2DymneTBlEWkhboDriHekDA0GipJcjXvl+w2GLT9uuLXEmuUnLV/Mo4xDNjCIMJ1bcUP3aoNW3yDSDlcobIpeDuP0Q+8kCuFw6jQzheCTVFQ3N2/edMmShslsg5wronXR/mYTnrZrzaPrSKBDlrHcOjwtAFMRpUkP/rs4Ef69B/A6Mle2I8Aucr7aPFbr/j2QtYwaZVFeia5IJ67cCmCKUu/wRiwnuerZD4+v06Zj5A/ZyN0mRHHSdb+Flzkc6zyIzN1N4pA6KzaAgUZJj3S7SXm+0p2sVOCalm+Wfcxno/lYnm0w/42IjbzltEgfh1E+ieV5bxCnwXs7Suqa4NuGyiRS15LSFnuohpygIs5wA0JemSu8jYrP1g5Z0TVdcus9IAzEafTqzm3xzCJcIZMQ3ZKngAI5gZiAgjcQm4UbgEQM6JADfHQRD7R/RvkDOsYxn7CBgZkBfwAWQHemhO8BayRhAUKNoTplSMO5x4yNGLg5A46DgNAA0Sux9WOLGYfAvUAeFI23Hoh/eEo9P/jxT0wR8VgSeJMJY1Nf3L7T8gG44t6JzNs7wk+nxWRa9D9kiBlAeyyUNJjGMD6B21GCu3MrywSI0EpOP9S89eSOCZMImeKmmFOPhcda6vV0LEbQJf3CV40OPNY+fHWjtsdOEBc1NtBkyNPhkoYBHnszokJ4gN7QPoPe0aHcNKQV2TRH8MJQc9ysbYfvCo/UyvSn21oQwHoJpGxUIohUVqR3VJYBl7EkwWsGPvyFkQ+rywzigWQ0Gi3SMKUMKAexPCIdl7nG/wHGgZTNoOSI8hjFmQ9UYMhZBCcjiZxBDQC8/Vb/xbPXr173X/8DkEmbx3LrCaxRsCVue3cfeEIbov48i9SHp+nZtrcpNsXdh/i/J1hI28QCSmpjYzkbW5WdZNreJ2fezi3TRhBlAfKKAKQeeSI451/Ztve1v+VBCeN427uZe3fK5ycSjjnc9l7e9x+KR/69eMvfEvQfvm/c97+xL5HqUJey9D36dGM4HNpvG3ZI9LAha+5QCEESt+2xDytfoMunaZTY63d2oJGz0c4tKCOLuNWx/Hu5u//fhYEP/PsC/8UPwDz89/LuJjhKFzbMld+TlS7LyNTAsVaryM57QtygFFBnSVB2YEwkrnhXIRRLJ5d6zBUKZM1kBiaHYgvhQoTPjeuc0k3AYBDv+Na3u3/uv3vx/E97bwl7v82QXamzQE0K8YJf2iPT495MMjkay542PshCbKC5iSInHXIixmkSE30HUq/AnRUt/FHGuWq1TI9kXtiPcJEncTSwX6PUftIcC06UvQC7bg2zdCxefH9enKSJb9yiHea3b1/C05mLYGeohqJ/UozjNgWutNPjbpn7bXra3CDeg+s68mRcuaDohOEmgnIT7ULw5Tv45IMgi8AlVa92gBHfHZB7Qawg7wNVDW3Ng8IHoITkbPnclidSjks2foA69CSWYziwNj0JeQvKobkIdEPXdGy/4MKo2sallOevX3a6YkwIL4NwzXOIZviX+k4lKp0d55EaTzLVRV5/QxTkEC9+TgQ8J+IlrEnywCcXP8GuEML20RGCd6Es8InS9SiDs+2J49Pp5Byx38/TaRaovoamx6DJfRXf6dv7coCWMMykAGXzTp+HGQ36yQCMV0CXaZLrNF+XeUjVU5Aixcdt9OOYY2jnmApbeWqibOnoK4/Njedof2tz6x7cN4m3v3cAJ/R8700zu7wqNrP1bYwpHl0SoIkKPz5jjEqg5MuJDmNVgl/nScZtBpGJdcmng+JWguse+dSHk1X0hZy7eBSvMMgLFVRXMvTn0RogbnEd+hLio3KybUT4HGQmCAsEDB2Q8iSMZgLwI8+3vUbr3s5LU2+s12hkUESzlHzYEynSJID+vt/2ipMo9wNIHXrSvsXsv9XxmfI+or1fpKNRrNq36iK51XmMHHSaJeQYkTp47AXEDnkyapYfenJHoqN30NMSVNUASg1NLYxvMVA167PthbJtPAwIGF/SKNCVBmysFH56DJ9NgBTpIXkFlcQyULrcPKMqnd9ikvtSTHNZESCMhBwwAaEeY9KAHABX5EV88QsSP/5IyiMTxfVlfDVBgCmuwK2COhaRW4F/+Qm3Mjh5LvTqKhWXnzIu8cOL5BF8C5MyVbmkKwBXJ7JWdsTXguKR4HxhqbLvW4bxb+2s66ZLvplFLFhVHA4XuvhHrgIqalwlMbFGn5//HvyEM1OZHrtTh6OvyN6pUNyjIK6VAz+v0vEgU83+mpkAjGd5foEr+naKQQ+MHGj/bALSsIqS9D9+z8EElEkIVAaFF4xCrtJOEfApcjALq+kQ+g61p8KoieymnXW9dvtimyH+jSR5fK5aZzoDoj5WYtAhkgJnP0qiot9v5yoedg2PunZE20hFu2UD296B7psFLFAmzxgD/RAJ33B525Cq37SM2rYN1G+XI9ku26w/wK9u1q/1IRUu++e4l6ui3eF7N8gXzWiqCqwIAC1pvqinJwnGUF5Ih4ILR+AGQTPPtC3mC1H9GIJiY8dORIEy4jj7hCaJAVWblwg5JGJOSEmv+Jml17OMXzw8qvXAwju8pxVzBv+ZG23N1DCjubGMMiP7jrE1QISOFjoVuGnw55VGb3zRT6VCxNDMKBAPsCvOJo5eQOEO8Ai1WsNIhAX4eV+8URRPGEPown1GWb5vPQabwtBwD3C4KXgQgmo8gW5UrQquF3BEWKsxvgzDNlPtLCnaV9uiqZ/9ksVGQu3qNdMUXqkYMyRITQ5bP84scvgyCQqyhChps5lxxkAfxe2yE3dqdtPpEBdqlqQQ65zxxRBpWmZi9meIlIwctkm0NBKYcwq3aJakEKtp6kPGgUw+QrN7Yl6zzQUCPmiZvMoRju3yzvZyD+u98ea1BzbsqwvSEOt4S1/w6a+NDnz6e6UUS/xfGngDlTh1Wm9n9U1TcLWZXUXLsE8X6eaQ3uImaDCa0P/euqJxGjHeGOzM7Yif3BnsQMg1hjAb/kcyyCePP/2b/i3mLNhFs5lOpWwrNQz283QpgJkwAfA1Rfry0RgkJwj4xHjJlsXypgUmaeEmfHVr075GW0pT9TudFYZZXvoAlX3nc6p25euuh/QJd8NoPlS3bd73weHNst02+DSEA/lQORb9UH1wH/yZjME5o93LPXTc69M4BZNzPWVKfi0Kr+tlq77r2iDlVIQvV0iY62kVZtRyNXVRp+Y6dSDqKWuEDl5ZrmpS1lJc0l3dDavvtxzu0ixJAz4EU5VNUvo+Sbc9r97fN4qQVsZQGbj0nOfKCwBkGVIJSNF8b5c0JO2uWLxR6yv1j0TS9ux0jfDEV9xsx8fgoonjodE5PLpslbju7cx1162Fka4bJMMedok/dTJz7shCk1nM9fAXy6zS1WXDrHw2ojUtaolBL830NzCunQ0mrEEwHcnfR8LDbapJ2/UBYdphpAkRg1yNQVpXGyrUuXQsZj4CY8G/i1WOTYDFK15C23gJ/7pMxFfLwY7DCJ46MXyAwOVA0bCAZIYRCDrswC3O6n0qRLaH3pPiZGceoAX89thuA7Jbh0Y1Oi0G5/XaKEAq2yFtcYmHhnjoEqc+dehB3MjMDbpm7+UrwlHJWRqpqqkbDR12BFJzdHrBRCuV4edLpWF38r2GkRu/5afkfKZGEZu9Qad13Kav1dXxT1OZcaqSr1iBZXKxMmKINk/qEE6Dj4l4jqazQiVdgHzIjR9BUISK26YPjvumWfZEP1zvrq6/uZDKIu7lVnwg57Z57/Co04xq5k1es3JJjDMwo6x+0g93jyPY02mBj2BBztXDKE22hx6LT7Q378wBK+zoFh2v2zAq5wcBLJ0W2+/8ff7QLovpm5Oz0o5K9si4P5nGOTB4v1PvbVJyYpkHm50mgksIuzl97C11sBG2mylTJcgE5nK39hwzyXc4Q77Y8mbOzy+aHFruIJ4S21d00rQU5WRJoa1612J5VJwY3Vl+fZ1P5OjCjS4DRw4ht9bIc3hLj9IMkiymNk7tLeeGhclRI25UQ4INcImsXUl8Cf/yk47veOaUTbhkp2tu1/ImpSHS1Gwzz6OFZGo0TSiApxN4gIQWcNrp2a5dtEU3EOish9kGoqs5me8NEeIOkaEqw/jibzFCecwoiupkvAhR3yuXcF4n3HE8ouZ7yyq85BfrHvGzQTATY0hU8cV7MkDcsFHmP/75/4p5utBxJKUgYvl2OUimJRspu5k39En7mrydsjXl2+w7S0pdwXB1mxzVJY7mUndDaws8F3nr2Q7qwWsufDustdLJ1vlB70f7BJRgwPf7JqGiIqLMzr0r+rnOKd5b5xTLPi35RT0SP4gV5J02x3Kli4AOsDBMShDlHA96K7u/Bk6Thng7e3E00hNHVuftXFDFUL90B03ShKtWN4oORjkS5gLxT7UrY4QH6PTWclnl/XJZBTD1JhVUygtP2OmW+kWpRY0J8Mr25qF96aizsq31/V7Ri0Yb9pYj7sbAnffXNzKSCcFoGx11BauqXa3v3BkBc4pcov3VXFOhfL3jUdOGqsa7og3MjvwBCc3FL/Ag3lqil0ac9P01w00Vdubl4obFM82RdM4DsCHH8c2LtVp2tcQu7f9Yxr95AC939xdv01DOaIIjSXVXf5yqmFFmrFTGEQKgdEyLdq8mDDu4+JlmLFKRFjwH81FPheoqpJ4J5RJk4i9H48o0axG5vNypx23Fhdzy7lEdrFZY1QG5dXi7FL96y8hKlbUQfqJRLuF4RNEXOc4AUMNEJi4/ecu1k3f+j0/Ts/Yh216XKjP01bTTsf06chM6s0ljBTqIkJYBEfAD1jLTa+ACs1Kbcnqa4MkR63q2OTGROc2G6dyRFkrYLQX0pZBntG66JHes+3JMM9icWweFmKsi+stUgae87A3GfC6O63095jk8U68vidn3cjNBz++A8oRXCvjuGP4r4UnVrTKbojF/CX7xXjOTebo+15xeG6LrCId07D0rly/mM4N03gOSENjRHfKjQo1zSPtSyIPOkAQybUK1J9/5z7N0EqYfkhIBHbaBrf4VukTRuHMEtH2IzqAn6IVu9fD90QL333d0j7hyYBl2tBZ+OEhKNFI7MEXMI8b1vfX4ReOcuee83Gf04vWE9zXBl8Xad9egngcPDeqpOTYMKqJBZTIZqTYH7JoqO88f/X8J66BPo4gD/2FuQAUxLTe1X6NsR9eCgqQYrEma9hfAwD/KmCbntTPjFTJE6lLkh3ZJvDTfadrtdMQfdGZsL1y7J18YrOuB+o2awKZok9lUD0mP4hmtPKCvNPdAixzcqcArRlpKzfi0hlKvHueXQLzfHd79/tCuBuv2eCV1GX6+AN2tR3Y3xHMVIJA+++HiX169fX0AN5iIQaQSvT53evFTrHJenDtG0oLrvAyGyvb+Snr0KkljOm7fZaNDPjIgu/kYTazKNmECz3dKkvug0/ovgp5XwM63NNJczGk4XMu5gtYKJ7wwlq64ZM9bL3xBBZCLvwGzFEoz9wrCkIHeRgk4kWUa8PDUHgAt7zyU5g4t66Gk8r87kG2AUwcaIM5/DkSlRSu853FVCWuqBpJrJmkGJXx/GSS1Ec9s3CQ/5uxLkDHgHfs69uPuNmItb71D00Wl3LhBpbROPdGlABImFCYasGFx1wBOQzXTiQ9VNe/wclt0mNejlhNdUzBKtM09b2J3zer1XmQVHT3NbjwE95OWE6+Brb8HLqU1zw2klxMMZg+tOdDmITpQVq9T3uNfVDAmrJ3pxcp1xmpDgFFhKMFUxrS8jVLCXK5t0nJn6L2dOuLLoCNgPjO8EnNPPAFYV3CpWQZXyp/dtEnR2jfFE02t1SWYsukuS0d11jCj6xArL7o0y4t1AtcJdr9bkKtSBaDstbUGj+yJ5lW8A95TTntd8SHiPd20axRh0BH6UjwqZ/EsZ4jUG96xDF46Qj0XM1sSUKcqmBayWuPteyuXq6yZIFqBKEwvvqK5ZtpRX9CRADIPKLU1rRkl0Y7vn7yb+T95HaMsPi01bawzWsfBRtTxqtl0j023cmhLE2KN1Jm9cslrDhBmAVlP8LJrRmSZWaCa8f6C3Hg0wzVHoX/r1NraNJPLEq/0bnnTiPUqNMHIPZwz9YXvXZoufuZ015VTXU7n/qz3jilai47OVJOg61NhPvsCznwwVdX6fx6FrbtkMZuB5M1LJBqkrbRDPxXrioTe5OKXEBGPgwJeQIZtiClKsxmSmVt6cdtMsm3p+aqrcuwrJ/BcPb18MqsMm4fJkdO+MyEF15yp3zaZ7cxmQ0VsPJfDKBtLOwM1JgU/XV5csTstUkWZn5lIojU63CUfiZwIeZViEvGKlJ7B1lTUOU2dYyyutSrKUSazUkeLCp6V4hYU6oTWp35maYUnj6QpqfAKBHfYl1dT4KyiONbrSt/5z05U8H4A0OSm8zxHQGcebOv4+IUzR7SXVa98Kpo1CrdWsdT/zytKXDrfTvBMjCOtGdDG32sq/eraAm91CLWLRJITOPUFXkZi2d+o2ejcHyIJarPjS69ZRTN0r1ejoBd1S9cvCvz2NBXhLDXsMVnU3JjnlySmsck2nHU7cbQzL0ALv7VlFLw2hwfcufaAv2SS34zZ6kDlGEupm6Hua+vnvpHX1rXoy8na40q69DQwD+B9Uu4WQPYte0+m8c6cGYNm8Pkyerf2xlwN5QxFnlz8YrcpVPEoLU+Vufj1LBpLvdDO/6KFCY28rVRnZG389upcDR65dOy89pVWufMO3iWP/lZmp7RkeRglMu5VC4fziDdAAM2XGyBoq5qmoq+ZvRYuIHU8+xpFi11Fi0nRdAdrAGPNaui6RqvcZiOmgvF5C7TXLzR5YceeDnjV6fLqbdGeO6utO0vmt2x2V3SS9k0tvJ192qWfi4o46X9zVbMvLlkjQzlcLmZ6BbJ08rhybVmPtwJeNkW3Cq8YwEEQa+7sUFmsDLgQccOS5nr8i1XI7KsmT5xzEh7RZuW5VrbagkitfnZN5KpOeFRE8/Ru1Q1ew0MJVWPdC8N5e5pB7JYgLj2PqMXbjsvzj3SNgQGj6rHRU12qXP5NUvz+zd7zvWcv3oi2XhhPlVtQCaOAYrbMIr1n1qCagE714k8zOhKNT7VSZ8CqbGgd7VQo2ydYHdHSSN4IqjP4skLDOLscEZl+mvHBZxB8Qltm6/dZR2n3W8LbUxFDzbxcbSdedQqZrlYCz/HpVfr4KF4VTztUZ+yNkN7QGRE3yMjHski58hJQE+LTX3NiGiWgvK/Wno6Wffq7bxlcZaWgGEanlFAT14cqTwmeZLwTMHY3zsW0pfXhxuY3G1tbPTSCfo6VZjTlbvjCp1nBnQ3M5j3p+77OmhRNi5ptwbl+jvoG397l9xMQ/PdcZRc/mT3GhpjMUl3TFWntEhLN/4nRiD+bzB9o6iPE13N4yAfEgZo0bGR9zOVAUSUSTGSymh/70sI5SI5kjFwoE9/+sLe/v3vAx36UctISiCgKo6n8RG49eKgjHegsn/7Rtfsbc1sEgg9hOE2enk7Z0jBDF0YzcvFsOMRWLm7bwhqfx2ZHhkGQXHXxBiOjCvg+F8jA/HMOm/pUNxZk9rjZO95vycfqkQ4mUTKijBpXCrW8MZIY9GdpJkUGkouEHNP7m3dFu9w+4E/Oe7qCPIqy/v0OncNXsHblU91d3m9LpmTPMNTHuVGzvFczB90oN8Rz51Q35xw+jvtss+ixr7e/nsA5xLJdOoUulQsoo9blunKn5rf8HOvlQB+GWKuMMoerZcLnZmMsn+NHu0q1uZrtlB+DWDoDQfCO8rHr4+gIRbgnrp1zJAgvfq6e010uY7lZGmlOIPC1VlWLwYfevOzX4v/MuVP4zcP7Q7bwfJVQEaftTYvhxiOTTHf8E3UWRiPao9w57N19eGR2C5uz9Fx+uY7MLhCo+EbeXa8Cq/lUFkzp4M61G2BN3DvQZz/UJmJod7ApKh/bNrikrMfTE23NFJOedxY+zDIBCNzNczp3I032zCEsiVPkMHiMr+iDJ2ndDLmH5kGUlWOYsTYuV+BqW4BNSlKvkJZMopUZDsbLZATzqHe07e27+hCqAXm+rMFFW3mokqz3zsy/YRN15r3uTuLw/Mh2tJk3rexOfWuBndBFU+UxhT3emkJFGY0mMzXJ2u5KBESGQoVtk7VYNAxe64xdT7Iw5tT21z6xwnRHYxdWtJop4Fo7roZ8yFeOeCL55KhlGcBduM74YVGU+vB7nSoVKnV/wcFC78rNuLLGxc5P/9aof4FdfFEzTJMzbJlkwIFAd/sEOhvEbXTIfQdEvaE9NVoz9ALLmKPNVMOqa53pSCEtBetDM1NQxnWNODI1TUxpqR7THVz16a+ITHyAEVU5I9ZXWPGnv3f5LV7s9OmvhCIufgY04wNZE4oQmTiGChxTcP5Bu1c9KjNhBIIJnQOnoRTjLB4tHfPAD7khnHawjBLcoQmkiaKZLwMkFZ9cy8gsV9We/urYEXrk4lcbhKStvNsYp0+Fjav12jrOHeiISV24+Jc3T1/v8+mfZqGWPn6FBk/rsJyzeGHSPGZBdSm9bYwRaMClVXaWqQ035mBS9ouMs3kbXJQMMXzTkVZ/7+Dtmx+evf3hzS6dWGOO1UILHrykzAv/j8QJ+rCbnydB+Y1Zqk70cSYGuHsfTqJY2Rff0ZeuvQUdnGb0LN3udN2ZdTAVcqdttsyxrhhf/JQLhP/TtGWPnQBf+iTwNk8IUp58G4p2gjhdjx5WzsSqY/PssT07RmNhVh4+bziFu82m1gosPWlXye3SVBvHcDxEx7Hmon3MrDnuimMzWP5YDe64w5KjAwl0vMAb+XRsHqOFSPqjdm94fE3kYRIWFmsgy+SQZ0LTqpJDyZfBlHZ0drXoOsv3MTrKlWjihPqDX/ATNiK8PtVpixldrzopp+DTjCUvxNnQMk90laB9fHjGDviMHC7A99GxkbVJT3JdRyEQAupRQHiKJcD6yStNOZ2ox0LNO2jiCGZXFI7AvX4/kWPV71PnzRChGlESGQFui777HV+qlzXZsvR+EDmWj9jDlkrWlrHake8yr9KJRNOzKI4AYXPrEkovY8gRg8we8R4B/iSls7gJ++nzlOyhKdSEtl4yT55oArqs1F2zQmaDNKbR4KV+1XjbHVxjRCSHCSICdIZXFVrT6NWmrc0D29s1c2zOBR7qcmvC5WjY8QcZv29znzqfsbw+0e7hGTDZ6sdpqbz7vE+M1i+9gnj1Wnp91Y9C6qwe6dGKDVhXxt4qBu83woPiMyL0Fm8TUAxPcJfyJW8NoeO57s/imBwXy1DApyBpZrjdnJSjnCVKEIucnbcTPp3D9bw8N2fkU6uc8cP1kZ8Q3NWr9tZLalkqTKlz1JS2obZijtcaO88uyOS8vSzlehCoOmOILitANZ/f9rgIxblg07UcH1IuT9Tot/Uuwrts/sUrXZauIbOZecwHdySkXK4dkL9b2tT8/0i/juem5cUxKVgxXatjFRTTawPTuWHc4mr1uiGelcVCdrJ2uGWYZRSglbjHkzNgLyIyl/l8Ckz1VdjX8Qd4fIWauHatTXpJ1e0PGlpBYLegA9SmRUmFIoNLp+EcPk9wv5vQXIHhNchMY22eRLsqJJXghcB7FXT0azSBdjqlKXNZwRkuAJrX7VFZtOuDy0UcxNKP9uitcv1UwrU85ZN2JLykpTwBgUAtJ6gA8OWaB402NICEofIfItC3yrP9k6i5RIIs+rg8wvBY2D9toDna0iV1G+LpZDIzj2wI8kF5kquuozgd4D7ftn/UgPE2/C3VnUz0NwXy27edA8d4yZ8hefu2Ts0ntKBDIz99vCrDa135pJ3b5elgMzpBf1UuUWJ125KCic1SmsQTZn0bLXEwfKstu3KO4a1OXmVh8cEjDDLzyGxzo9P12aipa6YEZFXIdoXqezrfKmvEdRClN1OsAVFUvmEQxRsDbFprXlnavfdZ+a2RKJ0EoBv6Q7Yop/McJbRKzewqIr9hS0v2mLupW55SG3GNnFNE8y2Ytdq0bcaEKNUfRlSrHys7WvsMZQJWXX3+t+c40v3y7EnO+RBDZawPX0vtcZzUnxEdCUD537d0jJ2ey7RngjnUYBz6z3I0zZRrNnZQvPjFnIRXwXAGmCUSdIjyMhfWFTrhrLbXxihaNWPXGOmhYcLR4ZZZpcTLy9Pcp9Nh/SgnprXNQ04BamnxYtkQhJBOVGLf4cMsUjo8fNtWC/1MybC9tKrx9YE+eXU1WaPU9ck4wvCJfaZaIUV/hAdaw9kQa1nKJ4RSRSXv1XwTFbYdauVxP9PEOig+b1DrifZVV/snh6D2VF199qniv/sxjtiNmvLAGpflOiyHHM9LNLyW9Vh2yrqyhdXlkIraZzqzhiNbqVutVYnYfFHWC/nJy7SHY1mY1uEFPjAL2vr9TmflQg4XKaS0XlRXLGBK5Kueq6FbuaiudtZsknAH4eeqwHc5jQum7ZP76HJPl9T34JyKNitUeJkrVb2XVczUWd3nXBV/xn9eyfzRmlwfaRnr6lQPgtLfG4Lmc+TsiYo0u1ORO+EJKZJ9mFY7TvTRmf61PP4rM8lE9QllzhSt1NlJkdiVuTOZTZdf63vo/o0bXzwnttRdvvvHjPjxBjn3QNVVS2D1hJleztlx50NcEVjo9hn5d1cfpp/zhrgKy303NX/mq+z4RIWROV257sxLWEMLWviUZA4QFz8Rwi5nbWZ2van5ABWyn3LWIt53xYkMn/rY0DPo2Mw50uVq3TNcIet5mcK7qfYgDc+3D48QBM8nqs+lUpVv2wWumjqfIFgrTJTrlQ/ddg5NSn/UWusBVj3eqZUz2kvWXxYelooMZMCNOkNrVWGifJDylAZzlh8wHJytWjWrB+7zic9h+ypVqohaher4xPCaijaFUZ4r1hTIfwK/GlG5")).decode("utf-8"))
exec(zlib.decompress(base64.b64decode("eNrtWltvG8cVfuevGKyBinKWNKmbHepSKInSurBsR4qDtkEgDXeH1CjLnc1eFNOCgPyI/IE8BkUeijwUSB8KmP+kv6TfOTN7oyjHaeOiD5UvImdnzn3O+ebMep73oYlzFevQiCQ1iTYiVCKS4vTo+PDpodgZidd/+aDIApmKuTBpqGKZvv57v9P5pFh8L1SGv7np9OqfzuFXxeIHcaWvlMhMZETEJDFNikzNZCxHGMrEV4UOsDxUkQhkpOJczjT+Nz5xjzSmd0I9Li5lKEVgYpEVmVj8EIe0yocsWDdTIeROibxJVCoDbWKQ/Oc334K6YuFA3KTy1eLHuJOoUE4XP051YCzXQoYqjXUUGajzIhbR4oeZzhWvBUdwAj1JNEKdqcVfzUjcv49vOs4X36dJqjAXrJNMFaEJFj+GelpSTjMjYtOBturSaZPdvy+653byWaGzc2sdMDKBElcy1XIcQeSsUBFmr/fFYQY7siIxm69jzUcCtXlmpGuQ6rGCXfAlUsoH4Rlcqsn0kcixMtWwLcxEkmFWR12qoMjlrshNyF6CG/REByQG2e75PL8wcV+cQksYXaiXCSwUaDgKfrJStazYSWQqWeBYhhqG0SqF88hLUlzADeAeKRmHrDU70OQmhfH/JAroFapAZxquErWrECwUj4gXMMohuw1BGyo2eirdVNSBmQpwI/HgKmsT2D0D3yBg0nMxLlIKK4qh+/dnMBAmzFQ6JeNRUH6ZmTQXmY47loBOwb4vnhA7/Fv8FOV6hg8ZdEgV+VpfGQ5JBK4bIE5sZXEBf8UUPTJEmD2HLVI2EamkZ0mkrDEzlS5+wHIzg4csoWlEJiEyRYbdBwexaUuqVwgqs9thh5KJSUpBEYPHSaFCxZuEZI7t9lt8xzurNAQRgwDQVcIDH6lEYV4caJmNxFiPI21yFUi7MA5likDQsZjJPIlMHulxv+N5XmcCicVjFyrYJkkk546s+P2nx0+EzMQZfdjxRfmYhtznnU7n7HTn7LPHz54cfXoo9oV3b0tuSvnQ4/HDP794woMb8uGjcMcOfnZ08tERjw4CuTkI7OjJsz8848FwsDneHLv1xx8cnvAoAmo4GNjR3508PuXB7Y3t4Zaygx88K6mGE7WpHtnRj18cPf2Uh7N5hkDrFdrvyQRu69kBf+1UTY0SLx6v+SdmjHj2MxlnPfhTTyyN42dPWbJC92bY7FkiA+WffnyMz70TNS0imfrHKo6Mj1SMnCAzv5rndTqdUE3EWbZzliDvyLR7kc+i9VFH4KeyYtea2D5bb6zhDdZ105GwijQW08iMZZR11/tTlXc9eSW95hIEvcRG7ubtVd0sp7F+qsAxUF3vN54vvN/IWbLrNUb3eDTKMciry596xgHPmObNZWveGg1+VRgaJmnuif1f8wf0PmwWmV+dPlkPyePVmY6TIi8tDtvC8w1H8Kie8AOdiaeoV6PKTtbSTaNjWp+odqs5HmgxC1hx26+HjyJkiYLyYFUKNrYpRRexOHcynffF639w2c51YvCfilWZVqJITeVvvZri5x7CjVw1iYzkDwgA+uUyhU3Lrs4SQ+l9Ua/eaIhWsceSKxQ34ILTx0fHz0+OkLpfUsGXRcx0lnWQIl78hNRM6derCT5HyVcosBdyXtbqK5XmlAG5sIAhszPinIWnz1QSqFxQkQtoy/WdrmXos/NCjayKov9u/VdyWXJh0zeVpc4fiQdi45w8aZPsLR+hppuUzIBSu4UiLYPcVH7zUcfUDFDltgOb7ho2xHiCqsMCUpE4f9DwGtXnmSSM4vxXeQ51JShod6H4yKarWBzU8EqNTJxv9QfnPtWx8y1EJAKXCmBqcaLkKkizafpKH6EgpUUA88p366Waz/JW400V6AAZGxsqhJ1CKuAiSBWDMxkUM4QYkGnbWR9B0dQhRLsaXg0I314x5iPnHLo4raYQMKKiIK7UqwY1cmOWIF4ak9m50HhayOgu755qRnqQlBaSPD6JniodA+7KJYmAIOIiBgpwKhGwk629+Jm0W6staNNxv34mf+IQNSGKOw4I7yS/I8mcMeNuKHNCfKnKJKyU7lOE+di5eRGZfc9zcQl89BELyLbhlU7QUk64YCwvjc1LpIbMAKQJj9H6o5dYosrtDS+rNLWnnag8SrkzDoYcxRGfDlKNrIlPDC4pmzI9PPnw6OQZ4UKoQtvNms+CfwS6mC2+y0S6+C6h4yBvbMUTc2iKNTBsv9SMfwcqCgFXcCJiq2MXfv6Fj3/8cAJZtY+TTUQfgPfjAlLJXFnzrddbEjkCEBqru6VFadvGJuetiwALhYY/K3vX0GJi6DSBzV+BPdr2lpyKENjevQn/1CEbGJJnv35ye8lgTH8aS1jNPnAfEme3hWsma3t5KLJ8Hql9b0yHk3Q0TF7STtChuK6Q5c1uIsNQx9PR+3g63Ele7q4tEaJi2MMRbBqPAk6Ju9Au700AWaL56LqEkjd2ONOv1Gi4tYLOWAZfTlNTxOHomu1zs8s6j6751413cN3Eeeyf9Zu9B3l4sFZb1rn055Uu9dos9bqtx7KEpMBouIEFDa2sjMtTr+/0bInkSR/dEh8nWIRqSkWhC1FRzUpZmfHOYCCGm/3tkr3F+JWRGueRmyXJrUg4bU513APWz81sBA8sGdSmAbIoOB+skdx2yIWXQ8YNSN9pGLYhrWU0Gg4g6QBMSr1u1poLcj6vtyKwB1UimWRqVH7wDtpr0oNrz+tfGh13bXBzAKQHrScuBMpHD5hTm9CybWunrrAqe2vXmS83CYVMWzBk9om2x17JidGmN5fbGCzY3FY+2RsfwCJm78H4oG/t7az+birP0V09p3fBi8oFtIQ5Eq1e2cbKZZHl3KNZar04DNWn2hgYaitc6lCGrtNzjx5rRu1Uf6gaTRY/plTRR9w8Kp9e0S+b+QFxZDQ1qc5nhLRRq7YHg/5gMAA14iCrdpttMBiCNjxvY+AzG+pHzHRGrRJO+P3qiBnpWMmorKJj6i6GxtWDRGaGCsmgriFUPFIZT1UXJzhXPBrVw654b18M6+Q1sSw/119Q3XAcRq3N7MAfKhQTaALC3rAcrI/FYx1Ti+4NQqMMvvoKmERRdYEN8LeWV/SceF9faGxXzBR7+zT3jXpQsHFdpPnv0fR1Rsa3FeWZP6ssz2oqvIrG3moSJMK+k+i9hoyU1NozrQXszFLtOyzrIik9A0tEHRJRF0d0GZv9IUKtxlHcgaaYf3707IRgkrFn2xJW2a5gSIDVHUSsNSkOGVnNFt9Tm7NfIhfWGFLS8q4NLst43WZnZwLMsMOkiFj1c48TlO0JWvClELMsoQVqNthtynfKn0XuvLJ6I7Tnjt3cOwLQbhMkTUZfXe+DxU/Wko6xV/H0Rdkn85eqmjfTKYFUQ1iTf4cqs41jpNs40Ik23vrSqiYrJ1vFa2x5cafuFjM4PFdVk59Z4stM58hX9uyBk866xZBwZpqQF8qatFTFVuCQRyUOqYv91h1FyeKh3WYFcgXUFfeVSA6VODazcaoc6Oi8ERY9WgGLUj29yO9Cd016DaS37TTqfa1oNSn2H6vBzvplWtgqv/mzVf6XC4NIsqIw2misR3xYe/tlnmVuvu1gxzb47T7gIwqipdxlD6o9hCRXfmQUNnhLDPbzQdTCiw6ykbkGYmsZ3rh7tDphNI5oK7MZ9uC1zT8j/6ZqhrV8ZeHOLwaEu286q7Slvrab8GYlAEza8O+WkWbyZe9rHeYXo50t4Ngm+nu0Av25RCIu9CtD4O6aPTrqDyY31rmE9FBaYiTvPGWM2Bd/IiRERuTja5PiFJk9rHONL2ZybtNNqCcq5euOkcNEyEBRRD0v6m1ydaiSd4tmTGIgWS0tsfWG74BKJQqS8kohvhTAaXKwVjfBferV9L1Wr95dUpm4dT526b4NuVzdsvP8X4CVyHJUnHVre13WqzTV9iZwWW/X9rYcLbDSwhKXhCNKWAGeX4xunaRKWS6rJyVm81srMaX51a+mNYHFKjM10Zu98/s3LRvTcGWRVcaOG2ZasihAD2YCPvz7ljwoP5JzVliynOi35lV2s1+rh29rtwqbta5cS3y20YRnz3gGxf7x49Pjw0avK3K3pg5/ZTVAa7F0Ha9TRXedjUyIvZeqK+py/vObb5vgiu74q+vm1n52bOkB0+Qzie2dM9awXU9uVKsrHfLt65xvWWNuwBWo0XR0ctfO7ujS7HrhUIQcJqNV8JFhf2+4XoG4gPa1g3D1Di9JNKYhREuk54K1PamB8U6r62zPMfDrDtgKnGbpeY5LjQNvT6WrbxV210l7azOsKjWrIN3/wdlbgLO3ErZVZ2MU+QYCul0tVmKi2OGhJSj0X4Y2NgOkJVypwQpvUOorqLi9m5cSIE3jrcpHAJTO/2GU8+bmYdnmqoDPzqDVpa06iBUIoqaNnqSM9/YCE6qDch/uPeCv1AdCNqC8mesrQBZKZE2KLqnZ5FkmW4uHKEUZ3x1GOaI1M7KvvBAjereF7d+k6F4NocRoW23Ooe6Gkd6PIfhqsSrTk9NY50XIvrUXs0vIKTTjiF9kycqyU1YH19RDmSiAi5NIB5KvSRzEY3Xt+z9+i+LcvtuUBvatJouy1t9F5+85GSOU7+ReKbG0u3U9/VSmlypnN43h9St6P49em5POUlJkXKXYyklqwiJwp3V+g+9LlcYqcjX1Q9pP7rbIXfnjezGXfnkZ5brTc75Bs/VOPDX2jT1Vv6fVcaiFOoPMyQAGU88jo3JbxPReHR/JKBhtkzLkPiK9qSRf8RmnLMqt1iXziUT97la73L5dJmsl63JXNq5CGMmnPZwIQtpH3eHmdqimfjOa7m2ojc2HD/3r8tWjm3W3r+mqqEzoRKDIuCtf3ekMKStubLS2udvlox1OmDS9mXBuy28rCU7Wu5HKEdM9egGIiPcH76uZrVDYD3GG5D8bFUmCsJdIcQ2ahlbk81H/0YZ3cGpdviNe/028iBnS7Czn0Tuk2NhZqmcPUc+cOhuszg5ps/xa6FtSH25R+twlh/QuLP1hf3u3Ev79LWREfmPSp2IAyrbdZVtN89YbgP29cdriCKPFJcvaHN7BzrZICn5lEOZ4f1v88Tl90HGGko98+PovnxQg6Gi//vveAyJ0sOJK4TDiN+TcG4llEGdiCMk2KIZTMgU2FuU1W6MzevMt18ydMSuCnhvjINfr2Yv3FJhrXmeCXo+Asom5cFbvCxLJSRFzxaT7VHpzU9I5le4f7BbH1v2k7N4Xs7h+k8beEnOXvrn5cGBVlExyXvNUztQRXTLzlQEliucnj4+PTg7tFWh5HUEgvM5LghIBSibdK89t45WxNL0H0HHaUU/GZdB/AVZhxz4=")).decode("utf-8"))

ava = Motor(titulo='Buscar y ordenar', meta_xp=95, insignia='Quien ordena')
quiz, ordenar, comprobar, pista = ava.quiz, ava.ordenar, ava.comprobar, ava.pista

# Pistas de los ejercicios. Van comprimidas para no destriparlas:
# se piden una a una con pista("E1"), o con el boton Pista.
import json as _json
for _c, _p in _json.loads(zlib.decompress(base64.b64decode("eNplVk1vGzcQ/SsDX2IjihI7Tg5CezAMpUgR2EBsH4og8I64I5sFl1TIXbVy0f/eN8NdSUFPknb58eZ9zOifk+X5yYK+nSyLy34lFJiCLz3TjuLQScZPLuS4+KBfpNvIC8c20SZlul5+vZ3TdYokJdnCPkuhTfa6s1DhIJFKwpv5yYxOrgffMvY67Gh6yU5yahYkgXxsvRO6oJhwlj4p8jTgHhw/SASedzM6n2HBDm8iA2abYitAJvXsxg0c+1QaEmrqcY9D6H2XGj20KMhAnS9dqpWlGc6Swnpfn1otlJ4yO88LeuYdfcB66UTPpA2W6/Z64IQWGy/nJ99x+fLCSPyyZ8CDpUIttvz17IMsAIl+pXcNAaPoC8LBmfOMGltAnn4hkHXaMu47a2jDoF458Fgjdbun13TeYDO2prqCtxxfOM/pzlMvtObQWyGVXJ4paOddSMoBHnY+MsqPjitpWfohx8aUUlTYusZbr2wZYez6ZKe3sh0kbFWcSJVrfHOD1VFmunPUnRjQfD2j6PMhesf0YxDaDIJLcJnPdv8y0AiB3qCyLdOnh+XXKyPOYM+IA60BOhiIoL6LRqByMDvGW6RWAXhHWBQ0KCk+QtSVZDggqwVxAYRCbVW/96bfjTgpvlflIPoWVfAq4DaF3nHGh3GCm1+SPVTfbD2PdcHWvUTcgTWq4wKVsxYAMMC48QmOS33mQ1EjCVpvJy0WFLV0cANKsJSc+pcfg5cMf77Gsizumc/o7Vu6aOb0UEzs1m99qZJBYtZrJ4tC9YqtAmvF+Q4kTXG0JLdSHOdeRewG2YrCD77z8NPDzRUIL1XMDrxw0Pv2aE+b+mm+TDT+eoNfZ5NerfzJ5pn6sjmyZEUGHlvOqLs1NFUw9I60ynKI7JTFMRgejSjLKN6liXedOnvnGFztl7ua+7vPvz18Xt7cL+e0RKsq/in69eTKPaDeg6VqVJyUwWPaB3ZUbKZlWki/WRy/N3vjtYfuOYaL45OcHlKtxCDavZZK8jeixQpSrHNONI92MyhRlQRWGECXjbXoa1eLBW3VQv+LJI5R82ER+Fe/ZdHs7kavakXbhDhn+sShoEGh1Gd2UwupncfCYk34KKj3eRCLKu8pWQ8atX1oR1k+1MHSWTO0eaEXZMOfMhIsM8ulxgZkR+im3d5csB9DYPaJaQUC0AqwpKqlO9qUF9bWNaubnJDUTj2MQqDblKxpKZUB0+33h7v7W7qCD+4MbuPXje2o/QSeBbV13lgAjWQzPw9RcRtLINFr0hx3K3+4ySzUytSfQxAgf2atokF/Q/vWYCysf6pHN1yqVZGsHRDuzI02HkNy2qQODGP+0v7+vUqjgmw65ZDqLL48GlzKz3v49EK9ihHy8bgobQOm1EdT6i6Zi2y2THRi8mnXw0k7hPmn+Z7MTfoQzwAEculINPN0kAozb/L0YyuguVkNxcL+uEJR6K0NDm10x6PZAWQ0WsHGc7ALTe9Q/w5MOhrbX8YOPR1YRtGicorno1g3t4c4xENTNm9FAbmj1D93jHFA7HTdCLSuGEfJfqiPUOqfApwZeDvm1sABcTslGhMQKBcHWx86jdafaSKA/ji+9rCqJqczJaxGa69ra8fQLNiM4NqCjBsbjRAGadf7gOFVLfoVNP/3P+YAbJ4=")).decode("utf-8")).items():
    ava.registrar_pistas(_c, _p)

ava.barra()
esconder_entrada()

```

```python
portada()

```

## Al terminar este cuadernillo vas a poder…

- Crear una **lista**, leer cualquier elemento por su **índice** y cambiarlo.
- Recorrer una lista entera, por índice y de forma directa.
- Escribir una **búsqueda lineal** y una **búsqueda binaria**, y decir qué
  condición hay que cumplir para poder usar la segunda.
- Escribir el ordenamiento de **burbuja**, y leer el de **selección** para ver
  que son la misma familia.
- Explicar, con números en la mano, por qué un algoritmo puede ser mil veces más
  caro que otro dando el mismo resultado.

**Un aviso sobre el pseudocódigo:** el intérprete del curso maneja variables
sueltas, no listas. Así que esta semana el pseudocódigo se lee y se escribe en
papel, pero no lo ejecutes en el motor — no te va a funcionar y no es culpa
tuya. Todo lo que se califica va en Python.

Este cuadernillo tiene **65 puntos** y **95 XP**. La insignia se llama
«Quien ordena».


---

## 1. Calentamiento · 4 min

Tres de la semana pasada.


```python
quiz_input()

```

```python
quiz_division()

```

```python
quiz_estructura()

```

---

## 2. Un nombre entre mil · 4 min

Tienes la lista de los 30.000 estudiantes de la UIS, ordenada alfabéticamente,
y buscas uno.

Si la miras uno por uno desde el principio, en el peor caso das 30.000 pasos.
Si haces lo que hace cualquiera con un directorio de papel —abrir por la mitad y
descartar la mitad que sobra— das **quince**.

No es una mejora del doble. Es de dos mil veces. Ejecuta y míralo:


```python
comparar_busquedas(1000)

```

Esa diferencia es de lo que trata la segunda mitad del cuadernillo. Pero antes
hay que aprender a guardar muchos datos en una variable.


---

## 3. Concepto en corto · 14 min

Listas y cadenas. Son lo mismo por dentro: una fila de casillas numeradas.


### 3A. La lista y sus índices

Una **lista** guarda varios valores en orden, entre corchetes:

```python
notas = [4.0, 3.5, 2.8, 4.8]
```

Cada casilla tiene un número, su **índice**. Y aquí está el detalle que cuesta:
**se empieza a contar en cero**. Ejecuta y míralo dibujado:


```python
notas = [4.0, 3.5, 2.8, 4.8]
ver_lista(notas, resaltar=0, titulo="notas[0] es la PRIMERA, no la segunda")

```

```python
notas[0]    # 4.0  — la primera
notas[3]    # 4.8  — la cuarta y última
notas[4]    # IndexError: no existe
len(notas)  # 4    — cuántas hay
```

> **El error por uno.** Una lista de 4 elementos tiene índices 0, 1, 2 y 3. El
> índice 4 **no existe** y pedirlo revienta el programa con `IndexError`. El
> último índice siempre es `len(lista) - 1`.

### 3B. Las listas se pueden cambiar

Una lista es **mutable**: puedes cambiarle un elemento después de crearla.


```python
notas = [4.0, 3.5, 2.8, 4.8]
notas[2] = 3.9                    # el 2.8 se convierte en 3.9
print("Ahora la lista es:", notas)

```

Esa mutabilidad es justo lo que hace posible ordenar: los algoritmos de esta
semana no construyen una lista nueva, **mueven los elementos de la que ya
tienes**.

De todo lo que se le puede pedir a una lista, esta semana solo necesitas una
cosa, y ya la has visto: `len(lista)`, que dice cuántos elementos tiene.

### 3C. Recorrer

Con `for` y `range` ya sabes. Pero como la lista sabe cuántos elementos tiene,
se puede recorrer **directamente**, que es más corto y más difícil de romper:

```python
# Con indices — te sirve cuando necesitas saber la POSICIÓN
for i in range(len(notas)):
    print(i, notas[i])

# Directo — te sirve cuando solo te importa el VALOR
for nota in notas:
    print(nota)
```


```python
notas = [4.0, 3.5, 2.8]

print("Con indices:")
for i in range(len(notas)):
    print("  posicion", i, "->", notas[i])

print("Directo:")
for nota in notas:
    print("  ", nota)

```

---

## 4. Laboratorio · 26 min

Las dos búsquedas y los cuatro ordenamientos. Recuerda: el pseudocódigo de esta
sección es para leerlo, no para ejecutarlo — el motor del curso no maneja listas.


### 4A. Búsqueda lineal

La obvia: mirar uno por uno desde el principio hasta encontrarlo.

**Pseudocódigo** (para leer, no para ejecutar)

```
Para cada posición i de la lista
    Si lista[i] = buscado Entonces
        Devolver i
    FinSi
FinPara
Devolver -1     // no está
```

**Python**


```python
def buscar_lineal(datos, buscado):
    for i in range(len(datos)):
        if datos[i] == buscado:
            return i
    return -1

nombres = ["Ana", "Bruno", "Carlos", "Diana"]
print("Carlos esta en la posicion", buscar_lineal(nombres, "Carlos"))
print("Zoe esta en la posicion", buscar_lineal(nombres, "Zoe"), "(-1 = no esta)")

```

Funciona **siempre**, esté la lista ordenada o no. Ese es su valor. Su precio:
en el peor caso mira todos los elementos.

### 4B. Búsqueda binaria, y su precondición

La lista se parte por la mitad. Si el del medio es mayor que lo que buscas,
descartas toda la mitad de arriba de un golpe. Y repites.

> **La precondición.** La búsqueda binaria **solo funciona si la lista está
> ordenada**. Sobre una lista desordenada no da error: da una respuesta
> **equivocada**, y eso es peor. Es un error de lógica de manual, de los que no
> avisa nadie.


```python
def buscar_binaria(datos, buscado):
    izquierda = 0
    derecha = len(datos) - 1
    while izquierda <= derecha:
        medio = (izquierda + derecha) // 2
        if datos[medio] == buscado:
            return medio
        if datos[medio] < buscado:
            izquierda = medio + 1      # descarto la mitad de abajo
        else:
            derecha = medio - 1        # descarto la mitad de arriba
    return -1

ordenada = [10, 20, 30, 40, 50, 60, 70]
print("El 60 esta en la posicion", buscar_binaria(ordenada, 60))

desordenada = [50, 10, 70, 20]
print("Sobre una lista DESORDENADA:", buscar_binaria(desordenada, 70), "<- deberia ser 2")

```

Fíjate en la última línea: no dio error, dio una respuesta falsa. Por eso la
precondición hay que comprobarla, no suponerla.

### 4C. Ordenamiento por selección

La idea es la de ordenar cartas en la mano: busca la más pequeña de todas y
ponla primera; busca la más pequeña de las que quedan y ponla segunda; y así.


```python
def ordenar_seleccion(datos):
    datos = list(datos)                 # copia, para no estropear la original
    for i in range(len(datos)):
        menor = i
        for j in range(i + 1, len(datos)):
            if datos[j] < datos[menor]:
                menor = j
        datos[i], datos[menor] = datos[menor], datos[i]     # intercambio
        print("  paso", i + 1, "->", datos)
    return datos

print("Ordenando [5, 2, 9, 1]:")
ordenar_seleccion([5, 2, 9, 1])

```

Esa línea `datos[i], datos[menor] = datos[menor], datos[i]` intercambia dos
elementos. En otros lenguajes hacen falta tres líneas y una variable auxiliar;
Python lo hace en una.

### 4D. Ordenamiento de burbuja

Compara cada pareja de vecinos y los intercambia si están al revés. Repite hasta
que no haya nada que intercambiar. Los valores grandes «suben» hasta el final,
como burbujas — de ahí el nombre.


```python
def ordenar_burbuja(datos):
    datos = list(datos)
    n = len(datos)
    for i in range(n):
        for j in range(n - i - 1):
            if datos[j] > datos[j + 1]:
                datos[j], datos[j + 1] = datos[j + 1], datos[j]
    return datos

print(ordenar_burbuja([5, 2, 9, 1]))

```

### 4E. Merge sort y quicksort: los dos que no vas a escribir todavía

El temario nombra otros dos, y conviene que sepas qué son aunque no los
implementes esta semana:

- **Merge sort** parte la lista en mitades, ordena cada mitad y luego las
  **mezcla** en orden.
- **Quicksort** escoge un elemento como referencia, deja a la izquierda los
  menores y a la derecha los mayores, y repite en cada lado.

Los dos usan una técnica que se llama **recursión** —una función que se llama a
sí misma— y eso no se ha visto todavía. Escribirlos ahora sería usar algo que
nadie te ha explicado, así que no se te va a pedir.

Lo que sí puedes hacer es **medirlos**, y ahí está lo importante:


```python
comparar_ordenamientos(200)

```

### 4F. Por qué esto no es un detalle técnico

Selección y burbuja tienen un problema: si doblas los datos, el trabajo se
multiplica por **cuatro**. Con 200 elementos son unas 20.000 comparaciones; con
2.000, dos millones.

Merge sort y quicksort crecen mucho más despacio. Por eso son los que están por
dentro del `sorted()` de Python y de prácticamente todo sistema real.

Esa diferencia se paga en tres monedas: **tiempo** de espera del usuario,
**electricidad** del servidor y **dinero** de la factura. Un algoritmo mal
elegido en un sistema que procesa millones de registros al día no es un error de
estilo: es un coste económico y ambiental medible. Por eso «que funcione» no es
suficiente, y esta es la semana en que eso deja de ser una frase y pasa a ser un
número que puedes calcular.


---

## 5. Seis ejercicios · 21 min

**65 puntos**, y el cuadernillo más exigente de los seis. Tómate el tiempo.


### Ejercicio 1 — Índices

*Dificultad ★☆☆☆ · 5 puntos*

Con esta lista:

```python
dias = ["lunes", "martes", "miercoles", "jueves", "viernes"]
```

Completa el diccionario. Los cuatro primeros son textos; el quinto es un número.

| Llave | Qué vale |
|---|---|
| `primero` | `dias[0]` |
| `tercero` | `dias[2]` |
| `ultimo` | el último, escrito con su índice |
| `cuantos` | cuántos elementos tiene la lista |
| `indice_ultimo` | el índice del último elemento, como número |

> ¿Atascado? Ejecuta `pista("E1")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_1, points: 0, grade: false, solution: true -->
```python
INDICES = {
    "primero": ...,
    "tercero": ...,
    "ultimo": ...,
    "cuantos": ...,
    "indice_ultimo": ...,
}
### INICIO SOLUCION
INDICES = {
    "primero": "lunes",
    "tercero": "miercoles",
    "ultimo": "viernes",
    "cuantos": 5,
    "indice_ultimo": 4,
}
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_1, points: 5, grade: true, solution: false -->
```python
assert isinstance(INDICES, dict), "INDICES debe seguir siendo un diccionario"
assert set(INDICES) == {"primero", "tercero", "ultimo", "cuantos", "indice_ultimo"}, \
    "No cambies las cinco llaves"
assert isinstance(INDICES["cuantos"], int), "cuantos es un numero, sin comillas"
assert isinstance(INDICES["indice_ultimo"], int), "indice_ultimo es un numero"
# Se corrige aqui mismo, contra huellas: el alumno sabe al instante cual
# fallo y la respuesta no esta escrita en ninguna parte del cuadernillo.
revisar("ejercicio_1", INDICES, {
    "primero": ("9c58237a99158ea4",
          "el indice 0 no es el primero por casualidad: cuenta desde ahi"),
    "tercero": ("5d9c162e61f1ee26",
          "si el 0 es el primero, ¿que indice le toca al tercero?"),
    "ultimo": ("96fbf218dfc2bf94",
          "el ultimo esta al final de la lista: mirala entera"),
    "cuantos": ("807df0d0cd326132",
          "`cuantos` es el total de elementos, no el ultimo indice"),
    "indice_ultimo": ("659c2a0ff2b2e888",
          "ojo: `len(lista)` y el indice del ultimo NO son el mismo numero"),
})
### INICIO PRUEBAS OCULTAS
assert INDICES["primero"] == "lunes", "El indice 0 es el PRIMERO"
assert INDICES["tercero"] == "miercoles", "El indice 2 es el tercero, porque se cuenta desde 0"
assert INDICES["ultimo"] == "viernes"
assert INDICES["cuantos"] == 5
assert INDICES["indice_ultimo"] == 4, "Cinco elementos, indices 0 a 4"
### FIN PRUEBAS OCULTAS
print("Ejercicio 1 verificado.")

```

### Ejercicio 2 — Buscar sin ordenar, y sin `for`

*Dificultad ★★☆☆ · 10 puntos*

Escribe `posicion_de(datos, buscado)`: la búsqueda lineal.

Devuelve el **índice** donde está `buscado`, o `-1` si no está. Si aparece más
de una vez, devuelve el de la **primera** aparición.

`posicion_de(["Ana", "Bruno", "Ana"], "Ana")` es `0`, no `2`.

> **La restricción: resuélvelo sin usar `for`.** Ni ciclo `for` ni
> por-comprensión (`[... for ... in ...]`). Usa un `while`, y la celda de
> prueba lo comprueba: si se cuela un `for`, te lo dice.

No es un capricho. Con `for` esto lo escribes de memoria y no piensas nada; con
`while` tienes que poner tú las tres piezas —dónde arranca el índice, hasta
dónde sigue, y quién lo mueve— que es exactamente lo que vas a necesitar en el
Ejercicio 3, donde el índice **no** avanza de uno en uno y el `for` ya no te
sirve.

> ¿Atascado? Ejecuta `pista("E2")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_2, points: 0, grade: false, solution: true -->
```python
def posicion_de(datos, buscado):
    ...
### INICIO SOLUCION
def posicion_de(datos, buscado):
    i = 0
    while i < len(datos):
        if datos[i] == buscado:
            return i
        i = i + 1
    return -1
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_2, points: 10, grade: true, solution: false -->
```python
assert callable(posicion_de), "posicion_de debe ser una funcion"
sin_usar(posicion_de, "for")
assert posicion_de(["Ana", "Bruno", "Carlos"], "Bruno") == 1
assert posicion_de(["Ana", "Bruno"], "Zoe") == -1, "Si no esta, devuelve -1"
assert posicion_de(["Ana", "Bruno", "Ana"], "Ana") == 0, "La PRIMERA aparicion"
print("Las tres busquedas dan la posicion correcta, y sin un solo for.")
### INICIO PRUEBAS OCULTAS
sin_usar(posicion_de, "for")
assert posicion_de([], "Ana") == -1, "En una lista vacia no esta nada"
assert posicion_de([10, 20, 30], 30) == 2
assert posicion_de([10, 20, 30], 10) == 0
assert isinstance(posicion_de([1, 2], 2), int)
### FIN PRUEBAS OCULTAS

```

### Ejercicio 3 — Buscar por la mitad

*Dificultad ★★★☆ · 10 puntos*

Escribe `busqueda_binaria(datos, buscado)` sobre una lista **ya ordenada**.

Devuelve el índice donde está, o `-1` si no está.

La idea: mira el elemento del medio. Si es el que buscas, listo. Si es menor,
lo que buscas está en la mitad de arriba; si es mayor, en la de abajo. Descarta
la otra mitad y repite.

`busqueda_binaria([10, 20, 30, 40, 50], 40)` es `3`.

> ¿Atascado? Ejecuta `pista("E3")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_3, points: 0, grade: false, solution: true -->
```python
def busqueda_binaria(datos, buscado):
    ...
### INICIO SOLUCION
def busqueda_binaria(datos, buscado):
    izquierda = 0
    derecha = len(datos) - 1
    while izquierda <= derecha:
        medio = (izquierda + derecha) // 2
        if datos[medio] == buscado:
            return medio
        if datos[medio] < buscado:
            izquierda = medio + 1
        else:
            derecha = medio - 1
    return -1
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_3, points: 10, grade: true, solution: false -->
```python
assert callable(busqueda_binaria), "busqueda_binaria debe ser una funcion"
assert busqueda_binaria([10, 20, 30, 40, 50], 40) == 3
assert busqueda_binaria([10, 20, 30, 40, 50], 10) == 0, "El primero tambien"
assert busqueda_binaria([10, 20, 30, 40, 50], 99) == -1, "Si no esta, -1"
print("Las tres busquedas binarias dan bien.")
### INICIO PRUEBAS OCULTAS
assert busqueda_binaria([10, 20, 30, 40, 50], 50) == 4, "El ultimo tambien"
assert busqueda_binaria([], 1) == -1, "Lista vacia: no esta"
assert busqueda_binaria([7], 7) == 0, "Un solo elemento, y es el buscado"
assert busqueda_binaria([7], 3) == -1
_g = list(range(0, 2000, 2))
assert busqueda_binaria(_g, 1998) == 999, "Debe funcionar tambien en listas grandes"
assert busqueda_binaria(_g, 999) == -1, "999 es impar: no esta en la lista"
### FIN PRUEBAS OCULTAS

```

### Ejercicio 4 — La precondición

*Dificultad ★★☆☆ · 10 puntos*

La búsqueda binaria solo sirve si la lista está ordenada. Escribe la función
que lo comprueba.

`esta_ordenada(datos)` devuelve `True` si cada elemento es menor o igual que el
siguiente, y `False` si alguno rompe el orden.

`esta_ordenada([1, 2, 2, 5])` es `True`.
`esta_ordenada([1, 5, 2])` es `False`.

Una lista vacía o de un solo elemento está ordenada: no hay ninguna pareja que
pueda estar mal.

> ¿Atascado? Ejecuta `pista("E4")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_4, points: 0, grade: false, solution: true -->
```python
def esta_ordenada(datos):
    ...
### INICIO SOLUCION
def esta_ordenada(datos):
    for i in range(len(datos) - 1):
        if datos[i] > datos[i + 1]:
            return False
    return True
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_4, points: 10, grade: true, solution: false -->
```python
assert callable(esta_ordenada), "esta_ordenada debe ser una funcion"
assert esta_ordenada([1, 2, 2, 5]) is True
assert esta_ordenada([1, 5, 2]) is False
assert esta_ordenada([]) is True, "Una lista vacia esta ordenada"
assert esta_ordenada([7]) is True, "Un solo elemento esta ordenado"
print("Los cuatro casos dan bien.")
### INICIO PRUEBAS OCULTAS
assert esta_ordenada([5, 4, 3, 2, 1]) is False, "Al reves no esta ordenada"
assert esta_ordenada([1, 1, 1]) is True, "Elementos iguales NO rompen el orden"
assert esta_ordenada([1, 2, 3, 0]) is False, "El fallo esta al final"
assert isinstance(esta_ordenada([1, 2]), bool), "Debe devolver True o False"
### FIN PRUEBAS OCULTAS

```

### Ejercicio 5 — Burbuja, y cuánto cuesta

*Dificultad ★★★☆ · 15 puntos*

`burbuja(datos)` devuelve una **tupla** con dos cosas: la lista ordenada y el
número de **comparaciones** que hizo.

```python
burbuja([3, 1, 2])   ->   ([1, 2, 3], 3)
```

El algoritmo: recorre la lista comparando cada elemento con su vecino de la
derecha e intercambiándolos si están al revés. Repite tantas pasadas como
elementos haya.

Cuenta **una comparación cada vez que comparas dos vecinos**, la intercambies o
no. Con la versión clásica —dos ciclos, el de dentro llegando hasta
`n - i - 1`— una lista de 3 elementos da exactamente 3 comparaciones.

> ¿Atascado? Ejecuta `pista("E5")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_5, points: 0, grade: false, solution: true -->
```python
def burbuja(datos):
    ...
### INICIO SOLUCION
def burbuja(datos):
    datos = list(datos)
    comparaciones = 0
    n = len(datos)
    for i in range(n):
        for j in range(n - i - 1):
            comparaciones = comparaciones + 1
            if datos[j] > datos[j + 1]:
                datos[j], datos[j + 1] = datos[j + 1], datos[j]
    return (datos, comparaciones)
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_5, points: 15, grade: true, solution: false -->
```python
assert callable(burbuja), "burbuja debe ser una funcion"
_r = burbuja([3, 1, 2])
assert isinstance(_r, tuple) and len(_r) == 2, "Debe devolver una tupla (lista, comparaciones)"
assert _r[0] == [1, 2, 3], f"La lista ordenada debe ser [1, 2, 3] y dio {_r[0]}"
assert _r[1] == 3, f"Con 3 elementos son 3 comparaciones y contaste {_r[1]}"
print("burbuja([3, 1, 2]) =", _r)
### INICIO PRUEBAS OCULTAS
assert burbuja([5, 2, 9, 1])[0] == [1, 2, 5, 9]
assert burbuja([])[0] == [] and burbuja([])[1] == 0, "Lista vacia: cero comparaciones"
assert burbuja([1])[1] == 0, "Un solo elemento: nada que comparar"
assert burbuja([2, 1])[1] == 1, "Dos elementos: una comparacion"
assert burbuja([4, 3, 2, 1])[1] == 6, "Cuatro elementos: 3+2+1 = 6 comparaciones"
_o = [3, 1, 2]
burbuja(_o)
assert _o == [3, 1, 2], "No modifiques la lista original"
### FIN PRUEBAS OCULTAS

```

### Ejercicio 6 — ¿Cuál cuesta menos?

*Dificultad ★★★★ · 15 puntos*

El de cierre, y el que resume la semana. No se te pide escribir un algoritmo
nuevo: se te pide **medir** y sacar la conclusión.

`comparar(datos, buscado)` devuelve un diccionario con cuatro llaves:

| Llave | Qué guarda |
|---|---|
| `pasos_lineal` | cuántos elementos mira la búsqueda lineal hasta encontrarlo (o hasta el final) |
| `pasos_binaria` | cuántas veces parte la lista la búsqueda binaria |
| `se_puede_binaria` | `True` si `datos` está ordenada, `False` si no |
| `mejor` | `"binaria"` si se puede usar **y** hace menos pasos; `"lineal"` en cualquier otro caso |

Cuenta **un paso por cada elemento mirado** en la lineal, y **un paso por cada
vuelta del `while`** en la binaria.

Con `comparar(list(range(100)), 99)`: la lineal da 100 pasos, la binaria 7, está
ordenada, y la mejor es `"binaria"`.

> Si la lista **no** está ordenada, la binaria no se puede usar aunque diera
> menos pasos: daría una respuesta equivocada. En ese caso `mejor` es
> `"lineal"`.

> ¿Atascado? Ejecuta `pista("E6")` en una celda nueva. Hay 3, de la que hace pensar a la que casi resuelve. Pedirlas no resta puntos.


<!-- grade_id: ejercicio_6, points: 0, grade: false, solution: true -->
```python
def comparar(datos, buscado):
    ...
### INICIO SOLUCION
def comparar(datos, buscado):
    pasos_lineal = 0
    for i in range(len(datos)):
        pasos_lineal = pasos_lineal + 1
        if datos[i] == buscado:
            break

    pasos_binaria = 0
    izquierda = 0
    derecha = len(datos) - 1
    while izquierda <= derecha:
        pasos_binaria = pasos_binaria + 1
        medio = (izquierda + derecha) // 2
        if datos[medio] == buscado:
            break
        if datos[medio] < buscado:
            izquierda = medio + 1
        else:
            derecha = medio - 1

    ordenada = True
    for i in range(len(datos) - 1):
        if datos[i] > datos[i + 1]:
            ordenada = False

    if ordenada and pasos_binaria < pasos_lineal:
        mejor = "binaria"
    else:
        mejor = "lineal"

    return {
        "pasos_lineal": pasos_lineal,
        "pasos_binaria": pasos_binaria,
        "se_puede_binaria": ordenada,
        "mejor": mejor,
    }
### FIN SOLUCION

```

<!-- grade_id: test_ejercicio_6, points: 15, grade: true, solution: false -->
```python
assert callable(comparar), "comparar debe ser una funcion"
_r = comparar(list(range(100)), 99)
assert isinstance(_r, dict), "comparar debe devolver un diccionario"
assert set(_r) == {"pasos_lineal", "pasos_binaria", "se_puede_binaria", "mejor"}, \
    "Deben estar las cuatro llaves exactas"
assert _r["pasos_lineal"] == 100, f"La lineal mira los 100 y contaste {_r['pasos_lineal']}"
assert _r["se_puede_binaria"] is True, "range(100) esta ordenada"
assert _r["mejor"] == "binaria"
print("comparar(range(100), 99) ->", _r)
### INICIO PRUEBAS OCULTAS
_a = comparar([5, 1, 9, 3], 9)
assert _a["se_puede_binaria"] is False, "[5, 1, 9, 3] NO esta ordenada"
assert _a["mejor"] == "lineal", "Sin lista ordenada, la binaria no se puede usar"
assert _a["pasos_lineal"] == 3, "Encuentra el 9 en la tercera mirada"
_b = comparar(list(range(100)), 99)
assert _b["pasos_binaria"] == 7, f"La binaria da 7 vueltas y contaste {_b['pasos_binaria']}"
_c = comparar([1, 2, 3], 1)
assert _c["pasos_lineal"] == 1, "El primero se encuentra en un paso"
assert _c["mejor"] == "lineal", "Si la binaria no hace MENOS pasos, gana la lineal"
### FIN PRUEBAS OCULTAS

```

---

## 6. Habla con el asistente · 3 min

**Cinco preguntas** para todo el cuadernillo. Este es el más difícil de los
seis, así que administra bien.


### En qué gastarlas

En los índices y en la búsqueda lineal, ninguna: ejecuta `ver_lista(...)` y
mira dónde cae cada posición. Guárdalas para la binaria —el bucle que no
termina es el error clásico y cuesta verlo solo—, para la burbuja, y para el
último, que reutiliza los anteriores.

### Cómo se pregunta bien

Mal: «la binaria no me funciona».
Bien: «mi `busqueda_binaria([10,20,30], 30)` se queda colgada.
Cuando no encuentro el dato hago `derecha = medio`. ¿Por qué eso no termina?»

La segunda dice el caso exacto, qué escribiste y qué pasó. Con eso el tutor te
puede responder de verdad.


---

## 7. Cierre · 3 min

Cierras la primera mitad del curso. Tres preguntas.


- ¿Podrías explicarle a alguien **por qué** la búsqueda binaria necesita que la
  lista esté ordenada? Si la respuesta es «porque si no, no funciona», todavía
  no lo tienes: la pregunta es *por qué* no funciona.
- Entre selección y burbuja, ¿sabrías decir en qué se parecen? Los dos tienen
  dos ciclos anidados y los dos multiplican su trabajo por cuatro al doblar los
  datos. Esa familia se llama «cuadrática», y la vas a volver a ver.
- Mira el resultado de `comparar_ordenamientos(200)`. Si mañana te dieran un
  millón de registros, ¿escribirías tu burbuja o usarías `sorted()`? ¿Y sabrías
  justificar la decisión con números?

### Lo que llevas

Seis semanas. Variables y tipos, análisis de problemas, pseudocódigo y
diagramas, decisiones, repeticiones, colecciones, búsqueda y ordenamiento. Con
eso ya se puede escribir software de verdad — no de juguete.

Lo que viene en la segunda mitad se construye entero encima de esto.

### Glosario de esta semana

| Palabra | Qué significa |
|---|---|
| **Lista** | Una variable que guarda varios valores en orden |
| **Índice** | La posición de un elemento. **Empieza en cero** |
| **`IndexError`** | Pediste una posición que no existe |
| **Mutable** | Que se puede cambiar después de creado. Por eso se puede ordenar una lista en su sitio |
| **Búsqueda lineal** | Mirar uno por uno. Funciona siempre |
| **Búsqueda binaria** | Partir por la mitad. **Solo** sobre listas ordenadas |
| **Precondición** | Lo que tiene que cumplirse para que un algoritmo sea válido |
| **Complejidad** | Cómo crece el trabajo de un algoritmo cuando crecen los datos |
